unsat
(assume a0 (and (and (and (and (and (and (and (and (and (and (or (or (or (or (or (= (op e0 e0) e0) (= (op e0 e0) e1)) (= (op e0 e0) e2)) (= (op e0 e0) e3)) (= (op e0 e0) e4)) (= (op e0 e0) e5)) (or (or (or (or (or (= (op e0 e1) e0) (= (op e0 e1) e1)) (= (op e0 e1) e2)) (= (op e0 e1) e3)) (= (op e0 e1) e4)) (= (op e0 e1) e5))) (or (or (or (or (or (= (op e0 e2) e0) (= (op e0 e2) e1)) (= (op e0 e2) e2)) (= (op e0 e2) e3)) (= (op e0 e2) e4)) (= (op e0 e2) e5))) (or (or (or (or (or (= (op e0 e3) e0) (= (op e0 e3) e1)) (= (op e0 e3) e2)) (= (op e0 e3) e3)) (= (op e0 e3) e4)) (= (op e0 e3) e5))) (or (or (or (or (or (= (op e0 e4) e0) (= (op e0 e4) e1)) (= (op e0 e4) e2)) (= (op e0 e4) e3)) (= (op e0 e4) e4)) (= (op e0 e4) e5))) (or (or (or (or (or (= (op e0 e5) e0) (= (op e0 e5) e1)) (= (op e0 e5) e2)) (= (op e0 e5) e3)) (= (op e0 e5) e4)) (= (op e0 e5) e5))) (and (and (and (and (and (or (or (or (or (or (= (op e1 e0) e0) (= (op e1 e0) e1)) (= (op e1 e0) e2)) (= (op e1 e0) e3)) (= (op e1 e0) e4)) (= (op e1 e0) e5)) (or (or (or (or (or (= (op e1 e1) e0) (= (op e1 e1) e1)) (= (op e1 e1) e2)) (= (op e1 e1) e3)) (= (op e1 e1) e4)) (= (op e1 e1) e5))) (or (or (or (or (or (= (op e1 e2) e0) (= (op e1 e2) e1)) (= (op e1 e2) e2)) (= (op e1 e2) e3)) (= (op e1 e2) e4)) (= (op e1 e2) e5))) (or (or (or (or (or (= (op e1 e3) e0) (= (op e1 e3) e1)) (= (op e1 e3) e2)) (= (op e1 e3) e3)) (= (op e1 e3) e4)) (= (op e1 e3) e5))) (or (or (or (or (or (= (op e1 e4) e0) (= (op e1 e4) e1)) (= (op e1 e4) e2)) (= (op e1 e4) e3)) (= (op e1 e4) e4)) (= (op e1 e4) e5))) (or (or (or (or (or (= (op e1 e5) e0) (= (op e1 e5) e1)) (= (op e1 e5) e2)) (= (op e1 e5) e3)) (= (op e1 e5) e4)) (= (op e1 e5) e5)))) (and (and (and (and (and (or (or (or (or (or (= (op e2 e0) e0) (= (op e2 e0) e1)) (= (op e2 e0) e2)) (= (op e2 e0) e3)) (= (op e2 e0) e4)) (= (op e2 e0) e5)) (or (or (or (or (or (= (op e2 e1) e0) (= (op e2 e1) e1)) (= (op e2 e1) e2)) (= (op e2 e1) e3)) (= (op e2 e1) e4)) (= (op e2 e1) e5))) (or (or (or (or (or (= (op e2 e2) e0) (= (op e2 e2) e1)) (= (op e2 e2) e2)) (= (op e2 e2) e3)) (= (op e2 e2) e4)) (= (op e2 e2) e5))) (or (or (or (or (or (= (op e2 e3) e0) (= (op e2 e3) e1)) (= (op e2 e3) e2)) (= (op e2 e3) e3)) (= (op e2 e3) e4)) (= (op e2 e3) e5))) (or (or (or (or (or (= (op e2 e4) e0) (= (op e2 e4) e1)) (= (op e2 e4) e2)) (= (op e2 e4) e3)) (= (op e2 e4) e4)) (= (op e2 e4) e5))) (or (or (or (or (or (= (op e2 e5) e0) (= (op e2 e5) e1)) (= (op e2 e5) e2)) (= (op e2 e5) e3)) (= (op e2 e5) e4)) (= (op e2 e5) e5)))) (and (and (and (and (and (or (or (or (or (or (= (op e3 e0) e0) (= (op e3 e0) e1)) (= (op e3 e0) e2)) (= (op e3 e0) e3)) (= (op e3 e0) e4)) (= (op e3 e0) e5)) (or (or (or (or (or (= (op e3 e1) e0) (= (op e3 e1) e1)) (= (op e3 e1) e2)) (= (op e3 e1) e3)) (= (op e3 e1) e4)) (= (op e3 e1) e5))) (or (or (or (or (or (= (op e3 e2) e0) (= (op e3 e2) e1)) (= (op e3 e2) e2)) (= (op e3 e2) e3)) (= (op e3 e2) e4)) (= (op e3 e2) e5))) (or (or (or (or (or (= (op e3 e3) e0) (= (op e3 e3) e1)) (= (op e3 e3) e2)) (= (op e3 e3) e3)) (= (op e3 e3) e4)) (= (op e3 e3) e5))) (or (or (or (or (or (= (op e3 e4) e0) (= (op e3 e4) e1)) (= (op e3 e4) e2)) (= (op e3 e4) e3)) (= (op e3 e4) e4)) (= (op e3 e4) e5))) (or (or (or (or (or (= (op e3 e5) e0) (= (op e3 e5) e1)) (= (op e3 e5) e2)) (= (op e3 e5) e3)) (= (op e3 e5) e4)) (= (op e3 e5) e5)))) (and (and (and (and (and (or (or (or (or (or (= (op e4 e0) e0) (= (op e4 e0) e1)) (= (op e4 e0) e2)) (= (op e4 e0) e3)) (= (op e4 e0) e4)) (= (op e4 e0) e5)) (or (or (or (or (or (= (op e4 e1) e0) (= (op e4 e1) e1)) (= (op e4 e1) e2)) (= (op e4 e1) e3)) (= (op e4 e1) e4)) (= (op e4 e1) e5))) (or (or (or (or (or (= (op e4 e2) e0) (= (op e4 e2) e1)) (= (op e4 e2) e2)) (= (op e4 e2) e3)) (= (op e4 e2) e4)) (= (op e4 e2) e5))) (or (or (or (or (or (= (op e4 e3) e0) (= (op e4 e3) e1)) (= (op e4 e3) e2)) (= (op e4 e3) e3)) (= (op e4 e3) e4)) (= (op e4 e3) e5))) (or (or (or (or (or (= (op e4 e4) e0) (= (op e4 e4) e1)) (= (op e4 e4) e2)) (= (op e4 e4) e3)) (= (op e4 e4) e4)) (= (op e4 e4) e5))) (or (or (or (or (or (= (op e4 e5) e0) (= (op e4 e5) e1)) (= (op e4 e5) e2)) (= (op e4 e5) e3)) (= (op e4 e5) e4)) (= (op e4 e5) e5)))) (and (and (and (and (and (or (or (or (or (or (= (op e5 e0) e0) (= (op e5 e0) e1)) (= (op e5 e0) e2)) (= (op e5 e0) e3)) (= (op e5 e0) e4)) (= (op e5 e0) e5)) (or (or (or (or (or (= (op e5 e1) e0) (= (op e5 e1) e1)) (= (op e5 e1) e2)) (= (op e5 e1) e3)) (= (op e5 e1) e4)) (= (op e5 e1) e5))) (or (or (or (or (or (= (op e5 e2) e0) (= (op e5 e2) e1)) (= (op e5 e2) e2)) (= (op e5 e2) e3)) (= (op e5 e2) e4)) (= (op e5 e2) e5))) (or (or (or (or (or (= (op e5 e3) e0) (= (op e5 e3) e1)) (= (op e5 e3) e2)) (= (op e5 e3) e3)) (= (op e5 e3) e4)) (= (op e5 e3) e5))) (or (or (or (or (or (= (op e5 e4) e0) (= (op e5 e4) e1)) (= (op e5 e4) e2)) (= (op e5 e4) e3)) (= (op e5 e4) e4)) (= (op e5 e4) e5))) (or (or (or (or (or (= (op e5 e5) e0) (= (op e5 e5) e1)) (= (op e5 e5) e2)) (= (op e5 e5) e3)) (= (op e5 e5) e4)) (= (op e5 e5) e5)))))
(assume a1 (and (and (and (and (and (and (and (and (and (and (and (or (or (or (or (or (= (op e0 e0) e0) (= (op e0 e1) e0)) (= (op e0 e2) e0)) (= (op e0 e3) e0)) (= (op e0 e4) e0)) (= (op e0 e5) e0)) (or (or (or (or (or (= (op e0 e0) e0) (= (op e1 e0) e0)) (= (op e2 e0) e0)) (= (op e3 e0) e0)) (= (op e4 e0) e0)) (= (op e5 e0) e0))) (and (or (or (or (or (or (= (op e0 e0) e1) (= (op e0 e1) e1)) (= (op e0 e2) e1)) (= (op e0 e3) e1)) (= (op e0 e4) e1)) (= (op e0 e5) e1)) (or (or (or (or (or (= (op e0 e0) e1) (= (op e1 e0) e1)) (= (op e2 e0) e1)) (= (op e3 e0) e1)) (= (op e4 e0) e1)) (= (op e5 e0) e1)))) (and (or (or (or (or (or (= (op e0 e0) e2) (= (op e0 e1) e2)) (= (op e0 e2) e2)) (= (op e0 e3) e2)) (= (op e0 e4) e2)) (= (op e0 e5) e2)) (or (or (or (or (or (= (op e0 e0) e2) (= (op e1 e0) e2)) (= (op e2 e0) e2)) (= (op e3 e0) e2)) (= (op e4 e0) e2)) (= (op e5 e0) e2)))) (and (or (or (or (or (or (= (op e0 e0) e3) (= (op e0 e1) e3)) (= (op e0 e2) e3)) (= (op e0 e3) e3)) (= (op e0 e4) e3)) (= (op e0 e5) e3)) (or (or (or (or (or (= (op e0 e0) e3) (= (op e1 e0) e3)) (= (op e2 e0) e3)) (= (op e3 e0) e3)) (= (op e4 e0) e3)) (= (op e5 e0) e3)))) (and (or (or (or (or (or (= (op e0 e0) e4) (= (op e0 e1) e4)) (= (op e0 e2) e4)) (= (op e0 e3) e4)) (= (op e0 e4) e4)) (= (op e0 e5) e4)) (or (or (or (or (or (= (op e0 e0) e4) (= (op e1 e0) e4)) (= (op e2 e0) e4)) (= (op e3 e0) e4)) (= (op e4 e0) e4)) (= (op e5 e0) e4)))) (and (or (or (or (or (or (= (op e0 e0) e5) (= (op e0 e1) e5)) (= (op e0 e2) e5)) (= (op e0 e3) e5)) (= (op e0 e4) e5)) (= (op e0 e5) e5)) (or (or (or (or (or (= (op e0 e0) e5) (= (op e1 e0) e5)) (= (op e2 e0) e5)) (= (op e3 e0) e5)) (= (op e4 e0) e5)) (= (op e5 e0) e5)))) (and (and (and (and (and (and (or (or (or (or (or (= (op e1 e0) e0) (= (op e1 e1) e0)) (= (op e1 e2) e0)) (= (op e1 e3) e0)) (= (op e1 e4) e0)) (= (op e1 e5) e0)) (or (or (or (or (or (= (op e0 e1) e0) (= (op e1 e1) e0)) (= (op e2 e1) e0)) (= (op e3 e1) e0)) (= (op e4 e1) e0)) (= (op e5 e1) e0))) (and (or (or (or (or (or (= (op e1 e0) e1) (= (op e1 e1) e1)) (= (op e1 e2) e1)) (= (op e1 e3) e1)) (= (op e1 e4) e1)) (= (op e1 e5) e1)) (or (or (or (or (or (= (op e0 e1) e1) (= (op e1 e1) e1)) (= (op e2 e1) e1)) (= (op e3 e1) e1)) (= (op e4 e1) e1)) (= (op e5 e1) e1)))) (and (or (or (or (or (or (= (op e1 e0) e2) (= (op e1 e1) e2)) (= (op e1 e2) e2)) (= (op e1 e3) e2)) (= (op e1 e4) e2)) (= (op e1 e5) e2)) (or (or (or (or (or (= (op e0 e1) e2) (= (op e1 e1) e2)) (= (op e2 e1) e2)) (= (op e3 e1) e2)) (= (op e4 e1) e2)) (= (op e5 e1) e2)))) (and (or (or (or (or (or (= (op e1 e0) e3) (= (op e1 e1) e3)) (= (op e1 e2) e3)) (= (op e1 e3) e3)) (= (op e1 e4) e3)) (= (op e1 e5) e3)) (or (or (or (or (or (= (op e0 e1) e3) (= (op e1 e1) e3)) (= (op e2 e1) e3)) (= (op e3 e1) e3)) (= (op e4 e1) e3)) (= (op e5 e1) e3)))) (and (or (or (or (or (or (= (op e1 e0) e4) (= (op e1 e1) e4)) (= (op e1 e2) e4)) (= (op e1 e3) e4)) (= (op e1 e4) e4)) (= (op e1 e5) e4)) (or (or (or (or (or (= (op e0 e1) e4) (= (op e1 e1) e4)) (= (op e2 e1) e4)) (= (op e3 e1) e4)) (= (op e4 e1) e4)) (= (op e5 e1) e4)))) (and (or (or (or (or (or (= (op e1 e0) e5) (= (op e1 e1) e5)) (= (op e1 e2) e5)) (= (op e1 e3) e5)) (= (op e1 e4) e5)) (= (op e1 e5) e5)) (or (or (or (or (or (= (op e0 e1) e5) (= (op e1 e1) e5)) (= (op e2 e1) e5)) (= (op e3 e1) e5)) (= (op e4 e1) e5)) (= (op e5 e1) e5))))) (and (and (and (and (and (and (or (or (or (or (or (= (op e2 e0) e0) (= (op e2 e1) e0)) (= (op e2 e2) e0)) (= (op e2 e3) e0)) (= (op e2 e4) e0)) (= (op e2 e5) e0)) (or (or (or (or (or (= (op e0 e2) e0) (= (op e1 e2) e0)) (= (op e2 e2) e0)) (= (op e3 e2) e0)) (= (op e4 e2) e0)) (= (op e5 e2) e0))) (and (or (or (or (or (or (= (op e2 e0) e1) (= (op e2 e1) e1)) (= (op e2 e2) e1)) (= (op e2 e3) e1)) (= (op e2 e4) e1)) (= (op e2 e5) e1)) (or (or (or (or (or (= (op e0 e2) e1) (= (op e1 e2) e1)) (= (op e2 e2) e1)) (= (op e3 e2) e1)) (= (op e4 e2) e1)) (= (op e5 e2) e1)))) (and (or (or (or (or (or (= (op e2 e0) e2) (= (op e2 e1) e2)) (= (op e2 e2) e2)) (= (op e2 e3) e2)) (= (op e2 e4) e2)) (= (op e2 e5) e2)) (or (or (or (or (or (= (op e0 e2) e2) (= (op e1 e2) e2)) (= (op e2 e2) e2)) (= (op e3 e2) e2)) (= (op e4 e2) e2)) (= (op e5 e2) e2)))) (and (or (or (or (or (or (= (op e2 e0) e3) (= (op e2 e1) e3)) (= (op e2 e2) e3)) (= (op e2 e3) e3)) (= (op e2 e4) e3)) (= (op e2 e5) e3)) (or (or (or (or (or (= (op e0 e2) e3) (= (op e1 e2) e3)) (= (op e2 e2) e3)) (= (op e3 e2) e3)) (= (op e4 e2) e3)) (= (op e5 e2) e3)))) (and (or (or (or (or (or (= (op e2 e0) e4) (= (op e2 e1) e4)) (= (op e2 e2) e4)) (= (op e2 e3) e4)) (= (op e2 e4) e4)) (= (op e2 e5) e4)) (or (or (or (or (or (= (op e0 e2) e4) (= (op e1 e2) e4)) (= (op e2 e2) e4)) (= (op e3 e2) e4)) (= (op e4 e2) e4)) (= (op e5 e2) e4)))) (and (or (or (or (or (or (= (op e2 e0) e5) (= (op e2 e1) e5)) (= (op e2 e2) e5)) (= (op e2 e3) e5)) (= (op e2 e4) e5)) (= (op e2 e5) e5)) (or (or (or (or (or (= (op e0 e2) e5) (= (op e1 e2) e5)) (= (op e2 e2) e5)) (= (op e3 e2) e5)) (= (op e4 e2) e5)) (= (op e5 e2) e5))))) (and (and (and (and (and (and (or (or (or (or (or (= (op e3 e0) e0) (= (op e3 e1) e0)) (= (op e3 e2) e0)) (= (op e3 e3) e0)) (= (op e3 e4) e0)) (= (op e3 e5) e0)) (or (or (or (or (or (= (op e0 e3) e0) (= (op e1 e3) e0)) (= (op e2 e3) e0)) (= (op e3 e3) e0)) (= (op e4 e3) e0)) (= (op e5 e3) e0))) (and (or (or (or (or (or (= (op e3 e0) e1) (= (op e3 e1) e1)) (= (op e3 e2) e1)) (= (op e3 e3) e1)) (= (op e3 e4) e1)) (= (op e3 e5) e1)) (or (or (or (or (or (= (op e0 e3) e1) (= (op e1 e3) e1)) (= (op e2 e3) e1)) (= (op e3 e3) e1)) (= (op e4 e3) e1)) (= (op e5 e3) e1)))) (and (or (or (or (or (or (= (op e3 e0) e2) (= (op e3 e1) e2)) (= (op e3 e2) e2)) (= (op e3 e3) e2)) (= (op e3 e4) e2)) (= (op e3 e5) e2)) (or (or (or (or (or (= (op e0 e3) e2) (= (op e1 e3) e2)) (= (op e2 e3) e2)) (= (op e3 e3) e2)) (= (op e4 e3) e2)) (= (op e5 e3) e2)))) (and (or (or (or (or (or (= (op e3 e0) e3) (= (op e3 e1) e3)) (= (op e3 e2) e3)) (= (op e3 e3) e3)) (= (op e3 e4) e3)) (= (op e3 e5) e3)) (or (or (or (or (or (= (op e0 e3) e3) (= (op e1 e3) e3)) (= (op e2 e3) e3)) (= (op e3 e3) e3)) (= (op e4 e3) e3)) (= (op e5 e3) e3)))) (and (or (or (or (or (or (= (op e3 e0) e4) (= (op e3 e1) e4)) (= (op e3 e2) e4)) (= (op e3 e3) e4)) (= (op e3 e4) e4)) (= (op e3 e5) e4)) (or (or (or (or (or (= (op e0 e3) e4) (= (op e1 e3) e4)) (= (op e2 e3) e4)) (= (op e3 e3) e4)) (= (op e4 e3) e4)) (= (op e5 e3) e4)))) (and (or (or (or (or (or (= (op e3 e0) e5) (= (op e3 e1) e5)) (= (op e3 e2) e5)) (= (op e3 e3) e5)) (= (op e3 e4) e5)) (= (op e3 e5) e5)) (or (or (or (or (or (= (op e0 e3) e5) (= (op e1 e3) e5)) (= (op e2 e3) e5)) (= (op e3 e3) e5)) (= (op e4 e3) e5)) (= (op e5 e3) e5))))) (and (and (and (and (and (and (or (or (or (or (or (= (op e4 e0) e0) (= (op e4 e1) e0)) (= (op e4 e2) e0)) (= (op e4 e3) e0)) (= (op e4 e4) e0)) (= (op e4 e5) e0)) (or (or (or (or (or (= (op e0 e4) e0) (= (op e1 e4) e0)) (= (op e2 e4) e0)) (= (op e3 e4) e0)) (= (op e4 e4) e0)) (= (op e5 e4) e0))) (and (or (or (or (or (or (= (op e4 e0) e1) (= (op e4 e1) e1)) (= (op e4 e2) e1)) (= (op e4 e3) e1)) (= (op e4 e4) e1)) (= (op e4 e5) e1)) (or (or (or (or (or (= (op e0 e4) e1) (= (op e1 e4) e1)) (= (op e2 e4) e1)) (= (op e3 e4) e1)) (= (op e4 e4) e1)) (= (op e5 e4) e1)))) (and (or (or (or (or (or (= (op e4 e0) e2) (= (op e4 e1) e2)) (= (op e4 e2) e2)) (= (op e4 e3) e2)) (= (op e4 e4) e2)) (= (op e4 e5) e2)) (or (or (or (or (or (= (op e0 e4) e2) (= (op e1 e4) e2)) (= (op e2 e4) e2)) (= (op e3 e4) e2)) (= (op e4 e4) e2)) (= (op e5 e4) e2)))) (and (or (or (or (or (or (= (op e4 e0) e3) (= (op e4 e1) e3)) (= (op e4 e2) e3)) (= (op e4 e3) e3)) (= (op e4 e4) e3)) (= (op e4 e5) e3)) (or (or (or (or (or (= (op e0 e4) e3) (= (op e1 e4) e3)) (= (op e2 e4) e3)) (= (op e3 e4) e3)) (= (op e4 e4) e3)) (= (op e5 e4) e3)))) (and (or (or (or (or (or (= (op e4 e0) e4) (= (op e4 e1) e4)) (= (op e4 e2) e4)) (= (op e4 e3) e4)) (= (op e4 e4) e4)) (= (op e4 e5) e4)) (or (or (or (or (or (= (op e0 e4) e4) (= (op e1 e4) e4)) (= (op e2 e4) e4)) (= (op e3 e4) e4)) (= (op e4 e4) e4)) (= (op e5 e4) e4)))) (and (or (or (or (or (or (= (op e4 e0) e5) (= (op e4 e1) e5)) (= (op e4 e2) e5)) (= (op e4 e3) e5)) (= (op e4 e4) e5)) (= (op e4 e5) e5)) (or (or (or (or (or (= (op e0 e4) e5) (= (op e1 e4) e5)) (= (op e2 e4) e5)) (= (op e3 e4) e5)) (= (op e4 e4) e5)) (= (op e5 e4) e5))))) (and (and (and (and (and (and (or (or (or (or (or (= (op e5 e0) e0) (= (op e5 e1) e0)) (= (op e5 e2) e0)) (= (op e5 e3) e0)) (= (op e5 e4) e0)) (= (op e5 e5) e0)) (or (or (or (or (or (= (op e0 e5) e0) (= (op e1 e5) e0)) (= (op e2 e5) e0)) (= (op e3 e5) e0)) (= (op e4 e5) e0)) (= (op e5 e5) e0))) (and (or (or (or (or (or (= (op e5 e0) e1) (= (op e5 e1) e1)) (= (op e5 e2) e1)) (= (op e5 e3) e1)) (= (op e5 e4) e1)) (= (op e5 e5) e1)) (or (or (or (or (or (= (op e0 e5) e1) (= (op e1 e5) e1)) (= (op e2 e5) e1)) (= (op e3 e5) e1)) (= (op e4 e5) e1)) (= (op e5 e5) e1)))) (and (or (or (or (or (or (= (op e5 e0) e2) (= (op e5 e1) e2)) (= (op e5 e2) e2)) (= (op e5 e3) e2)) (= (op e5 e4) e2)) (= (op e5 e5) e2)) (or (or (or (or (or (= (op e0 e5) e2) (= (op e1 e5) e2)) (= (op e2 e5) e2)) (= (op e3 e5) e2)) (= (op e4 e5) e2)) (= (op e5 e5) e2)))) (and (or (or (or (or (or (= (op e5 e0) e3) (= (op e5 e1) e3)) (= (op e5 e2) e3)) (= (op e5 e3) e3)) (= (op e5 e4) e3)) (= (op e5 e5) e3)) (or (or (or (or (or (= (op e0 e5) e3) (= (op e1 e5) e3)) (= (op e2 e5) e3)) (= (op e3 e5) e3)) (= (op e4 e5) e3)) (= (op e5 e5) e3)))) (and (or (or (or (or (or (= (op e5 e0) e4) (= (op e5 e1) e4)) (= (op e5 e2) e4)) (= (op e5 e3) e4)) (= (op e5 e4) e4)) (= (op e5 e5) e4)) (or (or (or (or (or (= (op e0 e5) e4) (= (op e1 e5) e4)) (= (op e2 e5) e4)) (= (op e3 e5) e4)) (= (op e4 e5) e4)) (= (op e5 e5) e4)))) (and (or (or (or (or (or (= (op e5 e0) e5) (= (op e5 e1) e5)) (= (op e5 e2) e5)) (= (op e5 e3) e5)) (= (op e5 e4) e5)) (= (op e5 e5) e5)) (or (or (or (or (or (= (op e0 e5) e5) (= (op e1 e5) e5)) (= (op e2 e5) e5)) (= (op e3 e5) e5)) (= (op e4 e5) e5)) (= (op e5 e5) e5))))))
(assume a2 (and (and (and (and (and (and (and (= (op unit e0) e0) (= (op e0 unit) e0)) (and (= (op unit e1) e1) (= (op e1 unit) e1))) (and (= (op unit e2) e2) (= (op e2 unit) e2))) (and (= (op unit e3) e3) (= (op e3 unit) e3))) (and (= (op unit e4) e4) (= (op e4 unit) e4))) (and (= (op unit e5) e5) (= (op e5 unit) e5))) (or (or (or (or (or (= unit e0) (= unit e1)) (= unit e2)) (= unit e3)) (= unit e4)) (= unit e5))))
(assume a3 (or (or (or (or (or (or (or (or (or (or (not (= (op e0 e0) (op e0 e0))) (not (= (op e1 e0) (op e0 e1)))) (not (= (op e2 e0) (op e0 e2)))) (not (= (op e3 e0) (op e0 e3)))) (not (= (op e4 e0) (op e0 e4)))) (not (= (op e5 e0) (op e0 e5)))) (or (or (or (or (or (not (= (op e0 e1) (op e1 e0))) (not (= (op e1 e1) (op e1 e1)))) (not (= (op e2 e1) (op e1 e2)))) (not (= (op e3 e1) (op e1 e3)))) (not (= (op e4 e1) (op e1 e4)))) (not (= (op e5 e1) (op e1 e5))))) (or (or (or (or (or (not (= (op e0 e2) (op e2 e0))) (not (= (op e1 e2) (op e2 e1)))) (not (= (op e2 e2) (op e2 e2)))) (not (= (op e3 e2) (op e2 e3)))) (not (= (op e4 e2) (op e2 e4)))) (not (= (op e5 e2) (op e2 e5))))) (or (or (or (or (or (not (= (op e0 e3) (op e3 e0))) (not (= (op e1 e3) (op e3 e1)))) (not (= (op e2 e3) (op e3 e2)))) (not (= (op e3 e3) (op e3 e3)))) (not (= (op e4 e3) (op e3 e4)))) (not (= (op e5 e3) (op e3 e5))))) (or (or (or (or (or (not (= (op e0 e4) (op e4 e0))) (not (= (op e1 e4) (op e4 e1)))) (not (= (op e2 e4) (op e4 e2)))) (not (= (op e3 e4) (op e4 e3)))) (not (= (op e4 e4) (op e4 e4)))) (not (= (op e5 e4) (op e4 e5))))) (or (or (or (or (or (not (= (op e0 e5) (op e5 e0))) (not (= (op e1 e5) (op e5 e1)))) (not (= (op e2 e5) (op e5 e2)))) (not (= (op e3 e5) (op e5 e3)))) (not (= (op e4 e5) (op e5 e4)))) (not (= (op e5 e5) (op e5 e5))))))
(assume a4 (or (or (or (or (or (and (and (and (and (and (not (= (op (op e0 e0) e0) e0)) (not (= (op (op e0 e1) e0) e1))) (not (= (op (op e0 e2) e0) e2))) (not (= (op (op e0 e3) e0) e3))) (not (= (op (op e0 e4) e0) e4))) (not (= (op (op e0 e5) e0) e5))) (and (and (and (and (and (not (= (op (op e1 e0) e1) e0)) (not (= (op (op e1 e1) e1) e1))) (not (= (op (op e1 e2) e1) e2))) (not (= (op (op e1 e3) e1) e3))) (not (= (op (op e1 e4) e1) e4))) (not (= (op (op e1 e5) e1) e5)))) (and (and (and (and (and (not (= (op (op e2 e0) e2) e0)) (not (= (op (op e2 e1) e2) e1))) (not (= (op (op e2 e2) e2) e2))) (not (= (op (op e2 e3) e2) e3))) (not (= (op (op e2 e4) e2) e4))) (not (= (op (op e2 e5) e2) e5)))) (and (and (and (and (and (not (= (op (op e3 e0) e3) e0)) (not (= (op (op e3 e1) e3) e1))) (not (= (op (op e3 e2) e3) e2))) (not (= (op (op e3 e3) e3) e3))) (not (= (op (op e3 e4) e3) e4))) (not (= (op (op e3 e5) e3) e5)))) (and (and (and (and (and (not (= (op (op e4 e0) e4) e0)) (not (= (op (op e4 e1) e4) e1))) (not (= (op (op e4 e2) e4) e2))) (not (= (op (op e4 e3) e4) e3))) (not (= (op (op e4 e4) e4) e4))) (not (= (op (op e4 e5) e4) e5)))) (and (and (and (and (and (not (= (op (op e5 e0) e5) e0)) (not (= (op (op e5 e1) e5) e1))) (not (= (op (op e5 e2) e5) e2))) (not (= (op (op e5 e3) e5) e3))) (not (= (op (op e5 e4) e5) e4))) (not (= (op (op e5 e5) e5) e5)))))
(assume a5 (and (and (and (and (and (or (or (or (or (or (= (op (op e0 e0) (op e0 e0)) e0) (= (op (op e0 e1) (op e0 e1)) e1)) (= (op (op e0 e2) (op e0 e2)) e2)) (= (op (op e0 e3) (op e0 e3)) e3)) (= (op (op e0 e4) (op e0 e4)) e4)) (= (op (op e0 e5) (op e0 e5)) e5)) (or (or (or (or (or (= (op (op e1 e0) (op e1 e0)) e0) (= (op (op e1 e1) (op e1 e1)) e1)) (= (op (op e1 e2) (op e1 e2)) e2)) (= (op (op e1 e3) (op e1 e3)) e3)) (= (op (op e1 e4) (op e1 e4)) e4)) (= (op (op e1 e5) (op e1 e5)) e5))) (or (or (or (or (or (= (op (op e2 e0) (op e2 e0)) e0) (= (op (op e2 e1) (op e2 e1)) e1)) (= (op (op e2 e2) (op e2 e2)) e2)) (= (op (op e2 e3) (op e2 e3)) e3)) (= (op (op e2 e4) (op e2 e4)) e4)) (= (op (op e2 e5) (op e2 e5)) e5))) (or (or (or (or (or (= (op (op e3 e0) (op e3 e0)) e0) (= (op (op e3 e1) (op e3 e1)) e1)) (= (op (op e3 e2) (op e3 e2)) e2)) (= (op (op e3 e3) (op e3 e3)) e3)) (= (op (op e3 e4) (op e3 e4)) e4)) (= (op (op e3 e5) (op e3 e5)) e5))) (or (or (or (or (or (= (op (op e4 e0) (op e4 e0)) e0) (= (op (op e4 e1) (op e4 e1)) e1)) (= (op (op e4 e2) (op e4 e2)) e2)) (= (op (op e4 e3) (op e4 e3)) e3)) (= (op (op e4 e4) (op e4 e4)) e4)) (= (op (op e4 e5) (op e4 e5)) e5))) (or (or (or (or (or (= (op (op e5 e0) (op e5 e0)) e0) (= (op (op e5 e1) (op e5 e1)) e1)) (= (op (op e5 e2) (op e5 e2)) e2)) (= (op (op e5 e3) (op e5 e3)) e3)) (= (op (op e5 e4) (op e5 e4)) e4)) (= (op (op e5 e5) (op e5 e5)) e5))))
(assume a6 (or (or (or (or (or (and (and (and (and (and (not (= (op (op e0 e0) (op e0 e0)) e0)) (not (= (op (op e1 e0) (op e1 e0)) e1))) (not (= (op (op e2 e0) (op e2 e0)) e2))) (not (= (op (op e3 e0) (op e3 e0)) e3))) (not (= (op (op e4 e0) (op e4 e0)) e4))) (not (= (op (op e5 e0) (op e5 e0)) e5))) (and (and (and (and (and (not (= (op (op e0 e1) (op e0 e1)) e0)) (not (= (op (op e1 e1) (op e1 e1)) e1))) (not (= (op (op e2 e1) (op e2 e1)) e2))) (not (= (op (op e3 e1) (op e3 e1)) e3))) (not (= (op (op e4 e1) (op e4 e1)) e4))) (not (= (op (op e5 e1) (op e5 e1)) e5)))) (and (and (and (and (and (not (= (op (op e0 e2) (op e0 e2)) e0)) (not (= (op (op e1 e2) (op e1 e2)) e1))) (not (= (op (op e2 e2) (op e2 e2)) e2))) (not (= (op (op e3 e2) (op e3 e2)) e3))) (not (= (op (op e4 e2) (op e4 e2)) e4))) (not (= (op (op e5 e2) (op e5 e2)) e5)))) (and (and (and (and (and (not (= (op (op e0 e3) (op e0 e3)) e0)) (not (= (op (op e1 e3) (op e1 e3)) e1))) (not (= (op (op e2 e3) (op e2 e3)) e2))) (not (= (op (op e3 e3) (op e3 e3)) e3))) (not (= (op (op e4 e3) (op e4 e3)) e4))) (not (= (op (op e5 e3) (op e5 e3)) e5)))) (and (and (and (and (and (not (= (op (op e0 e4) (op e0 e4)) e0)) (not (= (op (op e1 e4) (op e1 e4)) e1))) (not (= (op (op e2 e4) (op e2 e4)) e2))) (not (= (op (op e3 e4) (op e3 e4)) e3))) (not (= (op (op e4 e4) (op e4 e4)) e4))) (not (= (op (op e5 e4) (op e5 e4)) e5)))) (and (and (and (and (and (not (= (op (op e0 e5) (op e0 e5)) e0)) (not (= (op (op e1 e5) (op e1 e5)) e1))) (not (= (op (op e2 e5) (op e2 e5)) e2))) (not (= (op (op e3 e5) (op e3 e5)) e3))) (not (= (op (op e4 e5) (op e4 e5)) e4))) (not (= (op (op e5 e5) (op e5 e5)) e5)))))
(assume a7 (and (and (and (and (and (and (or (not (= (op e0 (op e0 e0)) e0)) (= e0 unit)) (or (not (= (op e1 (op e1 e1)) e1)) (= e1 unit))) (or (not (= (op e2 (op e2 e2)) e2)) (= e2 unit))) (or (not (= (op e3 (op e3 e3)) e3)) (= e3 unit))) (or (not (= (op e4 (op e4 e4)) e4)) (= e4 unit))) (or (not (= (op e5 (op e5 e5)) e5)) (= e5 unit))) (and (and (and (and (and (or (or (or (or (or (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit)))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit)))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit)))) (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit)))) (or (or (or (or (or (and (= (op e0 (op e0 e1)) e1) (not (= (op e0 e1) unit))) (and (= (op e1 (op e1 e1)) e1) (not (= (op e1 e1) unit)))) (and (= (op e2 (op e2 e1)) e1) (not (= (op e2 e1) unit)))) (and (= (op e3 (op e3 e1)) e1) (not (= (op e3 e1) unit)))) (and (= (op e4 (op e4 e1)) e1) (not (= (op e4 e1) unit)))) (and (= (op e5 (op e5 e1)) e1) (not (= (op e5 e1) unit))))) (or (or (or (or (or (and (= (op e0 (op e0 e2)) e2) (not (= (op e0 e2) unit))) (and (= (op e1 (op e1 e2)) e2) (not (= (op e1 e2) unit)))) (and (= (op e2 (op e2 e2)) e2) (not (= (op e2 e2) unit)))) (and (= (op e3 (op e3 e2)) e2) (not (= (op e3 e2) unit)))) (and (= (op e4 (op e4 e2)) e2) (not (= (op e4 e2) unit)))) (and (= (op e5 (op e5 e2)) e2) (not (= (op e5 e2) unit))))) (or (or (or (or (or (and (= (op e0 (op e0 e3)) e3) (not (= (op e0 e3) unit))) (and (= (op e1 (op e1 e3)) e3) (not (= (op e1 e3) unit)))) (and (= (op e2 (op e2 e3)) e3) (not (= (op e2 e3) unit)))) (and (= (op e3 (op e3 e3)) e3) (not (= (op e3 e3) unit)))) (and (= (op e4 (op e4 e3)) e3) (not (= (op e4 e3) unit)))) (and (= (op e5 (op e5 e3)) e3) (not (= (op e5 e3) unit))))) (or (or (or (or (or (and (= (op e0 (op e0 e4)) e4) (not (= (op e0 e4) unit))) (and (= (op e1 (op e1 e4)) e4) (not (= (op e1 e4) unit)))) (and (= (op e2 (op e2 e4)) e4) (not (= (op e2 e4) unit)))) (and (= (op e3 (op e3 e4)) e4) (not (= (op e3 e4) unit)))) (and (= (op e4 (op e4 e4)) e4) (not (= (op e4 e4) unit)))) (and (= (op e5 (op e5 e4)) e4) (not (= (op e5 e4) unit))))) (or (or (or (or (or (and (= (op e0 (op e0 e5)) e5) (not (= (op e0 e5) unit))) (and (= (op e1 (op e1 e5)) e5) (not (= (op e1 e5) unit)))) (and (= (op e2 (op e2 e5)) e5) (not (= (op e2 e5) unit)))) (and (= (op e3 (op e3 e5)) e5) (not (= (op e3 e5) unit)))) (and (= (op e4 (op e4 e5)) e5) (not (= (op e4 e5) unit)))) (and (= (op e5 (op e5 e5)) e5) (not (= (op e5 e5) unit)))))))
(assume a8 (= unit e0))
(assume a9 (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= (op e0 e0) (op e1 e0))) (not (= (op e0 e0) (op e2 e0)))) (not (= (op e0 e0) (op e3 e0)))) (not (= (op e0 e0) (op e4 e0)))) (not (= (op e0 e0) (op e5 e0)))) (not (= (op e1 e0) (op e2 e0)))) (not (= (op e1 e0) (op e3 e0)))) (not (= (op e1 e0) (op e4 e0)))) (not (= (op e1 e0) (op e5 e0)))) (not (= (op e2 e0) (op e3 e0)))) (not (= (op e2 e0) (op e4 e0)))) (not (= (op e2 e0) (op e5 e0)))) (not (= (op e3 e0) (op e4 e0)))) (not (= (op e3 e0) (op e5 e0)))) (not (= (op e4 e0) (op e5 e0)))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= (op e0 e1) (op e1 e1))) (not (= (op e0 e1) (op e2 e1)))) (not (= (op e0 e1) (op e3 e1)))) (not (= (op e0 e1) (op e4 e1)))) (not (= (op e0 e1) (op e5 e1)))) (not (= (op e1 e1) (op e2 e1)))) (not (= (op e1 e1) (op e3 e1)))) (not (= (op e1 e1) (op e4 e1)))) (not (= (op e1 e1) (op e5 e1)))) (not (= (op e2 e1) (op e3 e1)))) (not (= (op e2 e1) (op e4 e1)))) (not (= (op e2 e1) (op e5 e1)))) (not (= (op e3 e1) (op e4 e1)))) (not (= (op e3 e1) (op e5 e1)))) (not (= (op e4 e1) (op e5 e1))))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= (op e0 e2) (op e1 e2))) (not (= (op e0 e2) (op e2 e2)))) (not (= (op e0 e2) (op e3 e2)))) (not (= (op e0 e2) (op e4 e2)))) (not (= (op e0 e2) (op e5 e2)))) (not (= (op e1 e2) (op e2 e2)))) (not (= (op e1 e2) (op e3 e2)))) (not (= (op e1 e2) (op e4 e2)))) (not (= (op e1 e2) (op e5 e2)))) (not (= (op e2 e2) (op e3 e2)))) (not (= (op e2 e2) (op e4 e2)))) (not (= (op e2 e2) (op e5 e2)))) (not (= (op e3 e2) (op e4 e2)))) (not (= (op e3 e2) (op e5 e2)))) (not (= (op e4 e2) (op e5 e2))))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= (op e0 e3) (op e1 e3))) (not (= (op e0 e3) (op e2 e3)))) (not (= (op e0 e3) (op e3 e3)))) (not (= (op e0 e3) (op e4 e3)))) (not (= (op e0 e3) (op e5 e3)))) (not (= (op e1 e3) (op e2 e3)))) (not (= (op e1 e3) (op e3 e3)))) (not (= (op e1 e3) (op e4 e3)))) (not (= (op e1 e3) (op e5 e3)))) (not (= (op e2 e3) (op e3 e3)))) (not (= (op e2 e3) (op e4 e3)))) (not (= (op e2 e3) (op e5 e3)))) (not (= (op e3 e3) (op e4 e3)))) (not (= (op e3 e3) (op e5 e3)))) (not (= (op e4 e3) (op e5 e3))))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= (op e0 e4) (op e1 e4))) (not (= (op e0 e4) (op e2 e4)))) (not (= (op e0 e4) (op e3 e4)))) (not (= (op e0 e4) (op e4 e4)))) (not (= (op e0 e4) (op e5 e4)))) (not (= (op e1 e4) (op e2 e4)))) (not (= (op e1 e4) (op e3 e4)))) (not (= (op e1 e4) (op e4 e4)))) (not (= (op e1 e4) (op e5 e4)))) (not (= (op e2 e4) (op e3 e4)))) (not (= (op e2 e4) (op e4 e4)))) (not (= (op e2 e4) (op e5 e4)))) (not (= (op e3 e4) (op e4 e4)))) (not (= (op e3 e4) (op e5 e4)))) (not (= (op e4 e4) (op e5 e4))))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= (op e0 e5) (op e1 e5))) (not (= (op e0 e5) (op e2 e5)))) (not (= (op e0 e5) (op e3 e5)))) (not (= (op e0 e5) (op e4 e5)))) (not (= (op e0 e5) (op e5 e5)))) (not (= (op e1 e5) (op e2 e5)))) (not (= (op e1 e5) (op e3 e5)))) (not (= (op e1 e5) (op e4 e5)))) (not (= (op e1 e5) (op e5 e5)))) (not (= (op e2 e5) (op e3 e5)))) (not (= (op e2 e5) (op e4 e5)))) (not (= (op e2 e5) (op e5 e5)))) (not (= (op e3 e5) (op e4 e5)))) (not (= (op e3 e5) (op e5 e5)))) (not (= (op e4 e5) (op e5 e5))))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= (op e0 e0) (op e0 e1))) (not (= (op e0 e0) (op e0 e2)))) (not (= (op e0 e0) (op e0 e3)))) (not (= (op e0 e0) (op e0 e4)))) (not (= (op e0 e0) (op e0 e5)))) (not (= (op e0 e1) (op e0 e2)))) (not (= (op e0 e1) (op e0 e3)))) (not (= (op e0 e1) (op e0 e4)))) (not (= (op e0 e1) (op e0 e5)))) (not (= (op e0 e2) (op e0 e3)))) (not (= (op e0 e2) (op e0 e4)))) (not (= (op e0 e2) (op e0 e5)))) (not (= (op e0 e3) (op e0 e4)))) (not (= (op e0 e3) (op e0 e5)))) (not (= (op e0 e4) (op e0 e5)))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= (op e1 e0) (op e1 e1))) (not (= (op e1 e0) (op e1 e2)))) (not (= (op e1 e0) (op e1 e3)))) (not (= (op e1 e0) (op e1 e4)))) (not (= (op e1 e0) (op e1 e5)))) (not (= (op e1 e1) (op e1 e2)))) (not (= (op e1 e1) (op e1 e3)))) (not (= (op e1 e1) (op e1 e4)))) (not (= (op e1 e1) (op e1 e5)))) (not (= (op e1 e2) (op e1 e3)))) (not (= (op e1 e2) (op e1 e4)))) (not (= (op e1 e2) (op e1 e5)))) (not (= (op e1 e3) (op e1 e4)))) (not (= (op e1 e3) (op e1 e5)))) (not (= (op e1 e4) (op e1 e5))))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= (op e2 e0) (op e2 e1))) (not (= (op e2 e0) (op e2 e2)))) (not (= (op e2 e0) (op e2 e3)))) (not (= (op e2 e0) (op e2 e4)))) (not (= (op e2 e0) (op e2 e5)))) (not (= (op e2 e1) (op e2 e2)))) (not (= (op e2 e1) (op e2 e3)))) (not (= (op e2 e1) (op e2 e4)))) (not (= (op e2 e1) (op e2 e5)))) (not (= (op e2 e2) (op e2 e3)))) (not (= (op e2 e2) (op e2 e4)))) (not (= (op e2 e2) (op e2 e5)))) (not (= (op e2 e3) (op e2 e4)))) (not (= (op e2 e3) (op e2 e5)))) (not (= (op e2 e4) (op e2 e5))))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= (op e3 e0) (op e3 e1))) (not (= (op e3 e0) (op e3 e2)))) (not (= (op e3 e0) (op e3 e3)))) (not (= (op e3 e0) (op e3 e4)))) (not (= (op e3 e0) (op e3 e5)))) (not (= (op e3 e1) (op e3 e2)))) (not (= (op e3 e1) (op e3 e3)))) (not (= (op e3 e1) (op e3 e4)))) (not (= (op e3 e1) (op e3 e5)))) (not (= (op e3 e2) (op e3 e3)))) (not (= (op e3 e2) (op e3 e4)))) (not (= (op e3 e2) (op e3 e5)))) (not (= (op e3 e3) (op e3 e4)))) (not (= (op e3 e3) (op e3 e5)))) (not (= (op e3 e4) (op e3 e5))))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= (op e4 e0) (op e4 e1))) (not (= (op e4 e0) (op e4 e2)))) (not (= (op e4 e0) (op e4 e3)))) (not (= (op e4 e0) (op e4 e4)))) (not (= (op e4 e0) (op e4 e5)))) (not (= (op e4 e1) (op e4 e2)))) (not (= (op e4 e1) (op e4 e3)))) (not (= (op e4 e1) (op e4 e4)))) (not (= (op e4 e1) (op e4 e5)))) (not (= (op e4 e2) (op e4 e3)))) (not (= (op e4 e2) (op e4 e4)))) (not (= (op e4 e2) (op e4 e5)))) (not (= (op e4 e3) (op e4 e4)))) (not (= (op e4 e3) (op e4 e5)))) (not (= (op e4 e4) (op e4 e5))))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= (op e5 e0) (op e5 e1))) (not (= (op e5 e0) (op e5 e2)))) (not (= (op e5 e0) (op e5 e3)))) (not (= (op e5 e0) (op e5 e4)))) (not (= (op e5 e0) (op e5 e5)))) (not (= (op e5 e1) (op e5 e2)))) (not (= (op e5 e1) (op e5 e3)))) (not (= (op e5 e1) (op e5 e4)))) (not (= (op e5 e1) (op e5 e5)))) (not (= (op e5 e2) (op e5 e3)))) (not (= (op e5 e2) (op e5 e4)))) (not (= (op e5 e2) (op e5 e5)))) (not (= (op e5 e3) (op e5 e4)))) (not (= (op e5 e3) (op e5 e5)))) (not (= (op e5 e4) (op e5 e5)))))))
(assume a10 (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= e0 e1)) (not (= e0 e2))) (not (= e0 e3))) (not (= e0 e4))) (not (= e0 e5))) (not (= e1 e2))) (not (= e1 e3))) (not (= e1 e4))) (not (= e1 e5))) (not (= e2 e3))) (not (= e2 e4))) (not (= e2 e5))) (not (= e3 e4))) (not (= e3 e5))) (not (= e4 e5))))
(step t0 (cl (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0)))) (not (= e2 (op e2 e0))) (not (= e0 (op e2 (op e2 e0))))) :rule and_neg)
(step t1 (cl (=> (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0)))) (= e2 (op e2 (op e2 e2)))) (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0))))) :rule implies_neg1)
(anchor :step t2)
(assume t2.a0 (= e2 (op e2 e0)))
(assume t2.a1 (= e0 (op e2 (op e2 e0))))
(step t2.t0 (cl (not (= (= (op e2 unit) e2) (= e2 (op e2 e0)))) (not (= (op e2 unit) e2)) (= e2 (op e2 e0))) :rule equiv_pos2)
(step t2.t1 (cl (= (= (op e2 unit) e2) (= e2 (op e2 unit)))) :rule rare_rewrite :args ("eq-symm" (op e2 unit) e2))
(step t2.t2 (cl (= e2 e2)) :rule refl)
(step t2.t3 (cl (= (op e2 unit) (op e2 e0))) :rule cong :premises (t2.t2 a8))
(step t2.t4 (cl (= (= e2 (op e2 unit)) (= e2 (op e2 e0)))) :rule cong :premises (t2.t2 t2.t3))
(step t2.t5 (cl (= (= (op e2 unit) e2) (= e2 (op e2 e0)))) :rule trans :premises (t2.t1 t2.t4))
(step t2.t6 (cl (and (and (and (and (and (and (= (op unit e0) e0) (= (op e0 unit) e0)) (and (= (op unit e1) e1) (= (op e1 unit) e1))) (and (= (op unit e2) e2) (= (op e2 unit) e2))) (and (= (op unit e3) e3) (= (op e3 unit) e3))) (and (= (op unit e4) e4) (= (op e4 unit) e4))) (and (= (op unit e5) e5) (= (op e5 unit) e5)))) :rule and :premises (a2) :args (0))
(step t2.t7 (cl (and (and (and (and (and (= (op unit e0) e0) (= (op e0 unit) e0)) (and (= (op unit e1) e1) (= (op e1 unit) e1))) (and (= (op unit e2) e2) (= (op e2 unit) e2))) (and (= (op unit e3) e3) (= (op e3 unit) e3))) (and (= (op unit e4) e4) (= (op e4 unit) e4)))) :rule and :premises (t2.t6) :args (0))
(step t2.t8 (cl (and (and (and (and (= (op unit e0) e0) (= (op e0 unit) e0)) (and (= (op unit e1) e1) (= (op e1 unit) e1))) (and (= (op unit e2) e2) (= (op e2 unit) e2))) (and (= (op unit e3) e3) (= (op e3 unit) e3)))) :rule and :premises (t2.t7) :args (0))
(step t2.t9 (cl (and (and (and (= (op unit e0) e0) (= (op e0 unit) e0)) (and (= (op unit e1) e1) (= (op e1 unit) e1))) (and (= (op unit e2) e2) (= (op e2 unit) e2)))) :rule and :premises (t2.t8) :args (0))
(step t2.t10 (cl (and (= (op unit e2) e2) (= (op e2 unit) e2))) :rule and :premises (t2.t9) :args (1))
(step t2.t11 (cl (= (op e2 unit) e2)) :rule and :premises (t2.t10) :args (1))
(step t2.t12 (cl (= e2 (op e2 e0))) :rule resolution :premises (t2.t0 t2.t5 t2.t11))
(step t2.t13 (cl (= (op e2 (op e2 e0)) e0)) :rule symm :premises (t2.a1))
(step t2.t14 (cl (= e0 (op e2 (op e2 e0)))) :rule symm :premises (t2.t13))
(step t2.t15 (cl (= (op e2 e0) e2)) :rule symm :premises (t2.t12))
(step t2.t16 (cl (= (op e2 (op e2 e0)) (op e2 e2))) :rule cong :premises (t2.t2 t2.t15))
(step t2.t17 (cl (= e0 (op e2 e2))) :rule trans :premises (t2.t14 t2.t16))
(step t2.t18 (cl (= (op e2 e0) (op e2 (op e2 e2)))) :rule cong :premises (t2.t2 t2.t17))
(step t2.t19 (cl (= e2 (op e2 (op e2 e2)))) :rule trans :premises (t2.t12 t2.t18))
(step t2 (cl (not (= e2 (op e2 e0))) (not (= e0 (op e2 (op e2 e0)))) (= e2 (op e2 (op e2 e2)))) :rule subproof :discharge (t2.a0 t2.a1))
(step t3 (cl (not (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0))))) (= e2 (op e2 e0))) :rule and_pos :args (0))
(step t4 (cl (not (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0))))) (= e0 (op e2 (op e2 e0)))) :rule and_pos :args (1))
(step t5 (cl (= e2 (op e2 (op e2 e2))) (not (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0))))) (not (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0)))))) :rule resolution :premises (t2 t3 t4))
(step t6 (cl (not (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0))))) (not (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0))))) (= e2 (op e2 (op e2 e2)))) :rule reordering :premises (t5))
(step t7 (cl (not (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0))))) (= e2 (op e2 (op e2 e2)))) :rule contraction :premises (t6))
(step t8 (cl (=> (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0)))) (= e2 (op e2 (op e2 e2)))) (= e2 (op e2 (op e2 e2)))) :rule resolution :premises (t1 t7))
(step t9 (cl (=> (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0)))) (= e2 (op e2 (op e2 e2)))) (not (= e2 (op e2 (op e2 e2))))) :rule implies_neg2)
(step t10 (cl (=> (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0)))) (= e2 (op e2 (op e2 e2)))) (=> (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0)))) (= e2 (op e2 (op e2 e2))))) :rule resolution :premises (t8 t9))
(step t11 (cl (=> (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0)))) (= e2 (op e2 (op e2 e2))))) :rule contraction :premises (t10))
(step t12 (cl (not (and (= e2 (op e2 e0)) (= e0 (op e2 (op e2 e0))))) (= e2 (op e2 (op e2 e2)))) :rule implies :premises (t11))
(step t13 (cl (not (= e2 (op e2 e0))) (not (= e0 (op e2 (op e2 e0)))) (= e2 (op e2 (op e2 e2)))) :rule resolution :premises (t0 t12))
(step t14 (cl (= e2 (op e2 (op e2 e2))) (not (= e2 (op e2 e0))) (not (= e0 (op e2 (op e2 e0))))) :rule reordering :premises (t13))
(step t15 (cl (not (and (= e0 (op e2 (op e2 e0))) (not (= e0 (op e2 e0))))) (= e0 (op e2 (op e2 e0)))) :rule and_pos :args (0))
(step t16 (cl (= e0 (op e2 (op e2 e0))) (not (and (= e0 (op e2 (op e2 e0))) (not (= e0 (op e2 e0)))))) :rule reordering :premises (t15))
(step t17 (cl (not (= (or (or (or (or (or (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit)))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit)))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit)))) (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit)))) (or (and (= e0 (op e5 (op e5 e0))) (not (= e0 (op e5 e0)))) (and (= e0 (op e4 (op e4 e0))) (not (= e0 (op e4 e0)))) (and (= e0 (op e3 (op e3 e0))) (not (= e0 (op e3 e0)))) (and (= e0 (op e2 (op e2 e0))) (not (= e0 (op e2 e0)))) (and (= e0 (op e0 (op e0 e0))) (not (= e0 (op e0 e0)))) (and (= e0 (op e1 (op e1 e0))) (not (= e0 (op e1 e0))))))) (not (or (or (or (or (or (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit)))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit)))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit)))) (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit))))) (or (and (= e0 (op e5 (op e5 e0))) (not (= e0 (op e5 e0)))) (and (= e0 (op e4 (op e4 e0))) (not (= e0 (op e4 e0)))) (and (= e0 (op e3 (op e3 e0))) (not (= e0 (op e3 e0)))) (and (= e0 (op e2 (op e2 e0))) (not (= e0 (op e2 e0)))) (and (= e0 (op e0 (op e0 e0))) (not (= e0 (op e0 e0)))) (and (= e0 (op e1 (op e1 e0))) (not (= e0 (op e1 e0)))))) :rule equiv_pos2)
(step t18 (cl (= (or (or (or (or (or (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit)))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit)))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit)))) (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit)))) (or (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit))) (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))))) :rule hole :args ("ACI_NORM" (= (or (or (or (or (or (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit)))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit)))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit)))) (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit)))) (or (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit))) (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))))))
(step t19 (cl (= (= (op e5 (op e5 e0)) e0) (= e0 (op e5 (op e5 e0))))) :rule rare_rewrite :args ("eq-symm" (op e5 (op e5 e0)) e0))
(step t20 (cl (= (= (op e5 e0) unit) (= unit (op e5 e0)))) :rule rare_rewrite :args ("eq-symm" (op e5 e0) unit))
(step t21 (cl (= (not (= (op e5 e0) unit)) (not (= unit (op e5 e0))))) :rule cong :premises (t20))
(step t22 (cl (= (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit))) (and (= e0 (op e5 (op e5 e0))) (not (= unit (op e5 e0)))))) :rule cong :premises (t19 t21))
(step t23 (cl (= (= (op e4 (op e4 e0)) e0) (= e0 (op e4 (op e4 e0))))) :rule rare_rewrite :args ("eq-symm" (op e4 (op e4 e0)) e0))
(step t24 (cl (= (= (op e4 e0) unit) (= unit (op e4 e0)))) :rule rare_rewrite :args ("eq-symm" (op e4 e0) unit))
(step t25 (cl (= (not (= (op e4 e0) unit)) (not (= unit (op e4 e0))))) :rule cong :premises (t24))
(step t26 (cl (= (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit))) (and (= e0 (op e4 (op e4 e0))) (not (= unit (op e4 e0)))))) :rule cong :premises (t23 t25))
(step t27 (cl (= (= (op e3 (op e3 e0)) e0) (= e0 (op e3 (op e3 e0))))) :rule rare_rewrite :args ("eq-symm" (op e3 (op e3 e0)) e0))
(step t28 (cl (= (= (op e3 e0) unit) (= unit (op e3 e0)))) :rule rare_rewrite :args ("eq-symm" (op e3 e0) unit))
(step t29 (cl (= (not (= (op e3 e0) unit)) (not (= unit (op e3 e0))))) :rule cong :premises (t28))
(step t30 (cl (= (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit))) (and (= e0 (op e3 (op e3 e0))) (not (= unit (op e3 e0)))))) :rule cong :premises (t27 t29))
(step t31 (cl (= (= (op e2 (op e2 e0)) e0) (= e0 (op e2 (op e2 e0))))) :rule rare_rewrite :args ("eq-symm" (op e2 (op e2 e0)) e0))
(step t32 (cl (= (= (op e2 e0) unit) (= unit (op e2 e0)))) :rule rare_rewrite :args ("eq-symm" (op e2 e0) unit))
(step t33 (cl (= (not (= (op e2 e0) unit)) (not (= unit (op e2 e0))))) :rule cong :premises (t32))
(step t34 (cl (= (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit))) (and (= e0 (op e2 (op e2 e0))) (not (= unit (op e2 e0)))))) :rule cong :premises (t31 t33))
(step t35 (cl (= (= (op e0 (op e0 e0)) e0) (= e0 (op e0 (op e0 e0))))) :rule rare_rewrite :args ("eq-symm" (op e0 (op e0 e0)) e0))
(step t36 (cl (= (= (op e0 e0) unit) (= unit (op e0 e0)))) :rule rare_rewrite :args ("eq-symm" (op e0 e0) unit))
(step t37 (cl (= (not (= (op e0 e0) unit)) (not (= unit (op e0 e0))))) :rule cong :premises (t36))
(step t38 (cl (= (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= e0 (op e0 (op e0 e0))) (not (= unit (op e0 e0)))))) :rule cong :premises (t35 t37))
(step t39 (cl (= (= (op e1 (op e1 e0)) e0) (= e0 (op e1 (op e1 e0))))) :rule rare_rewrite :args ("eq-symm" (op e1 (op e1 e0)) e0))
(step t40 (cl (= (= (op e1 e0) unit) (= unit (op e1 e0)))) :rule rare_rewrite :args ("eq-symm" (op e1 e0) unit))
(step t41 (cl (= (not (= (op e1 e0) unit)) (not (= unit (op e1 e0))))) :rule cong :premises (t40))
(step t42 (cl (= (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit))) (and (= e0 (op e1 (op e1 e0))) (not (= unit (op e1 e0)))))) :rule cong :premises (t39 t41))
(step t43 (cl (= (or (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit))) (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))) (or (and (= e0 (op e5 (op e5 e0))) (not (= unit (op e5 e0)))) (and (= e0 (op e4 (op e4 e0))) (not (= unit (op e4 e0)))) (and (= e0 (op e3 (op e3 e0))) (not (= unit (op e3 e0)))) (and (= e0 (op e2 (op e2 e0))) (not (= unit (op e2 e0)))) (and (= e0 (op e0 (op e0 e0))) (not (= unit (op e0 e0)))) (and (= e0 (op e1 (op e1 e0))) (not (= unit (op e1 e0))))))) :rule cong :premises (t22 t26 t30 t34 t38 t42))
(step t44 (cl (= (or (or (or (or (or (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit)))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit)))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit)))) (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit)))) (or (and (= e0 (op e5 (op e5 e0))) (not (= unit (op e5 e0)))) (and (= e0 (op e4 (op e4 e0))) (not (= unit (op e4 e0)))) (and (= e0 (op e3 (op e3 e0))) (not (= unit (op e3 e0)))) (and (= e0 (op e2 (op e2 e0))) (not (= unit (op e2 e0)))) (and (= e0 (op e0 (op e0 e0))) (not (= unit (op e0 e0)))) (and (= e0 (op e1 (op e1 e0))) (not (= unit (op e1 e0))))))) :rule trans :premises (t18 t43))
(step t45 (cl (= (= e0 (op e5 (op e5 e0))) (= e0 (op e5 (op e5 e0))))) :rule refl)
(step t46 (cl (= (op e5 e0) (op e5 e0))) :rule refl)
(step t47 (cl (= (= unit (op e5 e0)) (= e0 (op e5 e0)))) :rule cong :premises (a8 t46))
(step t48 (cl (= (not (= unit (op e5 e0))) (not (= e0 (op e5 e0))))) :rule cong :premises (t47))
(step t49 (cl (= (and (= e0 (op e5 (op e5 e0))) (not (= unit (op e5 e0)))) (and (= e0 (op e5 (op e5 e0))) (not (= e0 (op e5 e0)))))) :rule cong :premises (t45 t48))
(step t50 (cl (= (= e0 (op e4 (op e4 e0))) (= e0 (op e4 (op e4 e0))))) :rule refl)
(step t51 (cl (= (op e4 e0) (op e4 e0))) :rule refl)
(step t52 (cl (= (= unit (op e4 e0)) (= e0 (op e4 e0)))) :rule cong :premises (a8 t51))
(step t53 (cl (= (not (= unit (op e4 e0))) (not (= e0 (op e4 e0))))) :rule cong :premises (t52))
(step t54 (cl (= (and (= e0 (op e4 (op e4 e0))) (not (= unit (op e4 e0)))) (and (= e0 (op e4 (op e4 e0))) (not (= e0 (op e4 e0)))))) :rule cong :premises (t50 t53))
(step t55 (cl (= (= e0 (op e3 (op e3 e0))) (= e0 (op e3 (op e3 e0))))) :rule refl)
(step t56 (cl (= (op e3 e0) (op e3 e0))) :rule refl)
(step t57 (cl (= (= unit (op e3 e0)) (= e0 (op e3 e0)))) :rule cong :premises (a8 t56))
(step t58 (cl (= (not (= unit (op e3 e0))) (not (= e0 (op e3 e0))))) :rule cong :premises (t57))
(step t59 (cl (= (and (= e0 (op e3 (op e3 e0))) (not (= unit (op e3 e0)))) (and (= e0 (op e3 (op e3 e0))) (not (= e0 (op e3 e0)))))) :rule cong :premises (t55 t58))
(step t60 (cl (= (= e0 (op e2 (op e2 e0))) (= e0 (op e2 (op e2 e0))))) :rule refl)
(step t61 (cl (= (op e2 e0) (op e2 e0))) :rule refl)
(step t62 (cl (= (= unit (op e2 e0)) (= e0 (op e2 e0)))) :rule cong :premises (a8 t61))
(step t63 (cl (= (not (= unit (op e2 e0))) (not (= e0 (op e2 e0))))) :rule cong :premises (t62))
(step t64 (cl (= (and (= e0 (op e2 (op e2 e0))) (not (= unit (op e2 e0)))) (and (= e0 (op e2 (op e2 e0))) (not (= e0 (op e2 e0)))))) :rule cong :premises (t60 t63))
(step t65 (cl (= (= e0 (op e0 (op e0 e0))) (= e0 (op e0 (op e0 e0))))) :rule refl)
(step t66 (cl (= (op e0 e0) (op e0 e0))) :rule refl)
(step t67 (cl (= (= unit (op e0 e0)) (= e0 (op e0 e0)))) :rule cong :premises (a8 t66))
(step t68 (cl (= (not (= unit (op e0 e0))) (not (= e0 (op e0 e0))))) :rule cong :premises (t67))
(step t69 (cl (= (and (= e0 (op e0 (op e0 e0))) (not (= unit (op e0 e0)))) (and (= e0 (op e0 (op e0 e0))) (not (= e0 (op e0 e0)))))) :rule cong :premises (t65 t68))
(step t70 (cl (= (= e0 (op e1 (op e1 e0))) (= e0 (op e1 (op e1 e0))))) :rule refl)
(step t71 (cl (= (op e1 e0) (op e1 e0))) :rule refl)
(step t72 (cl (= (= unit (op e1 e0)) (= e0 (op e1 e0)))) :rule cong :premises (a8 t71))
(step t73 (cl (= (not (= unit (op e1 e0))) (not (= e0 (op e1 e0))))) :rule cong :premises (t72))
(step t74 (cl (= (and (= e0 (op e1 (op e1 e0))) (not (= unit (op e1 e0)))) (and (= e0 (op e1 (op e1 e0))) (not (= e0 (op e1 e0)))))) :rule cong :premises (t70 t73))
(step t75 (cl (= (or (and (= e0 (op e5 (op e5 e0))) (not (= unit (op e5 e0)))) (and (= e0 (op e4 (op e4 e0))) (not (= unit (op e4 e0)))) (and (= e0 (op e3 (op e3 e0))) (not (= unit (op e3 e0)))) (and (= e0 (op e2 (op e2 e0))) (not (= unit (op e2 e0)))) (and (= e0 (op e0 (op e0 e0))) (not (= unit (op e0 e0)))) (and (= e0 (op e1 (op e1 e0))) (not (= unit (op e1 e0))))) (or (and (= e0 (op e5 (op e5 e0))) (not (= e0 (op e5 e0)))) (and (= e0 (op e4 (op e4 e0))) (not (= e0 (op e4 e0)))) (and (= e0 (op e3 (op e3 e0))) (not (= e0 (op e3 e0)))) (and (= e0 (op e2 (op e2 e0))) (not (= e0 (op e2 e0)))) (and (= e0 (op e0 (op e0 e0))) (not (= e0 (op e0 e0)))) (and (= e0 (op e1 (op e1 e0))) (not (= e0 (op e1 e0))))))) :rule cong :premises (t49 t54 t59 t64 t69 t74))
(step t76 (cl (= (or (or (or (or (or (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit)))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit)))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit)))) (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit)))) (or (and (= e0 (op e5 (op e5 e0))) (not (= e0 (op e5 e0)))) (and (= e0 (op e4 (op e4 e0))) (not (= e0 (op e4 e0)))) (and (= e0 (op e3 (op e3 e0))) (not (= e0 (op e3 e0)))) (and (= e0 (op e2 (op e2 e0))) (not (= e0 (op e2 e0)))) (and (= e0 (op e0 (op e0 e0))) (not (= e0 (op e0 e0)))) (and (= e0 (op e1 (op e1 e0))) (not (= e0 (op e1 e0))))))) :rule trans :premises (t44 t75))
(step t77 (cl (and (and (and (and (and (or (or (or (or (or (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit)))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit)))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit)))) (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit)))) (or (or (or (or (or (and (= (op e0 (op e0 e1)) e1) (not (= (op e0 e1) unit))) (and (= (op e1 (op e1 e1)) e1) (not (= (op e1 e1) unit)))) (and (= (op e2 (op e2 e1)) e1) (not (= (op e2 e1) unit)))) (and (= (op e3 (op e3 e1)) e1) (not (= (op e3 e1) unit)))) (and (= (op e4 (op e4 e1)) e1) (not (= (op e4 e1) unit)))) (and (= (op e5 (op e5 e1)) e1) (not (= (op e5 e1) unit))))) (or (or (or (or (or (and (= (op e0 (op e0 e2)) e2) (not (= (op e0 e2) unit))) (and (= (op e1 (op e1 e2)) e2) (not (= (op e1 e2) unit)))) (and (= (op e2 (op e2 e2)) e2) (not (= (op e2 e2) unit)))) (and (= (op e3 (op e3 e2)) e2) (not (= (op e3 e2) unit)))) (and (= (op e4 (op e4 e2)) e2) (not (= (op e4 e2) unit)))) (and (= (op e5 (op e5 e2)) e2) (not (= (op e5 e2) unit))))) (or (or (or (or (or (and (= (op e0 (op e0 e3)) e3) (not (= (op e0 e3) unit))) (and (= (op e1 (op e1 e3)) e3) (not (= (op e1 e3) unit)))) (and (= (op e2 (op e2 e3)) e3) (not (= (op e2 e3) unit)))) (and (= (op e3 (op e3 e3)) e3) (not (= (op e3 e3) unit)))) (and (= (op e4 (op e4 e3)) e3) (not (= (op e4 e3) unit)))) (and (= (op e5 (op e5 e3)) e3) (not (= (op e5 e3) unit))))) (or (or (or (or (or (and (= (op e0 (op e0 e4)) e4) (not (= (op e0 e4) unit))) (and (= (op e1 (op e1 e4)) e4) (not (= (op e1 e4) unit)))) (and (= (op e2 (op e2 e4)) e4) (not (= (op e2 e4) unit)))) (and (= (op e3 (op e3 e4)) e4) (not (= (op e3 e4) unit)))) (and (= (op e4 (op e4 e4)) e4) (not (= (op e4 e4) unit)))) (and (= (op e5 (op e5 e4)) e4) (not (= (op e5 e4) unit))))) (or (or (or (or (or (and (= (op e0 (op e0 e5)) e5) (not (= (op e0 e5) unit))) (and (= (op e1 (op e1 e5)) e5) (not (= (op e1 e5) unit)))) (and (= (op e2 (op e2 e5)) e5) (not (= (op e2 e5) unit)))) (and (= (op e3 (op e3 e5)) e5) (not (= (op e3 e5) unit)))) (and (= (op e4 (op e4 e5)) e5) (not (= (op e4 e5) unit)))) (and (= (op e5 (op e5 e5)) e5) (not (= (op e5 e5) unit)))))) :rule and :premises (a7) :args (1))
(step t78 (cl (and (and (and (and (or (or (or (or (or (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit)))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit)))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit)))) (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit)))) (or (or (or (or (or (and (= (op e0 (op e0 e1)) e1) (not (= (op e0 e1) unit))) (and (= (op e1 (op e1 e1)) e1) (not (= (op e1 e1) unit)))) (and (= (op e2 (op e2 e1)) e1) (not (= (op e2 e1) unit)))) (and (= (op e3 (op e3 e1)) e1) (not (= (op e3 e1) unit)))) (and (= (op e4 (op e4 e1)) e1) (not (= (op e4 e1) unit)))) (and (= (op e5 (op e5 e1)) e1) (not (= (op e5 e1) unit))))) (or (or (or (or (or (and (= (op e0 (op e0 e2)) e2) (not (= (op e0 e2) unit))) (and (= (op e1 (op e1 e2)) e2) (not (= (op e1 e2) unit)))) (and (= (op e2 (op e2 e2)) e2) (not (= (op e2 e2) unit)))) (and (= (op e3 (op e3 e2)) e2) (not (= (op e3 e2) unit)))) (and (= (op e4 (op e4 e2)) e2) (not (= (op e4 e2) unit)))) (and (= (op e5 (op e5 e2)) e2) (not (= (op e5 e2) unit))))) (or (or (or (or (or (and (= (op e0 (op e0 e3)) e3) (not (= (op e0 e3) unit))) (and (= (op e1 (op e1 e3)) e3) (not (= (op e1 e3) unit)))) (and (= (op e2 (op e2 e3)) e3) (not (= (op e2 e3) unit)))) (and (= (op e3 (op e3 e3)) e3) (not (= (op e3 e3) unit)))) (and (= (op e4 (op e4 e3)) e3) (not (= (op e4 e3) unit)))) (and (= (op e5 (op e5 e3)) e3) (not (= (op e5 e3) unit))))) (or (or (or (or (or (and (= (op e0 (op e0 e4)) e4) (not (= (op e0 e4) unit))) (and (= (op e1 (op e1 e4)) e4) (not (= (op e1 e4) unit)))) (and (= (op e2 (op e2 e4)) e4) (not (= (op e2 e4) unit)))) (and (= (op e3 (op e3 e4)) e4) (not (= (op e3 e4) unit)))) (and (= (op e4 (op e4 e4)) e4) (not (= (op e4 e4) unit)))) (and (= (op e5 (op e5 e4)) e4) (not (= (op e5 e4) unit)))))) :rule and :premises (t77) :args (0))
(step t79 (cl (and (and (and (or (or (or (or (or (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit)))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit)))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit)))) (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit)))) (or (or (or (or (or (and (= (op e0 (op e0 e1)) e1) (not (= (op e0 e1) unit))) (and (= (op e1 (op e1 e1)) e1) (not (= (op e1 e1) unit)))) (and (= (op e2 (op e2 e1)) e1) (not (= (op e2 e1) unit)))) (and (= (op e3 (op e3 e1)) e1) (not (= (op e3 e1) unit)))) (and (= (op e4 (op e4 e1)) e1) (not (= (op e4 e1) unit)))) (and (= (op e5 (op e5 e1)) e1) (not (= (op e5 e1) unit))))) (or (or (or (or (or (and (= (op e0 (op e0 e2)) e2) (not (= (op e0 e2) unit))) (and (= (op e1 (op e1 e2)) e2) (not (= (op e1 e2) unit)))) (and (= (op e2 (op e2 e2)) e2) (not (= (op e2 e2) unit)))) (and (= (op e3 (op e3 e2)) e2) (not (= (op e3 e2) unit)))) (and (= (op e4 (op e4 e2)) e2) (not (= (op e4 e2) unit)))) (and (= (op e5 (op e5 e2)) e2) (not (= (op e5 e2) unit))))) (or (or (or (or (or (and (= (op e0 (op e0 e3)) e3) (not (= (op e0 e3) unit))) (and (= (op e1 (op e1 e3)) e3) (not (= (op e1 e3) unit)))) (and (= (op e2 (op e2 e3)) e3) (not (= (op e2 e3) unit)))) (and (= (op e3 (op e3 e3)) e3) (not (= (op e3 e3) unit)))) (and (= (op e4 (op e4 e3)) e3) (not (= (op e4 e3) unit)))) (and (= (op e5 (op e5 e3)) e3) (not (= (op e5 e3) unit)))))) :rule and :premises (t78) :args (0))
(step t80 (cl (and (and (or (or (or (or (or (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit)))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit)))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit)))) (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit)))) (or (or (or (or (or (and (= (op e0 (op e0 e1)) e1) (not (= (op e0 e1) unit))) (and (= (op e1 (op e1 e1)) e1) (not (= (op e1 e1) unit)))) (and (= (op e2 (op e2 e1)) e1) (not (= (op e2 e1) unit)))) (and (= (op e3 (op e3 e1)) e1) (not (= (op e3 e1) unit)))) (and (= (op e4 (op e4 e1)) e1) (not (= (op e4 e1) unit)))) (and (= (op e5 (op e5 e1)) e1) (not (= (op e5 e1) unit))))) (or (or (or (or (or (and (= (op e0 (op e0 e2)) e2) (not (= (op e0 e2) unit))) (and (= (op e1 (op e1 e2)) e2) (not (= (op e1 e2) unit)))) (and (= (op e2 (op e2 e2)) e2) (not (= (op e2 e2) unit)))) (and (= (op e3 (op e3 e2)) e2) (not (= (op e3 e2) unit)))) (and (= (op e4 (op e4 e2)) e2) (not (= (op e4 e2) unit)))) (and (= (op e5 (op e5 e2)) e2) (not (= (op e5 e2) unit)))))) :rule and :premises (t79) :args (0))
(step t81 (cl (and (or (or (or (or (or (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit)))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit)))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit)))) (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit)))) (or (or (or (or (or (and (= (op e0 (op e0 e1)) e1) (not (= (op e0 e1) unit))) (and (= (op e1 (op e1 e1)) e1) (not (= (op e1 e1) unit)))) (and (= (op e2 (op e2 e1)) e1) (not (= (op e2 e1) unit)))) (and (= (op e3 (op e3 e1)) e1) (not (= (op e3 e1) unit)))) (and (= (op e4 (op e4 e1)) e1) (not (= (op e4 e1) unit)))) (and (= (op e5 (op e5 e1)) e1) (not (= (op e5 e1) unit)))))) :rule and :premises (t80) :args (0))
(step t82 (cl (or (or (or (or (or (and (= (op e0 (op e0 e0)) e0) (not (= (op e0 e0) unit))) (and (= (op e1 (op e1 e0)) e0) (not (= (op e1 e0) unit)))) (and (= (op e2 (op e2 e0)) e0) (not (= (op e2 e0) unit)))) (and (= (op e3 (op e3 e0)) e0) (not (= (op e3 e0) unit)))) (and (= (op e4 (op e4 e0)) e0) (not (= (op e4 e0) unit)))) (and (= (op e5 (op e5 e0)) e0) (not (= (op e5 e0) unit))))) :rule and :premises (t81) :args (0))
(step t83 (cl (or (and (= e0 (op e5 (op e5 e0))) (not (= e0 (op e5 e0)))) (and (= e0 (op e4 (op e4 e0))) (not (= e0 (op e4 e0)))) (and (= e0 (op e3 (op e3 e0))) (not (= e0 (op e3 e0)))) (and (= e0 (op e2 (op e2 e0))) (not (= e0 (op e2 e0)))) (and (= e0 (op e0 (op e0 e0))) (not (= e0 (op e0 e0)))) (and (= e0 (op e1 (op e1 e0))) (not (= e0 (op e1 e0)))))) :rule resolution :premises (t17 t76 t82))
(step t84 (cl (and (= e0 (op e5 (op e5 e0))) (not (= e0 (op e5 e0)))) (and (= e0 (op e4 (op e4 e0))) (not (= e0 (op e4 e0)))) (and (= e0 (op e3 (op e3 e0))) (not (= e0 (op e3 e0)))) (and (= e0 (op e2 (op e2 e0))) (not (= e0 (op e2 e0)))) (and (= e0 (op e0 (op e0 e0))) (not (= e0 (op e0 e0)))) (and (= e0 (op e1 (op e1 e0))) (not (= e0 (op e1 e0))))) :rule or :premises (t83))
(step t85 (cl (and (= e0 (op e1 (op e1 e0))) (not (= e0 (op e1 e0)))) (and (= e0 (op e0 (op e0 e0))) (not (= e0 (op e0 e0)))) (and (= e0 (op e2 (op e2 e0))) (not (= e0 (op e2 e0)))) (and (= e0 (op e3 (op e3 e0))) (not (= e0 (op e3 e0)))) (and (= e0 (op e4 (op e4 e0))) (not (= e0 (op e4 e0)))) (and (= e0 (op e5 (op e5 e0))) (not (= e0 (op e5 e0))))) :rule reordering :premises (t84))
(step t86 (cl (not (and (= e0 (op e1 (op e1 e0))) (not (= e0 (op e1 e0))))) (= e0 (op e1 (op e1 e0)))) :rule and_pos :args (0))
(step t87 (cl (= e0 (op e1 (op e1 e0))) (not (and (= e0 (op e1 (op e1 e0))) (not (= e0 (op e1 e0)))))) :rule reordering :premises (t86))
(step t88 (cl (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0)))) (not (= e1 (op e1 e0))) (not (= e0 (op e1 (op e1 e0))))) :rule and_neg)
(step t89 (cl (=> (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0)))) (= e1 (op e1 (op e1 e1)))) (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0))))) :rule implies_neg1)
(anchor :step t90)
(assume t90.a0 (= e1 (op e1 e0)))
(assume t90.a1 (= e0 (op e1 (op e1 e0))))
(step t90.t0 (cl (not (= (= (op e1 unit) e1) (= e1 (op e1 e0)))) (not (= (op e1 unit) e1)) (= e1 (op e1 e0))) :rule equiv_pos2)
(step t90.t1 (cl (= (= (op e1 unit) e1) (= e1 (op e1 unit)))) :rule rare_rewrite :args ("eq-symm" (op e1 unit) e1))
(step t90.t2 (cl (= e1 e1)) :rule refl)
(step t90.t3 (cl (= (op e1 unit) (op e1 e0))) :rule cong :premises (t90.t2 a8))
(step t90.t4 (cl (= (= e1 (op e1 unit)) (= e1 (op e1 e0)))) :rule cong :premises (t90.t2 t90.t3))
(step t90.t5 (cl (= (= (op e1 unit) e1) (= e1 (op e1 e0)))) :rule trans :premises (t90.t1 t90.t4))
(step t90.t6 (cl (and (and (and (and (and (and (= (op unit e0) e0) (= (op e0 unit) e0)) (and (= (op unit e1) e1) (= (op e1 unit) e1))) (and (= (op unit e2) e2) (= (op e2 unit) e2))) (and (= (op unit e3) e3) (= (op e3 unit) e3))) (and (= (op unit e4) e4) (= (op e4 unit) e4))) (and (= (op unit e5) e5) (= (op e5 unit) e5)))) :rule and :premises (a2) :args (0))
(step t90.t7 (cl (and (and (and (and (and (= (op unit e0) e0) (= (op e0 unit) e0)) (and (= (op unit e1) e1) (= (op e1 unit) e1))) (and (= (op unit e2) e2) (= (op e2 unit) e2))) (and (= (op unit e3) e3) (= (op e3 unit) e3))) (and (= (op unit e4) e4) (= (op e4 unit) e4)))) :rule and :premises (t90.t6) :args (0))
(step t90.t8 (cl (and (and (and (and (= (op unit e0) e0) (= (op e0 unit) e0)) (and (= (op unit e1) e1) (= (op e1 unit) e1))) (and (= (op unit e2) e2) (= (op e2 unit) e2))) (and (= (op unit e3) e3) (= (op e3 unit) e3)))) :rule and :premises (t90.t7) :args (0))
(step t90.t9 (cl (and (and (and (= (op unit e0) e0) (= (op e0 unit) e0)) (and (= (op unit e1) e1) (= (op e1 unit) e1))) (and (= (op unit e2) e2) (= (op e2 unit) e2)))) :rule and :premises (t90.t8) :args (0))
(step t90.t10 (cl (and (and (= (op unit e0) e0) (= (op e0 unit) e0)) (and (= (op unit e1) e1) (= (op e1 unit) e1)))) :rule and :premises (t90.t9) :args (0))
(step t90.t11 (cl (and (= (op unit e1) e1) (= (op e1 unit) e1))) :rule and :premises (t90.t10) :args (1))
(step t90.t12 (cl (= (op e1 unit) e1)) :rule and :premises (t90.t11) :args (1))
(step t90.t13 (cl (= e1 (op e1 e0))) :rule resolution :premises (t90.t0 t90.t5 t90.t12))
(step t90.t14 (cl (= (op e1 (op e1 e0)) e0)) :rule symm :premises (t90.a1))
(step t90.t15 (cl (= e0 (op e1 (op e1 e0)))) :rule symm :premises (t90.t14))
(step t90.t16 (cl (= (op e1 e0) e1)) :rule symm :premises (t90.t13))
(step t90.t17 (cl (= (op e1 (op e1 e0)) (op e1 e1))) :rule cong :premises (t90.t2 t90.t16))
(step t90.t18 (cl (= e0 (op e1 e1))) :rule trans :premises (t90.t15 t90.t17))
(step t90.t19 (cl (= (op e1 e0) (op e1 (op e1 e1)))) :rule cong :premises (t90.t2 t90.t18))
(step t90.t20 (cl (= e1 (op e1 (op e1 e1)))) :rule trans :premises (t90.t13 t90.t19))
(step t90 (cl (not (= e1 (op e1 e0))) (not (= e0 (op e1 (op e1 e0)))) (= e1 (op e1 (op e1 e1)))) :rule subproof :discharge (t90.a0 t90.a1))
(step t91 (cl (not (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0))))) (= e1 (op e1 e0))) :rule and_pos :args (0))
(step t92 (cl (not (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0))))) (= e0 (op e1 (op e1 e0)))) :rule and_pos :args (1))
(step t93 (cl (= e1 (op e1 (op e1 e1))) (not (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0))))) (not (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0)))))) :rule resolution :premises (t90 t91 t92))
(step t94 (cl (not (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0))))) (not (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0))))) (= e1 (op e1 (op e1 e1)))) :rule reordering :premises (t93))
(step t95 (cl (not (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0))))) (= e1 (op e1 (op e1 e1)))) :rule contraction :premises (t94))
(step t96 (cl (=> (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0)))) (= e1 (op e1 (op e1 e1)))) (= e1 (op e1 (op e1 e1)))) :rule resolution :premises (t89 t95))
(step t97 (cl (=> (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0)))) (= e1 (op e1 (op e1 e1)))) (not (= e1 (op e1 (op e1 e1))))) :rule implies_neg2)
(step t98 (cl (=> (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0)))) (= e1 (op e1 (op e1 e1)))) (=> (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0)))) (= e1 (op e1 (op e1 e1))))) :rule resolution :premises (t96 t97))
(step t99 (cl (=> (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0)))) (= e1 (op e1 (op e1 e1))))) :rule contraction :premises (t98))
(step t100 (cl (not (and (= e1 (op e1 e0)) (= e0 (op e1 (op e1 e0))))) (= e1 (op e1 (op e1 e1)))) :rule implies :premises (t99))
(step t101 (cl (not (= e1 (op e1 e0))) (not (= e0 (op e1 (op e1 e0)))) (= e1 (op e1 (op e1 e1)))) :rule resolution :premises (t88 t100))
(step t102 (cl (= e1 (op e1 (op e1 e1))) (not (= e1 (op e1 e0))) (not (= e0 (op e1 (op e1 e0))))) :rule reordering :premises (t101))
(step t103 (cl (not (= (or (not (= (op e1 (op e1 e1)) e1)) (= e1 unit)) (or (not (= e1 (op e1 (op e1 e1)))) (= e1 e0)))) (not (or (not (= (op e1 (op e1 e1)) e1)) (= e1 unit))) (or (not (= e1 (op e1 (op e1 e1)))) (= e1 e0))) :rule equiv_pos2)
(step t104 (cl (= (= (op e1 (op e1 e1)) e1) (= e1 (op e1 (op e1 e1))))) :rule rare_rewrite :args ("eq-symm" (op e1 (op e1 e1)) e1))
(step t105 (cl (= (not (= (op e1 (op e1 e1)) e1)) (not (= e1 (op e1 (op e1 e1)))))) :rule cong :premises (t104))
(step t106 (cl (= (= e1 unit) (= unit e1))) :rule rare_rewrite :args ("eq-symm" e1 unit))
(step t107 (cl (= (or (not (= (op e1 (op e1 e1)) e1)) (= e1 unit)) (or (not (= e1 (op e1 (op e1 e1)))) (= unit e1)))) :rule cong :premises (t105 t106))
(step t108 (cl (= (not (= e1 (op e1 (op e1 e1)))) (not (= e1 (op e1 (op e1 e1)))))) :rule refl)
(step t109 (cl (= e1 e1)) :rule refl)
(step t110 (cl (= (= unit e1) (= e0 e1))) :rule cong :premises (a8 t109))
(step t111 (cl (= (or (not (= e1 (op e1 (op e1 e1)))) (= unit e1)) (or (not (= e1 (op e1 (op e1 e1)))) (= e0 e1)))) :rule cong :premises (t108 t110))
(step t112 (cl (= (= e0 e1) (= e1 e0))) :rule rare_rewrite :args ("eq-symm" e0 e1))
(step t113 (cl (= (or (not (= e1 (op e1 (op e1 e1)))) (= e0 e1)) (or (not (= e1 (op e1 (op e1 e1)))) (= e1 e0)))) :rule cong :premises (t108 t112))
(step t114 (cl (= (or (not (= e1 (op e1 (op e1 e1)))) (= unit e1)) (or (not (= e1 (op e1 (op e1 e1)))) (= e1 e0)))) :rule trans :premises (t111 t113))
(step t115 (cl (= (or (not (= (op e1 (op e1 e1)) e1)) (= e1 unit)) (or (not (= e1 (op e1 (op e1 e1)))) (= e1 e0)))) :rule trans :premises (t107 t114))
(step t116 (cl (and (and (and (and (and (or (not (= (op e0 (op e0 e0)) e0)) (= e0 unit)) (or (not (= (op e1 (op e1 e1)) e1)) (= e1 unit))) (or (not (= (op e2 (op e2 e2)) e2)) (= e2 unit))) (or (not (= (op e3 (op e3 e3)) e3)) (= e3 unit))) (or (not (= (op e4 (op e4 e4)) e4)) (= e4 unit))) (or (not (= (op e5 (op e5 e5)) e5)) (= e5 unit)))) :rule and :premises (a7) :args (0))
(step t117 (cl (and (and (and (and (or (not (= (op e0 (op e0 e0)) e0)) (= e0 unit)) (or (not (= (op e1 (op e1 e1)) e1)) (= e1 unit))) (or (not (= (op e2 (op e2 e2)) e2)) (= e2 unit))) (or (not (= (op e3 (op e3 e3)) e3)) (= e3 unit))) (or (not (= (op e4 (op e4 e4)) e4)) (= e4 unit)))) :rule and :premises (t116) :args (0))
(step t118 (cl (and (and (and (or (not (= (op e0 (op e0 e0)) e0)) (= e0 unit)) (or (not (= (op e1 (op e1 e1)) e1)) (= e1 unit))) (or (not (= (op e2 (op e2 e2)) e2)) (= e2 unit))) (or (not (= (op e3 (op e3 e3)) e3)) (= e3 unit)))) :rule and :premises (t117) :args (0))
(step t119 (cl (and (and (or (not (= (op e0 (op e0 e0)) e0)) (= e0 unit)) (or (not (= (op e1 (op e1 e1)) e1)) (= e1 unit))) (or (not (= (op e2 (op e2 e2)) e2)) (= e2 unit)))) :rule and :premises (t118) :args (0))
(step t120 (cl (and (or (not (= (op e0 (op e0 e0)) e0)) (= e0 unit)) (or (not (= (op e1 (op e1 e1)) e1)) (= e1 unit)))) :rule and :premises (t119) :args (0))
(step t121 (cl (or (not (= (op e1 (op e1 e1)) e1)) (= e1 unit))) :rule and :premises (t120) :args (1))
(step t122 (cl (or (not (= e1 (op e1 (op e1 e1)))) (= e1 e0))) :rule resolution :premises (t103 t115 t121))
(step t123 (cl (not (= e1 (op e1 (op e1 e1)))) (= e1 e0)) :rule or :premises (t122))
(step t124 (cl (= e1 e0) (not (= e1 (op e1 (op e1 e1))))) :rule reordering :premises (t123))
(step t125 (cl (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= e0 e1)) (not (= e0 e2))) (not (= e0 e3))) (not (= e0 e4))) (not (= e0 e5))) (not (= e1 e2))) (not (= e1 e3))) (not (= e1 e4))) (not (= e1 e5))) (not (= e2 e3))) (not (= e2 e4))) (not (= e2 e5))) (not (= e3 e4))) (not (= e3 e5)))) :rule and :premises (a10) :args (0))
(step t126 (cl (and (and (and (and (and (and (and (and (and (and (and (and (not (= e0 e1)) (not (= e0 e2))) (not (= e0 e3))) (not (= e0 e4))) (not (= e0 e5))) (not (= e1 e2))) (not (= e1 e3))) (not (= e1 e4))) (not (= e1 e5))) (not (= e2 e3))) (not (= e2 e4))) (not (= e2 e5))) (not (= e3 e4)))) :rule and :premises (t125) :args (0))
(step t127 (cl (and (and (and (and (and (and (and (and (and (and (and (not (= e0 e1)) (not (= e0 e2))) (not (= e0 e3))) (not (= e0 e4))) (not (= e0 e5))) (not (= e1 e2))) (not (= e1 e3))) (not (= e1 e4))) (not (= e1 e5))) (not (= e2 e3))) (not (= e2 e4))) (not (= e2 e5)))) :rule and :premises (t126) :args (0))
(step t128 (cl (and (and (and (and (and (and (and (and (and (and (not (= e0 e1)) (not (= e0 e2))) (not (= e0 e3))) (not (= e0 e4))) (not (= e0 e5))) (not (= e1 e2))) (not (= e1 e3))) (not (= e1 e4))) (not (= e1 e5))) (not (= e2 e3))) (not (= e2 e4)))) :rule and :premises (t127) :args (0))
(step t129 (cl (and (and (and (and (and (and (and (and (and (not (= e0 e1)) (not (= e0 e2))) (not (= e0 e3))) (not (= e0 e4))) (not (= e0 e5))) (not (= e1 e2))) (not (= e1 e3))) (not (= e1 e4))) (not (= e1 e5))) (not (= e2 e3)))) :rule and :premises (t128) :args (0))
(step t130 (cl (and (and (and (and (and (and (and (and (not (= e0 e1)) (not (= e0 e2))) (not (= e0 e3))) (not (= e0 e4))) (not (= e0 e5))) (not (= e1 e2))) (not (= e1 e3))) (not (= e1 e4))) (not (= e1 e5)))) :rule and :premises (t129) :args (0))
(step t131 (cl (and (and (and (and (and (and (and (not (= e0 e1)) (not (= e0 e2))) (not (= e0 e3))) (not (= e0 e4))) (not (= e0 e5))) (not (= e1 e2))) (not (= e1 e3))) (not (= e1 e4)))) :rule and :premises (t130) :args (0))
(step t132 (cl (and (and (and (and (and (and (not (= e0 e1)) (not (= e0 e2))) (not (= e0 e3))) (not (= e0 e4))) (not (= e0 e5))) (not (= e1 e2))) (not (= e1 e3)))) :rule and :premises (t131) :args (0))
(step t133 (cl (and (and (and (and (and (not (= e0 e1)) (not (= e0 e2))) (not (= e0 e3))) (not (= e0 e4))) (not (= e0 e5))) (not (= e1 e2)))) :rule and :premises (t132) :args (0))
(step t134 (cl (and (and (and (and (not (= e0 e1)) (not (= e0 e2))) (not (= e0 e3))) (not (= e0 e4))) (not (= e0 e5)))) :rule and :premises (t133) :args (0))
(step t135 (cl (and (and (and (not (= e0 e1)) (not (= e0 e2))) (not (= e0 e3))) (not (= e0 e4)))) :rule and :premises (t134) :args (0))
(step t136 (cl (and (and (not (= e0 e1)) (not (= e0 e2))) (not (= e0 e3)))) :rule and :premises (t135) :args (0))
(step t137 (cl (and (not (= e0 e1)) (not (= e0 e2)))) :rule and :premises (t136) :args (0))
(step t138 (cl (not (= e0 e1))) :rule and :premises (t137) :args (0))
(step t139 (cl (not (= e1 e0))) :rule not_symm :premises (t138))
(step t140 (cl (not (= e1 (op e1 (op e1 e1))))) :rule resolution :premises (t124 t139))
(step t141 (cl (not (= (= (op e1 unit) e1) (= e1 (op e1 e0)))) (not (= (op e1 unit) e1)) (= e1 (op e1 e0))) :rule equiv_pos2)
(step t142 (cl (= (= (op e1 unit) e1) (= e1 (op e1 unit)))) :rule rare_rewrite :args ("eq-symm" (op e1 unit) e1))
(step t143 (cl (= (op e1 unit) (op e1 e0))) :rule cong :premises (t109 a8))
(step t144 (cl (= (= e1 (op e1 unit)) (= e1 (op e1 e0)))) :rule cong :premises (t109 t143))
(step t145 (cl (= (= (op e1 unit) e1) (= e1 (op e1 e0)))) :rule trans :premises (t142 t144))
(step t146 (cl (and (and (and (and (and (and (= (op unit e0) e0) (= (op e0 unit) e0)) (and (= (op unit e1) e1) (= (op e1 unit) e1))) (and (= (op unit e2) e2) (= (op e2 unit) e2))) (and (= (op unit e3) e3) (= (op e3 unit) e3))) (and (= (op unit e4) e4) (= (op e4 unit) e4))) (and (= (op unit e5) e5) (= (op e5 unit) e5)))) :rule and :premises (a2) :args (0))
(step t147 (cl (and (and (and (and (and (= (op unit e0) e0) (= (op e0 unit) e0)) (and (= (op unit e1) e1) (= (op e1 unit) e1))) (and (= (op unit e2) e2) (= (op e2 unit) e2))) (and (= (op unit e3) e3) (= (op e3 unit) e3))) (and (= (op unit e4) e4) (= (op e4 unit) e4)))) :rule and :premises (t146) :args (0))
(step t148 (cl (and (and (and (and (= (op unit e0) e0) (= (op e0 unit) e0)) (and (= (op unit e1) e1) (= (op e1 unit) e1))) (and (= (op unit e2) e2) (= (op e2 unit) e2))) (and (= (op unit e3) e3) (= (op e3 unit) e3)))) :rule and :premises (t147) :args (0))
(step t149 (cl (and (and (and (= (op unit e0) e0) (= (op e0 unit) e0)) (and (= (op unit e1) e1) (= (op e1 unit) e1))) (and (= (op unit e2) e2) (= (op e2 unit) e2)))) :rule and :premises (t148) :args (0))
(step t150 (cl (and (and (= (op unit e0) e0) (= (op e0 unit) e0)) (and (= (op unit e1) e1) (= (op e1 unit) e1)))) :rule and :premises (t149) :args (0))
(step t151 (cl (and (= (op unit e1) e1) (= (op e1 unit) e1))) :rule and :premises (t150) :args (1))
(step t152 (cl (= (op e1 unit) e1)) :rule and :premises (t151) :args (1))
(step t153 (cl (= e1 (op e1 e0))) :rule resolution :premises (t141 t145 t152))
(step t154 (cl (not (= e0 (op e1 (op e1 e0))))) :rule resolution :premises (t102 t140 t153))
(step t155 (cl (not (and (= e0 (op e1 (op e1 e0))) (not (= e0 (op e1 e0)))))) :rule resolution :premises (t87 t154))
(step t156 (cl (not (and (= e0 (op e0 (op e0 e0))) (not (= e0 (op e0 e0))))) (not (= e0 (op e0 e0)))) :rule and_pos :args (1))
(step t157 (cl (not (= e0 (op e0 e0))) (not (and (= e0 (op e0 (op e0 e0))) (not (= e0 (op e0 e0)))))) :rule reordering :premises (t156))
(step t158 (cl (not (= (= (op unit e0) e0) (= e0 (op e0 e0)))) (not (= (op unit e0) e0)) (= e0 (op e0 e0))) :rule equiv_pos2)
(step t159 (cl (= (= (op unit e0) e0) (= e0 (op unit e0)))) :rule rare_rewrite :args ("eq-symm" (op unit e0) e0))
(step t160 (cl (= e0 e0)) :rule refl)
(step t161 (cl (= (op unit e0) (op e0 e0))) :rule cong :premises (a8 t160))
(step t162 (cl (= (= e0 (op unit e0)) (= e0 (op e0 e0)))) :rule cong :premises (t160 t161))
(step t163 (cl (= (= (op unit e0) e0) (= e0 (op e0 e0)))) :rule trans :premises (t159 t162))
(step t164 (cl (and (= (op unit e0) e0) (= (op e0 unit) e0))) :rule and :premises (t150) :args (0))
(step t165 (cl (= (op unit e0) e0)) :rule and :premises (t164) :args (0))
(step t166 (cl (= e0 (op e0 e0))) :rule resolution :premises (t158 t163 t165))
(step t167 (cl (not (and (= e0 (op e0 (op e0 e0))) (not (= e0 (op e0 e0)))))) :rule resolution :premises (t157 t166))
(step t168 (cl (not (and (= e0 (op e3 (op e3 e0))) (not (= e0 (op e3 e0))))) (= e0 (op e3 (op e3 e0)))) :rule and_pos :args (0))
(step t169 (cl (= e0 (op e3 (op e3 e0))) (not (and (= e0 (op e3 (op e3 e0))) (not (= e0 (op e3 e0)))))) :rule reordering :premises (t168))
(step t170 (cl (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0)))) (not (= e3 (op e3 e0))) (not (= e0 (op e3 (op e3 e0))))) :rule and_neg)
(step t171 (cl (=> (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0)))) (= e3 (op e3 (op e3 e3)))) (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0))))) :rule implies_neg1)
(anchor :step t172)
(assume t172.a0 (= e3 (op e3 e0)))
(assume t172.a1 (= e0 (op e3 (op e3 e0))))
(step t172.t0 (cl (not (= (= (op e3 unit) e3) (= e3 (op e3 e0)))) (not (= (op e3 unit) e3)) (= e3 (op e3 e0))) :rule equiv_pos2)
(step t172.t1 (cl (= (= (op e3 unit) e3) (= e3 (op e3 unit)))) :rule rare_rewrite :args ("eq-symm" (op e3 unit) e3))
(step t172.t2 (cl (= e3 e3)) :rule refl)
(step t172.t3 (cl (= (op e3 unit) (op e3 e0))) :rule cong :premises (t172.t2 a8))
(step t172.t4 (cl (= (= e3 (op e3 unit)) (= e3 (op e3 e0)))) :rule cong :premises (t172.t2 t172.t3))
(step t172.t5 (cl (= (= (op e3 unit) e3) (= e3 (op e3 e0)))) :rule trans :premises (t172.t1 t172.t4))
(step t172.t6 (cl (and (= (op unit e3) e3) (= (op e3 unit) e3))) :rule and :premises (t148) :args (1))
(step t172.t7 (cl (= (op e3 unit) e3)) :rule and :premises (t172.t6) :args (1))
(step t172.t8 (cl (= e3 (op e3 e0))) :rule resolution :premises (t172.t0 t172.t5 t172.t7))
(step t172.t9 (cl (= (op e3 (op e3 e0)) e0)) :rule symm :premises (t172.a1))
(step t172.t10 (cl (= e0 (op e3 (op e3 e0)))) :rule symm :premises (t172.t9))
(step t172.t11 (cl (= (op e3 e0) e3)) :rule symm :premises (t172.t8))
(step t172.t12 (cl (= (op e3 (op e3 e0)) (op e3 e3))) :rule cong :premises (t172.t2 t172.t11))
(step t172.t13 (cl (= e0 (op e3 e3))) :rule trans :premises (t172.t10 t172.t12))
(step t172.t14 (cl (= (op e3 e0) (op e3 (op e3 e3)))) :rule cong :premises (t172.t2 t172.t13))
(step t172.t15 (cl (= e3 (op e3 (op e3 e3)))) :rule trans :premises (t172.t8 t172.t14))
(step t172 (cl (not (= e3 (op e3 e0))) (not (= e0 (op e3 (op e3 e0)))) (= e3 (op e3 (op e3 e3)))) :rule subproof :discharge (t172.a0 t172.a1))
(step t173 (cl (not (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0))))) (= e3 (op e3 e0))) :rule and_pos :args (0))
(step t174 (cl (not (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0))))) (= e0 (op e3 (op e3 e0)))) :rule and_pos :args (1))
(step t175 (cl (= e3 (op e3 (op e3 e3))) (not (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0))))) (not (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0)))))) :rule resolution :premises (t172 t173 t174))
(step t176 (cl (not (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0))))) (not (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0))))) (= e3 (op e3 (op e3 e3)))) :rule reordering :premises (t175))
(step t177 (cl (not (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0))))) (= e3 (op e3 (op e3 e3)))) :rule contraction :premises (t176))
(step t178 (cl (=> (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0)))) (= e3 (op e3 (op e3 e3)))) (= e3 (op e3 (op e3 e3)))) :rule resolution :premises (t171 t177))
(step t179 (cl (=> (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0)))) (= e3 (op e3 (op e3 e3)))) (not (= e3 (op e3 (op e3 e3))))) :rule implies_neg2)
(step t180 (cl (=> (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0)))) (= e3 (op e3 (op e3 e3)))) (=> (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0)))) (= e3 (op e3 (op e3 e3))))) :rule resolution :premises (t178 t179))
(step t181 (cl (=> (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0)))) (= e3 (op e3 (op e3 e3))))) :rule contraction :premises (t180))
(step t182 (cl (not (and (= e3 (op e3 e0)) (= e0 (op e3 (op e3 e0))))) (= e3 (op e3 (op e3 e3)))) :rule implies :premises (t181))
(step t183 (cl (not (= e3 (op e3 e0))) (not (= e0 (op e3 (op e3 e0)))) (= e3 (op e3 (op e3 e3)))) :rule resolution :premises (t170 t182))
(step t184 (cl (= e3 (op e3 (op e3 e3))) (not (= e3 (op e3 e0))) (not (= e0 (op e3 (op e3 e0))))) :rule reordering :premises (t183))
(step t185 (cl (not (= (or (not (= (op e3 (op e3 e3)) e3)) (= e3 unit)) (or (not (= e3 (op e3 (op e3 e3)))) (= e3 e0)))) (not (or (not (= (op e3 (op e3 e3)) e3)) (= e3 unit))) (or (not (= e3 (op e3 (op e3 e3)))) (= e3 e0))) :rule equiv_pos2)
(step t186 (cl (= (= (op e3 (op e3 e3)) e3) (= e3 (op e3 (op e3 e3))))) :rule rare_rewrite :args ("eq-symm" (op e3 (op e3 e3)) e3))
(step t187 (cl (= (not (= (op e3 (op e3 e3)) e3)) (not (= e3 (op e3 (op e3 e3)))))) :rule cong :premises (t186))
(step t188 (cl (= (= e3 unit) (= unit e3))) :rule rare_rewrite :args ("eq-symm" e3 unit))
(step t189 (cl (= (or (not (= (op e3 (op e3 e3)) e3)) (= e3 unit)) (or (not (= e3 (op e3 (op e3 e3)))) (= unit e3)))) :rule cong :premises (t187 t188))
(step t190 (cl (= (not (= e3 (op e3 (op e3 e3)))) (not (= e3 (op e3 (op e3 e3)))))) :rule refl)
(step t191 (cl (= e3 e3)) :rule refl)
(step t192 (cl (= (= unit e3) (= e0 e3))) :rule cong :premises (a8 t191))
(step t193 (cl (= (or (not (= e3 (op e3 (op e3 e3)))) (= unit e3)) (or (not (= e3 (op e3 (op e3 e3)))) (= e0 e3)))) :rule cong :premises (t190 t192))
(step t194 (cl (= (= e0 e3) (= e3 e0))) :rule rare_rewrite :args ("eq-symm" e0 e3))
(step t195 (cl (= (or (not (= e3 (op e3 (op e3 e3)))) (= e0 e3)) (or (not (= e3 (op e3 (op e3 e3)))) (= e3 e0)))) :rule cong :premises (t190 t194))
(step t196 (cl (= (or (not (= e3 (op e3 (op e3 e3)))) (= unit e3)) (or (not (= e3 (op e3 (op e3 e3)))) (= e3 e0)))) :rule trans :premises (t193 t195))
(step t197 (cl (= (or (not (= (op e3 (op e3 e3)) e3)) (= e3 unit)) (or (not (= e3 (op e3 (op e3 e3)))) (= e3 e0)))) :rule trans :premises (t189 t196))
(step t198 (cl (or (not (= (op e3 (op e3 e3)) e3)) (= e3 unit))) :rule and :premises (t118) :args (1))
(step t199 (cl (or (not (= e3 (op e3 (op e3 e3)))) (= e3 e0))) :rule resolution :premises (t185 t197 t198))
(step t200 (cl (not (= e3 (op e3 (op e3 e3)))) (= e3 e0)) :rule or :premises (t199))
(step t201 (cl (= e3 e0) (not (= e3 (op e3 (op e3 e3))))) :rule reordering :premises (t200))
(step t202 (cl (not (= e0 e3))) :rule and :premises (t136) :args (1))
(step t203 (cl (not (= e3 e0))) :rule not_symm :premises (t202))
(step t204 (cl (not (= e3 (op e3 (op e3 e3))))) :rule resolution :premises (t201 t203))
(step t205 (cl (not (= (= (op e3 unit) e3) (= e3 (op e3 e0)))) (not (= (op e3 unit) e3)) (= e3 (op e3 e0))) :rule equiv_pos2)
(step t206 (cl (= (= (op e3 unit) e3) (= e3 (op e3 unit)))) :rule rare_rewrite :args ("eq-symm" (op e3 unit) e3))
(step t207 (cl (= (op e3 unit) (op e3 e0))) :rule cong :premises (t191 a8))
(step t208 (cl (= (= e3 (op e3 unit)) (= e3 (op e3 e0)))) :rule cong :premises (t191 t207))
(step t209 (cl (= (= (op e3 unit) e3) (= e3 (op e3 e0)))) :rule trans :premises (t206 t208))
(step t210 (cl (and (= (op unit e3) e3) (= (op e3 unit) e3))) :rule and :premises (t148) :args (1))
(step t211 (cl (= (op e3 unit) e3)) :rule and :premises (t210) :args (1))
(step t212 (cl (= e3 (op e3 e0))) :rule resolution :premises (t205 t209 t211))
(step t213 (cl (not (= e0 (op e3 (op e3 e0))))) :rule resolution :premises (t184 t204 t212))
(step t214 (cl (not (and (= e0 (op e3 (op e3 e0))) (not (= e0 (op e3 e0)))))) :rule resolution :premises (t169 t213))
(step t215 (cl (not (and (= e0 (op e4 (op e4 e0))) (not (= e0 (op e4 e0))))) (= e0 (op e4 (op e4 e0)))) :rule and_pos :args (0))
(step t216 (cl (= e0 (op e4 (op e4 e0))) (not (and (= e0 (op e4 (op e4 e0))) (not (= e0 (op e4 e0)))))) :rule reordering :premises (t215))
(step t217 (cl (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0)))) (not (= e4 (op e4 e0))) (not (= e0 (op e4 (op e4 e0))))) :rule and_neg)
(step t218 (cl (=> (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0)))) (= e4 (op e4 (op e4 e4)))) (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0))))) :rule implies_neg1)
(anchor :step t219)
(assume t219.a0 (= e4 (op e4 e0)))
(assume t219.a1 (= e0 (op e4 (op e4 e0))))
(step t219.t0 (cl (not (= (= (op e4 unit) e4) (= e4 (op e4 e0)))) (not (= (op e4 unit) e4)) (= e4 (op e4 e0))) :rule equiv_pos2)
(step t219.t1 (cl (= (= (op e4 unit) e4) (= e4 (op e4 unit)))) :rule rare_rewrite :args ("eq-symm" (op e4 unit) e4))
(step t219.t2 (cl (= e4 e4)) :rule refl)
(step t219.t3 (cl (= (op e4 unit) (op e4 e0))) :rule cong :premises (t219.t2 a8))
(step t219.t4 (cl (= (= e4 (op e4 unit)) (= e4 (op e4 e0)))) :rule cong :premises (t219.t2 t219.t3))
(step t219.t5 (cl (= (= (op e4 unit) e4) (= e4 (op e4 e0)))) :rule trans :premises (t219.t1 t219.t4))
(step t219.t6 (cl (and (= (op unit e4) e4) (= (op e4 unit) e4))) :rule and :premises (t147) :args (1))
(step t219.t7 (cl (= (op e4 unit) e4)) :rule and :premises (t219.t6) :args (1))
(step t219.t8 (cl (= e4 (op e4 e0))) :rule resolution :premises (t219.t0 t219.t5 t219.t7))
(step t219.t9 (cl (= (op e4 (op e4 e0)) e0)) :rule symm :premises (t219.a1))
(step t219.t10 (cl (= e0 (op e4 (op e4 e0)))) :rule symm :premises (t219.t9))
(step t219.t11 (cl (= (op e4 e0) e4)) :rule symm :premises (t219.t8))
(step t219.t12 (cl (= (op e4 (op e4 e0)) (op e4 e4))) :rule cong :premises (t219.t2 t219.t11))
(step t219.t13 (cl (= e0 (op e4 e4))) :rule trans :premises (t219.t10 t219.t12))
(step t219.t14 (cl (= (op e4 e0) (op e4 (op e4 e4)))) :rule cong :premises (t219.t2 t219.t13))
(step t219.t15 (cl (= e4 (op e4 (op e4 e4)))) :rule trans :premises (t219.t8 t219.t14))
(step t219 (cl (not (= e4 (op e4 e0))) (not (= e0 (op e4 (op e4 e0)))) (= e4 (op e4 (op e4 e4)))) :rule subproof :discharge (t219.a0 t219.a1))
(step t220 (cl (not (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0))))) (= e4 (op e4 e0))) :rule and_pos :args (0))
(step t221 (cl (not (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0))))) (= e0 (op e4 (op e4 e0)))) :rule and_pos :args (1))
(step t222 (cl (= e4 (op e4 (op e4 e4))) (not (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0))))) (not (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0)))))) :rule resolution :premises (t219 t220 t221))
(step t223 (cl (not (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0))))) (not (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0))))) (= e4 (op e4 (op e4 e4)))) :rule reordering :premises (t222))
(step t224 (cl (not (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0))))) (= e4 (op e4 (op e4 e4)))) :rule contraction :premises (t223))
(step t225 (cl (=> (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0)))) (= e4 (op e4 (op e4 e4)))) (= e4 (op e4 (op e4 e4)))) :rule resolution :premises (t218 t224))
(step t226 (cl (=> (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0)))) (= e4 (op e4 (op e4 e4)))) (not (= e4 (op e4 (op e4 e4))))) :rule implies_neg2)
(step t227 (cl (=> (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0)))) (= e4 (op e4 (op e4 e4)))) (=> (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0)))) (= e4 (op e4 (op e4 e4))))) :rule resolution :premises (t225 t226))
(step t228 (cl (=> (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0)))) (= e4 (op e4 (op e4 e4))))) :rule contraction :premises (t227))
(step t229 (cl (not (and (= e4 (op e4 e0)) (= e0 (op e4 (op e4 e0))))) (= e4 (op e4 (op e4 e4)))) :rule implies :premises (t228))
(step t230 (cl (not (= e4 (op e4 e0))) (not (= e0 (op e4 (op e4 e0)))) (= e4 (op e4 (op e4 e4)))) :rule resolution :premises (t217 t229))
(step t231 (cl (= e4 (op e4 (op e4 e4))) (not (= e4 (op e4 e0))) (not (= e0 (op e4 (op e4 e0))))) :rule reordering :premises (t230))
(step t232 (cl (not (= (or (not (= (op e4 (op e4 e4)) e4)) (= e4 unit)) (or (not (= e4 (op e4 (op e4 e4)))) (= e4 e0)))) (not (or (not (= (op e4 (op e4 e4)) e4)) (= e4 unit))) (or (not (= e4 (op e4 (op e4 e4)))) (= e4 e0))) :rule equiv_pos2)
(step t233 (cl (= (= (op e4 (op e4 e4)) e4) (= e4 (op e4 (op e4 e4))))) :rule rare_rewrite :args ("eq-symm" (op e4 (op e4 e4)) e4))
(step t234 (cl (= (not (= (op e4 (op e4 e4)) e4)) (not (= e4 (op e4 (op e4 e4)))))) :rule cong :premises (t233))
(step t235 (cl (= (= e4 unit) (= unit e4))) :rule rare_rewrite :args ("eq-symm" e4 unit))
(step t236 (cl (= (or (not (= (op e4 (op e4 e4)) e4)) (= e4 unit)) (or (not (= e4 (op e4 (op e4 e4)))) (= unit e4)))) :rule cong :premises (t234 t235))
(step t237 (cl (= (not (= e4 (op e4 (op e4 e4)))) (not (= e4 (op e4 (op e4 e4)))))) :rule refl)
(step t238 (cl (= e4 e4)) :rule refl)
(step t239 (cl (= (= unit e4) (= e0 e4))) :rule cong :premises (a8 t238))
(step t240 (cl (= (or (not (= e4 (op e4 (op e4 e4)))) (= unit e4)) (or (not (= e4 (op e4 (op e4 e4)))) (= e0 e4)))) :rule cong :premises (t237 t239))
(step t241 (cl (= (= e0 e4) (= e4 e0))) :rule rare_rewrite :args ("eq-symm" e0 e4))
(step t242 (cl (= (or (not (= e4 (op e4 (op e4 e4)))) (= e0 e4)) (or (not (= e4 (op e4 (op e4 e4)))) (= e4 e0)))) :rule cong :premises (t237 t241))
(step t243 (cl (= (or (not (= e4 (op e4 (op e4 e4)))) (= unit e4)) (or (not (= e4 (op e4 (op e4 e4)))) (= e4 e0)))) :rule trans :premises (t240 t242))
(step t244 (cl (= (or (not (= (op e4 (op e4 e4)) e4)) (= e4 unit)) (or (not (= e4 (op e4 (op e4 e4)))) (= e4 e0)))) :rule trans :premises (t236 t243))
(step t245 (cl (or (not (= (op e4 (op e4 e4)) e4)) (= e4 unit))) :rule and :premises (t117) :args (1))
(step t246 (cl (or (not (= e4 (op e4 (op e4 e4)))) (= e4 e0))) :rule resolution :premises (t232 t244 t245))
(step t247 (cl (not (= e4 (op e4 (op e4 e4)))) (= e4 e0)) :rule or :premises (t246))
(step t248 (cl (= e4 e0) (not (= e4 (op e4 (op e4 e4))))) :rule reordering :premises (t247))
(step t249 (cl (not (= e0 e4))) :rule and :premises (t135) :args (1))
(step t250 (cl (not (= e4 e0))) :rule not_symm :premises (t249))
(step t251 (cl (not (= e4 (op e4 (op e4 e4))))) :rule resolution :premises (t248 t250))
(step t252 (cl (not (= (= (op e4 unit) e4) (= e4 (op e4 e0)))) (not (= (op e4 unit) e4)) (= e4 (op e4 e0))) :rule equiv_pos2)
(step t253 (cl (= (= (op e4 unit) e4) (= e4 (op e4 unit)))) :rule rare_rewrite :args ("eq-symm" (op e4 unit) e4))
(step t254 (cl (= (op e4 unit) (op e4 e0))) :rule cong :premises (t238 a8))
(step t255 (cl (= (= e4 (op e4 unit)) (= e4 (op e4 e0)))) :rule cong :premises (t238 t254))
(step t256 (cl (= (= (op e4 unit) e4) (= e4 (op e4 e0)))) :rule trans :premises (t253 t255))
(step t257 (cl (and (= (op unit e4) e4) (= (op e4 unit) e4))) :rule and :premises (t147) :args (1))
(step t258 (cl (= (op e4 unit) e4)) :rule and :premises (t257) :args (1))
(step t259 (cl (= e4 (op e4 e0))) :rule resolution :premises (t252 t256 t258))
(step t260 (cl (not (= e0 (op e4 (op e4 e0))))) :rule resolution :premises (t231 t251 t259))
(step t261 (cl (not (and (= e0 (op e4 (op e4 e0))) (not (= e0 (op e4 e0)))))) :rule resolution :premises (t216 t260))
(step t262 (cl (not (and (= e0 (op e5 (op e5 e0))) (not (= e0 (op e5 e0))))) (= e0 (op e5 (op e5 e0)))) :rule and_pos :args (0))
(step t263 (cl (= e0 (op e5 (op e5 e0))) (not (and (= e0 (op e5 (op e5 e0))) (not (= e0 (op e5 e0)))))) :rule reordering :premises (t262))
(step t264 (cl (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0)))) (not (= e5 (op e5 e0))) (not (= e0 (op e5 (op e5 e0))))) :rule and_neg)
(step t265 (cl (=> (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0)))) (= e5 (op e5 (op e5 e5)))) (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0))))) :rule implies_neg1)
(anchor :step t266)
(assume t266.a0 (= e5 (op e5 e0)))
(assume t266.a1 (= e0 (op e5 (op e5 e0))))
(step t266.t0 (cl (not (= (= (op e5 unit) e5) (= e5 (op e5 e0)))) (not (= (op e5 unit) e5)) (= e5 (op e5 e0))) :rule equiv_pos2)
(step t266.t1 (cl (= (= (op e5 unit) e5) (= e5 (op e5 unit)))) :rule rare_rewrite :args ("eq-symm" (op e5 unit) e5))
(step t266.t2 (cl (= e5 e5)) :rule refl)
(step t266.t3 (cl (= (op e5 unit) (op e5 e0))) :rule cong :premises (t266.t2 a8))
(step t266.t4 (cl (= (= e5 (op e5 unit)) (= e5 (op e5 e0)))) :rule cong :premises (t266.t2 t266.t3))
(step t266.t5 (cl (= (= (op e5 unit) e5) (= e5 (op e5 e0)))) :rule trans :premises (t266.t1 t266.t4))
(step t266.t6 (cl (and (= (op unit e5) e5) (= (op e5 unit) e5))) :rule and :premises (t146) :args (1))
(step t266.t7 (cl (= (op e5 unit) e5)) :rule and :premises (t266.t6) :args (1))
(step t266.t8 (cl (= e5 (op e5 e0))) :rule resolution :premises (t266.t0 t266.t5 t266.t7))
(step t266.t9 (cl (= (op e5 (op e5 e0)) e0)) :rule symm :premises (t266.a1))
(step t266.t10 (cl (= e0 (op e5 (op e5 e0)))) :rule symm :premises (t266.t9))
(step t266.t11 (cl (= (op e5 e0) e5)) :rule symm :premises (t266.t8))
(step t266.t12 (cl (= (op e5 (op e5 e0)) (op e5 e5))) :rule cong :premises (t266.t2 t266.t11))
(step t266.t13 (cl (= e0 (op e5 e5))) :rule trans :premises (t266.t10 t266.t12))
(step t266.t14 (cl (= (op e5 e0) (op e5 (op e5 e5)))) :rule cong :premises (t266.t2 t266.t13))
(step t266.t15 (cl (= e5 (op e5 (op e5 e5)))) :rule trans :premises (t266.t8 t266.t14))
(step t266 (cl (not (= e5 (op e5 e0))) (not (= e0 (op e5 (op e5 e0)))) (= e5 (op e5 (op e5 e5)))) :rule subproof :discharge (t266.a0 t266.a1))
(step t267 (cl (not (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0))))) (= e5 (op e5 e0))) :rule and_pos :args (0))
(step t268 (cl (not (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0))))) (= e0 (op e5 (op e5 e0)))) :rule and_pos :args (1))
(step t269 (cl (= e5 (op e5 (op e5 e5))) (not (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0))))) (not (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0)))))) :rule resolution :premises (t266 t267 t268))
(step t270 (cl (not (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0))))) (not (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0))))) (= e5 (op e5 (op e5 e5)))) :rule reordering :premises (t269))
(step t271 (cl (not (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0))))) (= e5 (op e5 (op e5 e5)))) :rule contraction :premises (t270))
(step t272 (cl (=> (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0)))) (= e5 (op e5 (op e5 e5)))) (= e5 (op e5 (op e5 e5)))) :rule resolution :premises (t265 t271))
(step t273 (cl (=> (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0)))) (= e5 (op e5 (op e5 e5)))) (not (= e5 (op e5 (op e5 e5))))) :rule implies_neg2)
(step t274 (cl (=> (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0)))) (= e5 (op e5 (op e5 e5)))) (=> (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0)))) (= e5 (op e5 (op e5 e5))))) :rule resolution :premises (t272 t273))
(step t275 (cl (=> (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0)))) (= e5 (op e5 (op e5 e5))))) :rule contraction :premises (t274))
(step t276 (cl (not (and (= e5 (op e5 e0)) (= e0 (op e5 (op e5 e0))))) (= e5 (op e5 (op e5 e5)))) :rule implies :premises (t275))
(step t277 (cl (not (= e5 (op e5 e0))) (not (= e0 (op e5 (op e5 e0)))) (= e5 (op e5 (op e5 e5)))) :rule resolution :premises (t264 t276))
(step t278 (cl (= e5 (op e5 (op e5 e5))) (not (= e5 (op e5 e0))) (not (= e0 (op e5 (op e5 e0))))) :rule reordering :premises (t277))
(step t279 (cl (not (= (or (not (= (op e5 (op e5 e5)) e5)) (= e5 unit)) (or (not (= e5 (op e5 (op e5 e5)))) (= e5 e0)))) (not (or (not (= (op e5 (op e5 e5)) e5)) (= e5 unit))) (or (not (= e5 (op e5 (op e5 e5)))) (= e5 e0))) :rule equiv_pos2)
(step t280 (cl (= (= (op e5 (op e5 e5)) e5) (= e5 (op e5 (op e5 e5))))) :rule rare_rewrite :args ("eq-symm" (op e5 (op e5 e5)) e5))
(step t281 (cl (= (not (= (op e5 (op e5 e5)) e5)) (not (= e5 (op e5 (op e5 e5)))))) :rule cong :premises (t280))
(step t282 (cl (= (= e5 unit) (= unit e5))) :rule rare_rewrite :args ("eq-symm" e5 unit))
(step t283 (cl (= (or (not (= (op e5 (op e5 e5)) e5)) (= e5 unit)) (or (not (= e5 (op e5 (op e5 e5)))) (= unit e5)))) :rule cong :premises (t281 t282))
(step t284 (cl (= (not (= e5 (op e5 (op e5 e5)))) (not (= e5 (op e5 (op e5 e5)))))) :rule refl)
(step t285 (cl (= e5 e5)) :rule refl)
(step t286 (cl (= (= unit e5) (= e0 e5))) :rule cong :premises (a8 t285))
(step t287 (cl (= (or (not (= e5 (op e5 (op e5 e5)))) (= unit e5)) (or (not (= e5 (op e5 (op e5 e5)))) (= e0 e5)))) :rule cong :premises (t284 t286))
(step t288 (cl (= (= e0 e5) (= e5 e0))) :rule rare_rewrite :args ("eq-symm" e0 e5))
(step t289 (cl (= (or (not (= e5 (op e5 (op e5 e5)))) (= e0 e5)) (or (not (= e5 (op e5 (op e5 e5)))) (= e5 e0)))) :rule cong :premises (t284 t288))
(step t290 (cl (= (or (not (= e5 (op e5 (op e5 e5)))) (= unit e5)) (or (not (= e5 (op e5 (op e5 e5)))) (= e5 e0)))) :rule trans :premises (t287 t289))
(step t291 (cl (= (or (not (= (op e5 (op e5 e5)) e5)) (= e5 unit)) (or (not (= e5 (op e5 (op e5 e5)))) (= e5 e0)))) :rule trans :premises (t283 t290))
(step t292 (cl (or (not (= (op e5 (op e5 e5)) e5)) (= e5 unit))) :rule and :premises (t116) :args (1))
(step t293 (cl (or (not (= e5 (op e5 (op e5 e5)))) (= e5 e0))) :rule resolution :premises (t279 t291 t292))
(step t294 (cl (not (= e5 (op e5 (op e5 e5)))) (= e5 e0)) :rule or :premises (t293))
(step t295 (cl (= e5 e0) (not (= e5 (op e5 (op e5 e5))))) :rule reordering :premises (t294))
(step t296 (cl (not (= e0 e5))) :rule and :premises (t134) :args (1))
(step t297 (cl (not (= e5 e0))) :rule not_symm :premises (t296))
(step t298 (cl (not (= e5 (op e5 (op e5 e5))))) :rule resolution :premises (t295 t297))
(step t299 (cl (not (= (= (op e5 unit) e5) (= e5 (op e5 e0)))) (not (= (op e5 unit) e5)) (= e5 (op e5 e0))) :rule equiv_pos2)
(step t300 (cl (= (= (op e5 unit) e5) (= e5 (op e5 unit)))) :rule rare_rewrite :args ("eq-symm" (op e5 unit) e5))
(step t301 (cl (= (op e5 unit) (op e5 e0))) :rule cong :premises (t285 a8))
(step t302 (cl (= (= e5 (op e5 unit)) (= e5 (op e5 e0)))) :rule cong :premises (t285 t301))
(step t303 (cl (= (= (op e5 unit) e5) (= e5 (op e5 e0)))) :rule trans :premises (t300 t302))
(step t304 (cl (and (= (op unit e5) e5) (= (op e5 unit) e5))) :rule and :premises (t146) :args (1))
(step t305 (cl (= (op e5 unit) e5)) :rule and :premises (t304) :args (1))
(step t306 (cl (= e5 (op e5 e0))) :rule resolution :premises (t299 t303 t305))
(step t307 (cl (not (= e0 (op e5 (op e5 e0))))) :rule resolution :premises (t278 t298 t306))
(step t308 (cl (not (and (= e0 (op e5 (op e5 e0))) (not (= e0 (op e5 e0)))))) :rule resolution :premises (t263 t307))
(step t309 (cl (and (= e0 (op e2 (op e2 e0))) (not (= e0 (op e2 e0))))) :rule resolution :premises (t85 t155 t167 t214 t261 t308))
(step t310 (cl (= e0 (op e2 (op e2 e0)))) :rule resolution :premises (t16 t309))
(step t311 (cl (not (= (or (not (= (op e2 (op e2 e2)) e2)) (= e2 unit)) (or (not (= e2 (op e2 (op e2 e2)))) (= e2 e0)))) (not (or (not (= (op e2 (op e2 e2)) e2)) (= e2 unit))) (or (not (= e2 (op e2 (op e2 e2)))) (= e2 e0))) :rule equiv_pos2)
(step t312 (cl (= (= (op e2 (op e2 e2)) e2) (= e2 (op e2 (op e2 e2))))) :rule rare_rewrite :args ("eq-symm" (op e2 (op e2 e2)) e2))
(step t313 (cl (= (not (= (op e2 (op e2 e2)) e2)) (not (= e2 (op e2 (op e2 e2)))))) :rule cong :premises (t312))
(step t314 (cl (= (= e2 unit) (= unit e2))) :rule rare_rewrite :args ("eq-symm" e2 unit))
(step t315 (cl (= (or (not (= (op e2 (op e2 e2)) e2)) (= e2 unit)) (or (not (= e2 (op e2 (op e2 e2)))) (= unit e2)))) :rule cong :premises (t313 t314))
(step t316 (cl (= (not (= e2 (op e2 (op e2 e2)))) (not (= e2 (op e2 (op e2 e2)))))) :rule refl)
(step t317 (cl (= e2 e2)) :rule refl)
(step t318 (cl (= (= unit e2) (= e0 e2))) :rule cong :premises (a8 t317))
(step t319 (cl (= (or (not (= e2 (op e2 (op e2 e2)))) (= unit e2)) (or (not (= e2 (op e2 (op e2 e2)))) (= e0 e2)))) :rule cong :premises (t316 t318))
(step t320 (cl (= (= e0 e2) (= e2 e0))) :rule rare_rewrite :args ("eq-symm" e0 e2))
(step t321 (cl (= (or (not (= e2 (op e2 (op e2 e2)))) (= e0 e2)) (or (not (= e2 (op e2 (op e2 e2)))) (= e2 e0)))) :rule cong :premises (t316 t320))
(step t322 (cl (= (or (not (= e2 (op e2 (op e2 e2)))) (= unit e2)) (or (not (= e2 (op e2 (op e2 e2)))) (= e2 e0)))) :rule trans :premises (t319 t321))
(step t323 (cl (= (or (not (= (op e2 (op e2 e2)) e2)) (= e2 unit)) (or (not (= e2 (op e2 (op e2 e2)))) (= e2 e0)))) :rule trans :premises (t315 t322))
(step t324 (cl (or (not (= (op e2 (op e2 e2)) e2)) (= e2 unit))) :rule and :premises (t119) :args (1))
(step t325 (cl (or (not (= e2 (op e2 (op e2 e2)))) (= e2 e0))) :rule resolution :premises (t311 t323 t324))
(step t326 (cl (not (= e2 (op e2 (op e2 e2)))) (= e2 e0)) :rule or :premises (t325))
(step t327 (cl (= e2 e0) (not (= e2 (op e2 (op e2 e2))))) :rule reordering :premises (t326))
(step t328 (cl (not (= e0 e2))) :rule and :premises (t137) :args (1))
(step t329 (cl (not (= e2 e0))) :rule not_symm :premises (t328))
(step t330 (cl (not (= e2 (op e2 (op e2 e2))))) :rule resolution :premises (t327 t329))
(step t331 (cl (not (= (= (op e2 unit) e2) (= e2 (op e2 e0)))) (not (= (op e2 unit) e2)) (= e2 (op e2 e0))) :rule equiv_pos2)
(step t332 (cl (= (= (op e2 unit) e2) (= e2 (op e2 unit)))) :rule rare_rewrite :args ("eq-symm" (op e2 unit) e2))
(step t333 (cl (= (op e2 unit) (op e2 e0))) :rule cong :premises (t317 a8))
(step t334 (cl (= (= e2 (op e2 unit)) (= e2 (op e2 e0)))) :rule cong :premises (t317 t333))
(step t335 (cl (= (= (op e2 unit) e2) (= e2 (op e2 e0)))) :rule trans :premises (t332 t334))
(step t336 (cl (and (= (op unit e2) e2) (= (op e2 unit) e2))) :rule and :premises (t149) :args (1))
(step t337 (cl (= (op e2 unit) e2)) :rule and :premises (t336) :args (1))
(step t338 (cl (= e2 (op e2 e0))) :rule resolution :premises (t331 t335 t337))
(step t339 (cl) :rule resolution :premises (t14 t310 t330 t338))
