unsat
(assume a0 (= y$a_q0_Bottle_1 (not y$1)))
(assume a1 (= y$a_q0_Bottle_2 (not y$3)))
(assume a2 (= y$a_q0_Bottle_3 (not y$5)))
(assume a3 (= y$a_q0_Bottle_4 (not y$7)))
(assume a4 (= y$a_q0_Bottle_5 (not y$9)))
(assume a5 (= y$a_q10_Bottle_1 (not y$11)))
(assume a6 (= y$a_q10_Bottle_3 (not y$13)))
(assume a7 (= y$a_q11 (not y$15)))
(assume a8 (= y$a_q12 (not y$17)))
(assume a9 (= y$a_q13 (not y$19)))
(assume a10 (= y$a_q14 (not y$21)))
(assume a11 (= y$a_q1_Bottle_1 (not y$23)))
(assume a12 (= y$a_q1_Bottle_2 (not y$25)))
(assume a13 (= y$a_q1_Bottle_3 (not y$27)))
(assume a14 (= y$a_q1_Bottle_4 (not y$29)))
(assume a15 (= y$a_q1_Bottle_5 (not y$31)))
(assume a16 (= y$a_q2_Bottle_1 (not y$33)))
(assume a17 (= y$a_q2_Bottle_2 (not y$35)))
(assume a18 (= y$a_q2_Bottle_3 (not y$37)))
(assume a19 (= y$a_q2_Bottle_4 (not y$39)))
(assume a20 (= y$a_q2_Bottle_5 (not y$41)))
(assume a21 (= y$a_q3_Bottle_1 (not y$43)))
(assume a22 (= y$a_q3_Bottle_2 (not y$45)))
(assume a23 (= y$a_q3_Bottle_3 (not y$47)))
(assume a24 (= y$a_q3_Bottle_4 (not y$49)))
(assume a25 (= y$a_q3_Bottle_5 (not y$51)))
(assume a26 (= y$a_q4_Bottle_1 (not y$53)))
(assume a27 (= y$a_q4_Bottle_2 (not y$55)))
(assume a28 (= y$a_q4_Bottle_3 (not y$57)))
(assume a29 (= y$a_q4_Bottle_4 (not y$59)))
(assume a30 (= y$a_q4_Bottle_5 (not y$61)))
(assume a31 (= y$a_q5_Bottle_1 (not y$63)))
(assume a32 (= y$a_q5_Bottle_2 (not y$65)))
(assume a33 (= y$a_q5_Bottle_3 (not y$67)))
(assume a34 (= y$a_q5_Bottle_4 (not y$69)))
(assume a35 (= y$a_q5_Bottle_5 (not y$71)))
(assume a36 (= y$a_q6_Bottle_1 (not y$73)))
(assume a37 (= y$a_q6_Bottle_2 (not y$75)))
(assume a38 (= y$a_q6_Bottle_3 (not y$77)))
(assume a39 (= y$a_q6_Bottle_4 (not y$79)))
(assume a40 (= y$a_q7_Bottle_1 (not y$81)))
(assume a41 (= y$a_q7_Bottle_2 (not y$83)))
(assume a42 (= y$a_q7_Bottle_3 (not y$85)))
(assume a43 (= y$a_q7_Bottle_4 (not y$87)))
(assume a44 (= y$a_q8_Bottle_1 (not y$89)))
(assume a45 (= y$a_q8_Bottle_2 (not y$91)))
(assume a46 (= y$a_q8_Bottle_3 (not y$93)))
(assume a47 (= y$a_q9_Bottle_1 (not y$95)))
(assume a48 (= y$a_q9_Bottle_3 (not y$97)))
(assume a49 (= y$a_q_Sink (not y$99)))
(assume a50 (= y$a_q_Source (not y$101)))
(assume a51 (= y$dve_invalid (not y$103)))
(assume a52 (= y$prop (not y$57110)))
(assume a53 (= y$id54_op (and y$a_q4_Bottle_1 y$103)))
(assume a54 (= y$id54_op (not y$55378)))
(assume a55 (= y$55379 (= y$prop y$55378)))
(assume a56 (= y$57117 (and y$1 y$3 y$5 y$7 y$9 y$11 y$13 y$15 y$17 y$19 y$21 y$23 y$25 y$27 y$29 y$31 y$33 y$35 y$37 y$39 y$41 y$43 y$45 y$47 y$49 y$51 y$53 y$55 y$57 y$59 y$61 y$63 y$65 y$67 y$69 y$71 y$73 y$75 y$77 y$79 y$81 y$83 y$85 y$87 y$89 y$91 y$93 y$95 y$97 y$99 y$101 y$103 y$57110 y$55379)))
(assume a57 y$57117)
(step t0 (cl (not (= y$id54_op (not y$55378))) y$id54_op (not (not y$55378))) :rule equiv_pos1)
(step t1 (cl (not y$id54_op) (and y$a_q4_Bottle_1 y$103)) :rule equiv1 :premises (a53))
(step t2 (cl (not (and y$a_q4_Bottle_1 y$103)) y$a_q4_Bottle_1) :rule and_pos :args (0))
(step t3 (cl (not y$a_q4_Bottle_1) (not y$53)) :rule equiv1 :premises (a26))
(step t4 (cl (not (= y$57117 (and y$1 y$3 y$5 y$7 y$9 y$11 y$13 y$15 y$17 y$19 y$21 y$23 y$25 y$27 y$29 y$31 y$33 y$35 y$37 y$39 y$41 y$43 y$45 y$47 y$49 y$51 y$53 y$55 y$57 y$59 y$61 y$63 y$65 y$67 y$69 y$71 y$73 y$75 y$77 y$79 y$81 y$83 y$85 y$87 y$89 y$91 y$93 y$95 y$97 y$99 y$101 y$103 y$57110 y$55379))) (not y$57117) (and y$1 y$3 y$5 y$7 y$9 y$11 y$13 y$15 y$17 y$19 y$21 y$23 y$25 y$27 y$29 y$31 y$33 y$35 y$37 y$39 y$41 y$43 y$45 y$47 y$49 y$51 y$53 y$55 y$57 y$59 y$61 y$63 y$65 y$67 y$69 y$71 y$73 y$75 y$77 y$79 y$81 y$83 y$85 y$87 y$89 y$91 y$93 y$95 y$97 y$99 y$101 y$103 y$57110 y$55379)) :rule equiv_pos2)
(step t5 (cl (and y$1 y$3 y$5 y$7 y$9 y$11 y$13 y$15 y$17 y$19 y$21 y$23 y$25 y$27 y$29 y$31 y$33 y$35 y$37 y$39 y$41 y$43 y$45 y$47 y$49 y$51 y$53 y$55 y$57 y$59 y$61 y$63 y$65 y$67 y$69 y$71 y$73 y$75 y$77 y$79 y$81 y$83 y$85 y$87 y$89 y$91 y$93 y$95 y$97 y$99 y$101 y$103 y$57110 y$55379)) :rule resolution :premises (t4 a56 a57))
(step t6 (cl y$53) :rule and :premises (t5) :args (26))
(step t7 (cl (not y$a_q4_Bottle_1)) :rule resolution :premises (t3 t6))
(step t8 (cl (not (and y$a_q4_Bottle_1 y$103))) :rule resolution :premises (t2 t7))
(step t9 (cl (not y$id54_op)) :rule resolution :premises (t1 t8))
(step t10 (cl (not (= y$55379 (= y$55378 y$prop))) (not y$55379) (= y$55378 y$prop)) :rule equiv_pos2)
(step t11 (cl (not (= (= y$55379 (= y$prop y$55378)) (= y$55379 (= y$55378 y$prop)))) (not (= y$55379 (= y$prop y$55378))) (= y$55379 (= y$55378 y$prop))) :rule equiv_pos2)
(step t12 (cl (= y$55379 y$55379)) :rule refl)
(step t13 (cl (= (= y$prop y$55378) (= y$55378 y$prop))) :rule rare_rewrite :args ("eq-symm" y$prop y$55378))
(step t14 (cl (= (= y$55379 (= y$prop y$55378)) (= y$55379 (= y$55378 y$prop)))) :rule cong :premises (t12 t13))
(step t15 (cl (= y$55379 (= y$55378 y$prop))) :rule resolution :premises (t11 t14 a55))
(step t16 (cl y$55379) :rule and :premises (t5) :args (53))
(step t17 (cl (= y$55378 y$prop)) :rule resolution :premises (t10 t15 t16))
(step t18 (cl (not y$55378) y$prop) :rule equiv1 :premises (t17))
(step t19 (cl (not y$prop) (not y$57110)) :rule equiv1 :premises (a52))
(step t20 (cl y$57110) :rule and :premises (t5) :args (52))
(step t21 (cl (not y$prop)) :rule resolution :premises (t19 t20))
(step t22 (cl (not y$55378)) :rule resolution :premises (t18 t21))
(step t23 (cl (not (= y$id54_op (not y$55378)))) :rule resolution :premises (t0 t9 t22))
(step t24 (cl) :rule resolution :premises (a54 t23))
