unsat
(assume a0 (not (= f1 f2)))
(assume a1 (! (not false) :named @p_1))
(assume a2 (! (= (! (f3 f4 (! (f5 (! (f6 f7 (! (f5 (! (f6 f8 (! (f5 f9 (! (f5 f10 (! (f5 f10 (! (f5 f11 f12) :named @p_2)) :named @p_3)) :named @p_4)) :named @p_5)) :named @p_6) f13) :named @p_7)) :named @p_8) f14) :named @p_9)) :named @p_10) f1) :named @p_11))
(assume a3 (! (= (! (f3 (! (f15 f16 @p_9) :named @p_12) f14) :named @p_13) f1) :named @p_14))
(assume a4 (not (= (f3 (! (f15 f16 f14) :named @p_15) f13) f1)))
(assume a5 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 f17 ?v0) f1) (=> (= (f3 f17 ?v1) f1) (= (f3 f17 (f5 (f6 f8 ?v0) ?v1)) f1)))))
(assume a6 (forall ((?v0 S3)) (=> (= (f3 f4 ?v0) f1) (=> (not (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12))))) (=> (not (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12))))) (= (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0) f1))))))
(assume a7 (forall ((?v0 S3)) (= (f5 (f6 f7 f14) (f5 f9 ?v0)) (f5 f9 (f5 (f6 f7 (f5 f11 f12)) ?v0)))))
(assume a8 (forall ((?v0 S3)) (= (f5 (f6 f7 (f5 f9 ?v0)) f14) (f5 f9 (f5 (f6 f7 ?v0) (f5 f11 f12))))))
(assume a9 (forall ((?v0 S3)) (= (= (f3 (f15 f16 f14) (f5 f9 ?v0)) f1) (= (f3 (f15 f16 (f5 f11 f12)) ?v0) f1))))
(assume a10 (forall ((?v0 S3)) (= (= (f3 (f15 f16 (f5 f9 ?v0)) f14) f1) (= (f3 (f15 f16 ?v0) (f5 f11 f12)) f1))))
(assume a11 (= (f5 (f6 f7 f14) f14) (! (f5 f9 @p_3) :named @p_16)))
(assume a12 (= (f18 (f19 f20 f21) f21) (f22 f23 @p_3)))
(assume a13 (forall ((?v0 S3)) (= (f5 (f6 f8 ?v0) (f5 f9 (f5 f10 (f5 f11 f12)))) (f5 (f6 f7 ?v0) ?v0))))
(assume a14 (forall ((?v0 S8)) (= (f18 (f19 f24 ?v0) (f22 f23 (f5 f10 (f5 f11 f12)))) (f18 (f19 f20 ?v0) ?v0))))
(assume a15 (forall ((?v0 S3)) (= (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f11 f12)))) ?v0) (f5 (f6 f7 ?v0) ?v0))))
(assume a16 (forall ((?v0 S8)) (= (f18 (f19 f24 (f22 f23 (f5 f10 (f5 f11 f12)))) ?v0) (f18 (f19 f20 ?v0) ?v0))))
(assume a17 (= (f3 f4 @p_16) f1))
(assume a18 (forall ((?v0 S3) (?v1 S3)) (= (= (f5 f9 ?v0) (f5 f9 ?v1)) (= ?v0 ?v1))))
(assume a19 (forall ((?v0 S3) (?v1 S8)) (= (= (f22 f23 ?v0) ?v1) (= ?v1 (f22 f23 ?v0)))))
(assume a20 (forall ((?v0 S3) (?v1 S3)) (= (= (f5 f9 ?v0) ?v1) (= ?v1 (f5 f9 ?v0)))))
(assume a21 (forall ((?v0 S3) (?v1 S3)) (= (= (f5 f11 ?v0) (f5 f11 ?v1)) (= ?v0 ?v1))))
(assume a22 (= (= f12 f12) true))
(assume a23 (forall ((?v0 S3) (?v1 S3)) (= (= (f5 f10 ?v0) (f5 f10 ?v1)) (= ?v0 ?v1))))
(assume a24 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v0) f1) (= ?v0 ?v1)))))
(assume a25 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v2) f1) (= (f3 (f15 f16 ?v0) ?v2) f1)))))
(assume a26 (forall ((?v0 S3) (?v1 S3)) (or (= (f3 (f15 f16 ?v0) ?v1) f1) (= (f3 (f15 f16 ?v1) ?v0) f1))))
(assume a27 (forall ((?v0 S3)) (= (f3 (f15 f16 ?v0) ?v0) f1)))
(assume a28 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (f5 (f6 f8 (f5 (f6 f8 ?v0) ?v1)) ?v2) (f5 (f6 f8 ?v0) (f5 (f6 f8 ?v1) ?v2)))))
(assume a29 (forall ((?v0 S3) (?v1 S3)) (= (f5 (f6 f8 ?v0) ?v1) (f5 (f6 f8 ?v1) ?v0))))
(assume a30 (forall ((?v0 S3)) (= (f5 f9 ?v0) ?v0)))
(assume a31 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (f5 (f6 f7 (f5 (f6 f7 ?v0) ?v1)) ?v2) (f5 (f6 f7 ?v0) (f5 (f6 f7 ?v1) ?v2)))))
(assume a32 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (f5 (f6 f7 ?v0) (f5 (f6 f7 ?v1) ?v2)) (f5 (f6 f7 ?v1) (f5 (f6 f7 ?v0) ?v2)))))
(assume a33 (forall ((?v0 S3) (?v1 S3)) (= (f5 (f6 f7 ?v0) ?v1) (f5 (f6 f7 ?v1) ?v0))))
(assume a34 (forall ((?v0 S3)) (= (= (f5 f11 ?v0) f12) false)))
(assume a35 (forall ((?v0 S3)) (= (= f12 (f5 f11 ?v0)) false)))
(assume a36 (forall ((?v0 S3) (?v1 S3)) (= (= (f5 f11 ?v0) (f5 f10 ?v1)) false)))
(assume a37 (forall ((?v0 S3) (?v1 S3)) (= (= (f5 f10 ?v0) (f5 f11 ?v1)) false)))
(assume a38 (forall ((?v0 S3)) (= (= (f5 f10 ?v0) f12) (= ?v0 f12))))
(assume a39 (forall ((?v0 S3)) (= (= f12 (f5 f10 ?v0)) (= f12 ?v0))))
(assume a40 (= (f5 f10 f12) f12))
(assume a41 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1))))
(assume a42 (= (= (f3 (f15 f16 f12) f12) f1) true))
(assume a43 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f10 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1))))
(assume a44 (forall ((?v0 S3)) (= (f5 (f6 f8 f12) ?v0) f12)))
(assume a45 (forall ((?v0 S3) (?v1 S3)) (= (f5 (f6 f8 (f5 f10 ?v0)) ?v1) (f5 f10 (f5 (f6 f8 ?v0) ?v1)))))
(assume a46 (forall ((?v0 S3)) (= (f5 (f6 f7 ?v0) f12) ?v0)))
(assume a47 (forall ((?v0 S3)) (= (f5 (f6 f7 f12) ?v0) ?v0)))
(assume a48 (forall ((?v0 S3) (?v1 S3)) (= (f5 (f6 f7 (f5 f10 ?v0)) (f5 f10 ?v1)) (f5 f10 (f5 (f6 f7 ?v0) ?v1)))))
(assume a49 (forall ((?v0 S3)) (= (f5 f10 ?v0) (f5 (f6 f7 ?v0) ?v0))))
(assume a50 (forall ((?v0 S3)) (= (f5 (f6 f8 ?v0) f14) ?v0)))
(assume a51 (forall ((?v0 S3)) (= (f5 (f6 f8 f14) ?v0) ?v0)))
(assume a52 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f9 ?v0)) (f5 f9 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1))))
(assume a53 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (= (f3 (f15 f16 (f5 (f6 f7 ?v2) ?v0)) (f5 (f6 f7 ?v2) ?v1)) f1))))
(assume a54 (forall ((?v0 S3) (?v1 S3)) (= (f5 (f6 f8 (f5 f9 ?v0)) (f5 f9 ?v1)) (f5 f9 (f5 (f6 f8 ?v0) ?v1)))))
(assume a55 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (f5 (f6 f8 (f5 (f6 f7 ?v0) ?v1)) ?v2) (f5 (f6 f7 (f5 (f6 f8 ?v0) ?v2)) (f5 (f6 f8 ?v1) ?v2)))))
(assume a56 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (f5 (f6 f8 ?v0) (f5 (f6 f7 ?v1) ?v2)) (f5 (f6 f7 (f5 (f6 f8 ?v0) ?v1)) (f5 (f6 f8 ?v0) ?v2)))))
(assume a57 (forall ((?v0 S3) (?v1 S3)) (= (f5 (f6 f7 (f5 f9 ?v0)) (f5 f9 ?v1)) (f5 f9 (f5 (f6 f7 ?v0) ?v1)))))
(assume a58 (forall ((?v0 S8) (?v1 S8) (?v2 S3)) (= (f18 (f19 f24 (f18 (f19 f20 ?v0) ?v1)) (f22 f23 ?v2)) (f18 (f19 f20 (f18 (f19 f24 ?v0) (f22 f23 ?v2))) (f18 (f19 f24 ?v1) (f22 f23 ?v2))))))
(assume a59 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (f5 (f6 f8 (f5 (f6 f7 ?v0) ?v1)) (f5 f9 ?v2)) (f5 (f6 f7 (f5 (f6 f8 ?v0) (f5 f9 ?v2))) (f5 (f6 f8 ?v1) (f5 f9 ?v2))))))
(assume a60 (forall ((?v0 S3) (?v1 S8) (?v2 S8)) (= (f18 (f19 f24 (f22 f23 ?v0)) (f18 (f19 f20 ?v1) ?v2)) (f18 (f19 f20 (f18 (f19 f24 (f22 f23 ?v0)) ?v1)) (f18 (f19 f24 (f22 f23 ?v0)) ?v2)))))
(assume a61 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (f5 (f6 f8 (f5 f9 ?v0)) (f5 (f6 f7 ?v1) ?v2)) (f5 (f6 f7 (f5 (f6 f8 (f5 f9 ?v0)) ?v1)) (f5 (f6 f8 (f5 f9 ?v0)) ?v2)))))
(assume a62 (forall ((?v0 S3)) (= (f5 (f6 f7 (f5 f9 f12)) ?v0) ?v0)))
(assume a63 (forall ((?v0 S3)) (= (f5 (f6 f7 ?v0) (f5 f9 f12)) ?v0)))
(assume a64 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (f5 (f6 f8 (f5 f9 ?v0)) (f5 (f6 f8 (f5 f9 ?v1)) ?v2)) (f5 (f6 f8 (f5 f9 (f5 (f6 f8 ?v0) ?v1))) ?v2))))
(assume a65 (forall ((?v0 S3) (?v1 S3)) (= (f5 f9 (f5 (f6 f8 ?v0) ?v1)) (f5 (f6 f8 (f5 f9 ?v0)) (f5 f9 ?v1)))))
(assume a66 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (f5 (f6 f7 (f5 f9 ?v0)) (f5 (f6 f7 (f5 f9 ?v1)) ?v2)) (f5 (f6 f7 (f5 f9 (f5 (f6 f7 ?v0) ?v1))) ?v2))))
(assume a67 (forall ((?v0 S3) (?v1 S3)) (= (f5 f9 (f5 (f6 f7 ?v0) ?v1)) (f5 (f6 f7 (f5 f9 ?v0)) (f5 f9 ?v1)))))
(assume a68 (forall ((?v0 S3)) (= (= (f3 (f15 f16 f12) (f5 f11 ?v0)) f1) (= (f3 (f15 f16 f12) ?v0) f1))))
(assume a69 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1))))
(assume a70 (forall ((?v0 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) f12) f1) (= (f3 (f15 f16 ?v0) f12) f1))))
(assume a71 (forall ((?v0 S3)) (= (= (f3 (f15 f16 f12) (f5 f10 ?v0)) f1) (= (f3 (f15 f16 f12) ?v0) f1))))
(assume a72 (forall ((?v0 S3) (?v1 S3)) (= (f5 (f6 f7 (f5 f11 ?v0)) (f5 f10 ?v1)) (f5 f11 (f5 (f6 f7 ?v0) ?v1)))))
(assume a73 (forall ((?v0 S3) (?v1 S3)) (= (f5 (f6 f7 (f5 f10 ?v0)) (f5 f11 ?v1)) (f5 f11 (f5 (f6 f7 ?v0) ?v1)))))
(assume a74 (forall ((?v0 S3)) (= (f5 f11 ?v0) (f5 (f6 f7 (f5 (f6 f7 f14) ?v0)) ?v0))))
(assume a75 (forall ((?v0 S3)) (= (f5 f9 (f5 f11 ?v0)) (f5 (f6 f7 (f5 (f6 f7 f14) (f5 f9 ?v0))) (f5 f9 ?v0)))))
(assume a76 (forall ((?v0 S3)) (= (f5 (f6 f8 (f5 f9 (f5 f11 f12))) ?v0) ?v0)))
(assume a77 (forall ((?v0 S3)) (= (f5 (f6 f8 ?v0) (f5 f9 (f5 f11 f12))) ?v0)))
(assume a78 (= (! (f22 f23 @p_2) :named @p_17) f21))
(assume a79 (= (! (f5 f9 @p_2) :named @p_18) f14))
(assume a80 (! (= f14 @p_18) :named @p_19))
(assume a81 (forall ((?v0 S3) (?v1 S3)) (= (f5 (f6 f8 (f5 f11 ?v0)) ?v1) (f5 (f6 f7 (f5 f10 (f5 (f6 f8 ?v0) ?v1))) ?v1))))
(assume a82 (forall ((?v0 S3)) (= (f5 (f6 f8 (f5 (f6 f7 f14) f14)) (f5 f9 ?v0)) (f5 f9 (f5 f10 ?v0)))))
(assume a83 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 f12) ?v0) f1) (=> (= (f3 (f15 f16 f12) ?v1) f1) (= (f18 (f19 f20 (f22 f23 ?v0)) (f22 f23 ?v1)) (f22 f23 (f5 (f6 f7 ?v0) ?v1)))))))
(assume a84 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 f12) ?v0) f1) (=> (= (f3 (f15 f16 f12) ?v1) f1) (= (f5 (f6 f7 (f5 f9 ?v0)) (f5 f9 ?v1)) (f5 f9 (f5 (f6 f7 ?v0) ?v1)))))))
(assume a85 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 f12) ?v0) f1) (=> (= (f3 (f15 f16 f12) ?v1) f1) (= (f18 (f19 f24 (f22 f23 ?v0)) (f22 f23 ?v1)) (f22 f23 (f5 (f6 f8 ?v0) ?v1)))))))
(assume a86 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 f12) ?v0) f1) (=> (= (f3 (f15 f16 f12) ?v1) f1) (= (f5 (f6 f8 (f5 f9 ?v0)) (f5 f9 ?v1)) (f5 f9 (f5 (f6 f8 ?v0) ?v1)))))))
(assume a87 (forall ((?v0 S3) (?v1 S3) (?v2 S2)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 ?v2 ?v0) f1) (=> (forall ((?v3 S3)) (=> (= (f3 (f15 f16 ?v0) ?v3) f1) (=> (= (f3 ?v2 ?v3) f1) (= (f3 ?v2 (f5 (f6 f7 ?v3) f14)) f1)))) (= (f3 ?v2 ?v1) f1))))))
(assume a88 (= (f3 (f15 f16 f13) f25) f1))
(assume a89 (forall ((?v0 S8) (?v1 S8)) (= (f18 (f19 f20 (f18 (f19 f24 ?v0) ?v1)) ?v1) (f18 (f19 f24 (f18 (f19 f20 ?v0) f21)) ?v1))))
(assume a90 (forall ((?v0 S3) (?v1 S3)) (= (f5 (f6 f7 (f5 (f6 f8 ?v0) ?v1)) ?v1) (f5 (f6 f8 (f5 (f6 f7 ?v0) f14)) ?v1))))
(assume a91 (forall ((?v0 S8) (?v1 S8)) (= (f18 (f19 f20 ?v0) (f18 (f19 f24 ?v1) ?v0)) (f18 (f19 f24 (f18 (f19 f20 ?v1) f21)) ?v0))))
(assume a92 (forall ((?v0 S3) (?v1 S3)) (= (f5 (f6 f7 ?v0) (f5 (f6 f8 ?v1) ?v0)) (f5 (f6 f8 (f5 (f6 f7 ?v1) f14)) ?v0))))
(assume a93 (forall ((?v0 S8)) (= (f18 (f19 f24 f26) ?v0) f26)))
(assume a94 (forall ((?v0 S3)) (= (f5 (f6 f8 f25) ?v0) f25)))
(assume a95 (forall ((?v0 S8)) (= (f18 (f19 f24 ?v0) f26) f26)))
(assume a96 (forall ((?v0 S3)) (= (f5 (f6 f8 ?v0) f25) f25)))
(assume a97 (forall ((?v0 S8)) (= (f18 (f19 f20 f26) ?v0) ?v0)))
(assume a98 (forall ((?v0 S3)) (= (f5 (f6 f7 f25) ?v0) ?v0)))
(assume a99 (forall ((?v0 S8)) (= (f18 (f19 f20 ?v0) f26) ?v0)))
(assume a100 (forall ((?v0 S3)) (= (f5 (f6 f7 ?v0) f25) ?v0)))
(assume a101 (forall ((?v0 S8) (?v1 S8)) (= (= ?v0 (f18 (f19 f20 ?v0) ?v1)) (= ?v1 f26))))
(assume a102 (forall ((?v0 S3) (?v1 S3)) (= (= ?v0 (f5 (f6 f7 ?v0) ?v1)) (= ?v1 f25))))
(assume a103 (forall ((?v0 S3)) (= (= (f5 (f6 f7 ?v0) ?v0) f25) (= ?v0 f25))))
(assume a104 (! (= f12 f25) :named @p_20))
(assume a105 (not (= f25 f14)))
(assume a106 (forall ((?v0 S3) (?v1 S3)) (= (= (f5 (f6 f7 (f5 (f6 f8 ?v0) ?v0)) (f5 (f6 f8 ?v1) ?v1)) f25) (and (= ?v0 f25) (= ?v1 f25)))))
(assume a107 (forall ((?v0 S8) (?v1 S8) (?v2 S8) (?v3 S8) (?v4 S8)) (=> (not (= ?v0 f26)) (=> (and (= ?v1 ?v2) (not (= ?v3 ?v4))) (not (= (f18 (f19 f20 ?v1) (f18 (f19 f24 ?v0) ?v3)) (f18 (f19 f20 ?v2) (f18 (f19 f24 ?v0) ?v4))))))))
(assume a108 (forall ((?v0 S3) (?v1 S3) (?v2 S3) (?v3 S3) (?v4 S3)) (=> (not (= ?v0 f25)) (=> (and (= ?v1 ?v2) (not (= ?v3 ?v4))) (not (= (f5 (f6 f7 ?v1) (f5 (f6 f8 ?v0) ?v3)) (f5 (f6 f7 ?v2) (f5 (f6 f8 ?v0) ?v4))))))))
(assume a109 (= f25 (! (f5 f9 f12) :named @p_21)))
(assume a110 (forall ((?v0 S3) (?v1 S3)) (= (f3 (f15 f16 f25) (f5 (f6 f7 (f5 (f6 f8 ?v0) ?v0)) (f5 (f6 f8 ?v1) ?v1))) f1)))
(assume a111 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 (f6 f7 (f5 (f6 f8 ?v0) ?v0)) (f5 (f6 f8 ?v1) ?v1))) f25) f1) (and (= ?v0 f25) (= ?v1 f25)))))
(assume a112 (= @p_21 f25))
(assume a113 (= (! (f22 f23 f12) :named @p_22) f26))
(assume a114 (forall ((?v0 S3)) (not (= (f5 (f6 f7 (f5 (f6 f7 f14) ?v0)) ?v0) f25))))
(assume a115 (forall ((?v0 S8) (?v1 S8)) (= (f18 (f19 f24 ?v0) ?v1) (f18 (f19 f24 ?v1) ?v0))))
(assume a116 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (f18 (f19 f24 ?v0) (f18 (f19 f24 ?v1) ?v2)) (f18 (f19 f24 ?v1) (f18 (f19 f24 ?v0) ?v2)))))
(assume a117 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (f5 (f6 f8 ?v0) (f5 (f6 f8 ?v1) ?v2)) (f5 (f6 f8 ?v1) (f5 (f6 f8 ?v0) ?v2)))))
(assume a118 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (f18 (f19 f24 ?v0) (f18 (f19 f24 ?v1) ?v2)) (f18 (f19 f24 (f18 (f19 f24 ?v0) ?v1)) ?v2))))
(assume a119 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (f5 (f6 f8 ?v0) (f5 (f6 f8 ?v1) ?v2)) (f5 (f6 f8 (f5 (f6 f8 ?v0) ?v1)) ?v2))))
(assume a120 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (f18 (f19 f24 (f18 (f19 f24 ?v0) ?v1)) ?v2) (f18 (f19 f24 ?v0) (f18 (f19 f24 ?v1) ?v2)))))
(assume a121 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (f18 (f19 f24 (f18 (f19 f24 ?v0) ?v1)) ?v2) (f18 (f19 f24 (f18 (f19 f24 ?v0) ?v2)) ?v1))))
(assume a122 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (f5 (f6 f8 (f5 (f6 f8 ?v0) ?v1)) ?v2) (f5 (f6 f8 (f5 (f6 f8 ?v0) ?v2)) ?v1))))
(assume a123 (forall ((?v0 S8) (?v1 S8) (?v2 S8) (?v3 S8)) (= (f18 (f19 f24 (f18 (f19 f24 ?v0) ?v1)) (f18 (f19 f24 ?v2) ?v3)) (f18 (f19 f24 ?v0) (f18 (f19 f24 ?v1) (f18 (f19 f24 ?v2) ?v3))))))
(assume a124 (forall ((?v0 S3) (?v1 S3) (?v2 S3) (?v3 S3)) (= (f5 (f6 f8 (f5 (f6 f8 ?v0) ?v1)) (f5 (f6 f8 ?v2) ?v3)) (f5 (f6 f8 ?v0) (f5 (f6 f8 ?v1) (f5 (f6 f8 ?v2) ?v3))))))
(assume a125 (forall ((?v0 S8) (?v1 S8) (?v2 S8) (?v3 S8)) (= (f18 (f19 f24 (f18 (f19 f24 ?v0) ?v1)) (f18 (f19 f24 ?v2) ?v3)) (f18 (f19 f24 ?v2) (f18 (f19 f24 (f18 (f19 f24 ?v0) ?v1)) ?v3)))))
(assume a126 (forall ((?v0 S3) (?v1 S3) (?v2 S3) (?v3 S3)) (= (f5 (f6 f8 (f5 (f6 f8 ?v0) ?v1)) (f5 (f6 f8 ?v2) ?v3)) (f5 (f6 f8 ?v2) (f5 (f6 f8 (f5 (f6 f8 ?v0) ?v1)) ?v3)))))
(assume a127 (forall ((?v0 S8) (?v1 S8) (?v2 S8) (?v3 S8)) (= (f18 (f19 f24 (f18 (f19 f24 ?v0) ?v1)) (f18 (f19 f24 ?v2) ?v3)) (f18 (f19 f24 (f18 (f19 f24 ?v0) ?v2)) (f18 (f19 f24 ?v1) ?v3)))))
(assume a128 (forall ((?v0 S3) (?v1 S3) (?v2 S3) (?v3 S3)) (= (f5 (f6 f8 (f5 (f6 f8 ?v0) ?v1)) (f5 (f6 f8 ?v2) ?v3)) (f5 (f6 f8 (f5 (f6 f8 ?v0) ?v2)) (f5 (f6 f8 ?v1) ?v3)))))
(assume a129 (forall ((?v0 S3)) (= (f5 f9 (f5 f10 ?v0)) (f5 (f6 f7 (f5 (f6 f7 f25) (f5 f9 ?v0))) (f5 f9 ?v0)))))
(assume a130 (forall ((?v0 S8) (?v1 S8)) (= (f18 (f19 f20 ?v0) ?v1) (f18 (f19 f20 ?v1) ?v0))))
(assume a131 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (f18 (f19 f20 ?v0) (f18 (f19 f20 ?v1) ?v2)) (f18 (f19 f20 ?v1) (f18 (f19 f20 ?v0) ?v2)))))
(assume a132 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (f18 (f19 f20 ?v0) (f18 (f19 f20 ?v1) ?v2)) (f18 (f19 f20 (f18 (f19 f20 ?v0) ?v1)) ?v2))))
(assume a133 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (f5 (f6 f7 ?v0) (f5 (f6 f7 ?v1) ?v2)) (f5 (f6 f7 (f5 (f6 f7 ?v0) ?v1)) ?v2))))
(assume a134 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (f18 (f19 f20 (f18 (f19 f20 ?v0) ?v1)) ?v2) (f18 (f19 f20 ?v0) (f18 (f19 f20 ?v1) ?v2)))))
(assume a135 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (f18 (f19 f20 (f18 (f19 f20 ?v0) ?v1)) ?v2) (f18 (f19 f20 (f18 (f19 f20 ?v0) ?v2)) ?v1))))
(assume a136 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (f5 (f6 f7 (f5 (f6 f7 ?v0) ?v1)) ?v2) (f5 (f6 f7 (f5 (f6 f7 ?v0) ?v2)) ?v1))))
(assume a137 (forall ((?v0 S8) (?v1 S8) (?v2 S8) (?v3 S8)) (= (f18 (f19 f20 (f18 (f19 f20 ?v0) ?v1)) (f18 (f19 f20 ?v2) ?v3)) (f18 (f19 f20 (f18 (f19 f20 ?v0) ?v2)) (f18 (f19 f20 ?v1) ?v3)))))
(assume a138 (forall ((?v0 S3) (?v1 S3) (?v2 S3) (?v3 S3)) (= (f5 (f6 f7 (f5 (f6 f7 ?v0) ?v1)) (f5 (f6 f7 ?v2) ?v3)) (f5 (f6 f7 (f5 (f6 f7 ?v0) ?v2)) (f5 (f6 f7 ?v1) ?v3)))))
(assume a139 (forall ((?v0 S3)) (= (= (f3 (f15 f16 f25) (f5 f9 ?v0)) f1) (= (f3 (f15 f16 f12) ?v0) f1))))
(assume a140 (forall ((?v0 S3)) (= (= (f3 (f15 f16 (f5 f9 ?v0)) f25) f1) (= (f3 (f15 f16 ?v0) f12) f1))))
(assume a141 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (f18 (f19 f24 ?v0) (f18 (f19 f20 ?v1) ?v2)) (f18 (f19 f20 (f18 (f19 f24 ?v0) ?v1)) (f18 (f19 f24 ?v0) ?v2)))))
(assume a142 (forall ((?v0 S8) (?v1 S8) (?v2 S8) (?v3 S8)) (= (and (not (= ?v0 ?v1)) (not (= ?v2 ?v3))) (not (= (f18 (f19 f20 (f18 (f19 f24 ?v0) ?v2)) (f18 (f19 f24 ?v1) ?v3)) (f18 (f19 f20 (f18 (f19 f24 ?v0) ?v3)) (f18 (f19 f24 ?v1) ?v2)))))))
(assume a143 (forall ((?v0 S3) (?v1 S3) (?v2 S3) (?v3 S3)) (= (and (not (= ?v0 ?v1)) (not (= ?v2 ?v3))) (not (= (f5 (f6 f7 (f5 (f6 f8 ?v0) ?v2)) (f5 (f6 f8 ?v1) ?v3)) (f5 (f6 f7 (f5 (f6 f8 ?v0) ?v3)) (f5 (f6 f8 ?v1) ?v2)))))))
(assume a144 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (f18 (f19 f24 (f18 (f19 f20 ?v0) ?v1)) ?v2) (f18 (f19 f20 (f18 (f19 f24 ?v0) ?v2)) (f18 (f19 f24 ?v1) ?v2)))))
(assume a145 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (f18 (f19 f20 (f18 (f19 f24 ?v0) ?v1)) (f18 (f19 f24 ?v2) ?v1)) (f18 (f19 f24 (f18 (f19 f20 ?v0) ?v2)) ?v1))))
(assume a146 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (f5 (f6 f7 (f5 (f6 f8 ?v0) ?v1)) (f5 (f6 f8 ?v2) ?v1)) (f5 (f6 f8 (f5 (f6 f7 ?v0) ?v2)) ?v1))))
(assume a147 (forall ((?v0 S8) (?v1 S8) (?v2 S8) (?v3 S8)) (= (= (f18 (f19 f20 (f18 (f19 f24 ?v0) ?v1)) (f18 (f19 f24 ?v2) ?v3)) (f18 (f19 f20 (f18 (f19 f24 ?v0) ?v3)) (f18 (f19 f24 ?v2) ?v1))) (or (= ?v0 ?v2) (= ?v1 ?v3)))))
(assume a148 (forall ((?v0 S3) (?v1 S3) (?v2 S3) (?v3 S3)) (= (= (f5 (f6 f7 (f5 (f6 f8 ?v0) ?v1)) (f5 (f6 f8 ?v2) ?v3)) (f5 (f6 f7 (f5 (f6 f8 ?v0) ?v3)) (f5 (f6 f8 ?v2) ?v1))) (or (= ?v0 ?v2) (= ?v1 ?v3)))))
(assume a149 (forall ((?v0 S8)) (= (f18 (f19 f24 f21) ?v0) ?v0)))
(assume a150 (forall ((?v0 S8)) (= (f18 (f19 f24 ?v0) f21) ?v0)))
(assume a151 (= f21 @p_17))
(assume a152 (forall ((?v0 S3) (?v1 S3)) (= (= (f27 (f28 f29 (f22 f23 ?v0)) (f22 f23 ?v1)) f1) (ite (= (f3 (f15 f16 ?v0) ?v1) f1) true (= (f3 (f15 f16 ?v0) f12) f1)))))
(assume a153 (forall ((?v0 S8)) (= (f18 (f19 f20 ?v0) ?v0) (f18 (f19 f24 (f18 (f19 f20 f21) f21)) ?v0))))
(assume a154 (forall ((?v0 S3)) (= (f5 (f6 f7 ?v0) ?v0) (f5 (f6 f8 (f5 (f6 f7 f14) f14)) ?v0))))
(assume a155 (! (= (! (f3 (! (f15 f16 f25) :named @p_23) @p_16) :named @p_24) f1) :named @p_25))
(assume a156 (forall ((?v0 S3)) (=> (= (f3 (f15 f16 f25) (f5 f9 ?v0)) f1) (and (= (f3 (f15 f16 f25) (f5 f9 (f5 f10 ?v0))) f1) (= (f3 (f15 f16 f25) (f5 f9 (f5 f11 ?v0))) f1)))))
(assume a157 (= (f3 @p_23 (! (f5 f9 (! (f5 f11 @p_2) :named @p_26)) :named @p_27)) f1))
(assume a158 (forall ((?v0 S3) (?v1 S3) (?v2 S3) (?v3 S3) (?v4 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v2) ?v1) f1) (=> (= (f3 (f15 f16 f25) ?v3) f1) (=> (= (f3 (f15 f16 f25) ?v4) f1) (=> (= (f5 (f6 f7 ?v3) ?v4) f14) (= (f3 (f15 f16 (f5 (f6 f7 (f5 (f6 f8 ?v3) ?v0)) (f5 (f6 f8 ?v4) ?v2))) ?v1) f1))))))))
(assume a159 (= (f3 @p_23 @p_21) f1))
(assume a160 (= f26 @p_22))
(assume a161 (forall ((?v0 S3)) (= (= (f22 f23 ?v0) f26) (= (f3 (f15 f16 ?v0) f12) f1))))
(assume a162 (forall ((?v0 S3)) (= (= f26 (f22 f23 ?v0)) (= (f3 (f15 f16 ?v0) f12) f1))))
(assume a163 (forall ((?v0 S8) (?v1 S8)) (=> (= (f18 (f19 f24 ?v0) ?v1) f26) (or (= ?v0 f26) (= ?v1 f26)))))
(assume a164 (forall ((?v0 S3) (?v1 S3)) (=> (= (f5 (f6 f8 ?v0) ?v1) f25) (or (= ?v0 f25) (= ?v1 f25)))))
(assume a165 (forall ((?v0 S8) (?v1 S8)) (=> (not (= ?v0 f26)) (=> (not (= ?v1 f26)) (not (= (f18 (f19 f24 ?v0) ?v1) f26))))))
(assume a166 (forall ((?v0 S3) (?v1 S3)) (=> (not (= ?v0 f25)) (=> (not (= ?v1 f25)) (not (= (f5 (f6 f8 ?v0) ?v1) f25))))))
(assume a167 (forall ((?v0 S3) (?v1 S3)) (= (= (f5 (f6 f8 ?v0) ?v1) f25) (or (= ?v0 f25) (= ?v1 f25)))))
(assume a168 (not (= f26 f21)))
(assume a169 (not (= f21 f26)))
(assume a170 (not (= f14 f25)))
(assume a171 (forall ((?v0 S8) (?v1 S8) (?v2 S8) (?v3 S8)) (= (f18 (f19 f20 (f18 (f19 f24 ?v0) ?v1)) (f18 (f19 f20 (f18 (f19 f24 ?v2) ?v1)) ?v3)) (f18 (f19 f20 (f18 (f19 f24 (f18 (f19 f20 ?v0) ?v2)) ?v1)) ?v3))))
(assume a172 (forall ((?v0 S3) (?v1 S3) (?v2 S3) (?v3 S3)) (= (f5 (f6 f7 (f5 (f6 f8 ?v0) ?v1)) (f5 (f6 f7 (f5 (f6 f8 ?v2) ?v1)) ?v3)) (f5 (f6 f7 (f5 (f6 f8 (f5 (f6 f7 ?v0) ?v2)) ?v1)) ?v3))))
(assume a173 (forall ((?v0 S3) (?v1 S1) (?v2 S1)) (=> (=> (= (f3 (f15 f16 f25) ?v0) f1) (= (= ?v1 f1) (= ?v2 f1))) (= (=> (= (f3 (f15 f16 f25) ?v0) f1) (= ?v1 f1)) (=> (= (f3 (f15 f16 f25) ?v0) f1) (= ?v2 f1))))))
(assume a174 (forall ((?v0 S3) (?v1 S1) (?v2 S1)) (=> (=> (= (f3 (f15 f16 f25) ?v0) f1) (= (= ?v1 f1) (= ?v2 f1))) (= (and (= (f3 (f15 f16 f25) ?v0) f1) (= ?v1 f1)) (and (= (f3 (f15 f16 f25) ?v0) f1) (= ?v2 f1))))))
(assume a175 (= (f3 @p_23 f25) f1))
(assume a176 (forall ((?v0 S8) (?v1 S8)) (=> (or (and (= (f27 (f28 f29 f26) ?v0) f1) (= (f27 (f28 f29 ?v1) f26) f1)) (and (= (f27 (f28 f29 ?v0) f26) f1) (= (f27 (f28 f29 f26) ?v1) f1))) (= (f27 (f28 f29 (f18 (f19 f24 ?v0) ?v1)) f26) f1))))
(assume a177 (forall ((?v0 S3) (?v1 S3)) (=> (or (and (= (f3 (f15 f16 f25) ?v0) f1) (= (f3 (f15 f16 ?v1) f25) f1)) (and (= (f3 (f15 f16 ?v0) f25) f1) (= (f3 (f15 f16 f25) ?v1) f1))) (= (f3 (f15 f16 (f5 (f6 f8 ?v0) ?v1)) f25) f1))))
(assume a178 (forall ((?v0 S3) (?v1 S3)) (=> (or (and (= (f3 (f15 f16 f25) ?v0) f1) (= (f3 (f15 f16 f25) ?v1) f1)) (and (= (f3 (f15 f16 ?v0) f25) f1) (= (f3 (f15 f16 ?v1) f25) f1))) (= (f3 (f15 f16 f25) (f5 (f6 f8 ?v0) ?v1)) f1))))
(assume a179 (forall ((?v0 S8) (?v1 S8) (?v2 S8) (?v3 S8)) (=> (= (f27 (f28 f29 ?v0) ?v1) f1) (=> (= (f27 (f28 f29 ?v2) ?v3) f1) (=> (= (f27 (f28 f29 f26) ?v1) f1) (=> (= (f27 (f28 f29 f26) ?v2) f1) (= (f27 (f28 f29 (f18 (f19 f24 ?v0) ?v2)) (f18 (f19 f24 ?v1) ?v3)) f1)))))))
(assume a180 (forall ((?v0 S3) (?v1 S3) (?v2 S3) (?v3 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v2) ?v3) f1) (=> (= (f3 (f15 f16 f25) ?v1) f1) (=> (= (f3 (f15 f16 f25) ?v2) f1) (= (f3 (f15 f16 (f5 (f6 f8 ?v0) ?v2)) (f5 (f6 f8 ?v1) ?v3)) f1)))))))
(assume a181 (forall ((?v0 S8) (?v1 S8) (?v2 S8) (?v3 S8)) (=> (= (f27 (f28 f29 ?v0) ?v1) f1) (=> (= (f27 (f28 f29 ?v2) ?v3) f1) (=> (= (f27 (f28 f29 f26) ?v0) f1) (=> (= (f27 (f28 f29 f26) ?v2) f1) (= (f27 (f28 f29 (f18 (f19 f24 ?v0) ?v2)) (f18 (f19 f24 ?v1) ?v3)) f1)))))))
(assume a182 (forall ((?v0 S3) (?v1 S3) (?v2 S3) (?v3 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v2) ?v3) f1) (=> (= (f3 (f15 f16 f25) ?v0) f1) (=> (= (f3 (f15 f16 f25) ?v2) f1) (= (f3 (f15 f16 (f5 (f6 f8 ?v0) ?v2)) (f5 (f6 f8 ?v1) ?v3)) f1)))))))
(assume a183 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v2) f25) f1) (= (f3 (f15 f16 (f5 (f6 f8 ?v2) ?v1)) (f5 (f6 f8 ?v2) ?v0)) f1)))))
(assume a184 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v2) f25) f1) (= (f3 (f15 f16 (f5 (f6 f8 ?v1) ?v2)) (f5 (f6 f8 ?v0) ?v2)) f1)))))
(assume a185 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f27 (f28 f29 ?v0) ?v1) f1) (=> (= (f27 (f28 f29 f26) ?v2) f1) (= (f27 (f28 f29 (f18 (f19 f24 ?v2) ?v0)) (f18 (f19 f24 ?v2) ?v1)) f1)))))
(assume a186 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 f25) ?v2) f1) (= (f3 (f15 f16 (f5 (f6 f8 ?v2) ?v0)) (f5 (f6 f8 ?v2) ?v1)) f1)))))
(assume a187 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f27 (f28 f29 ?v0) ?v1) f1) (=> (= (f27 (f28 f29 f26) ?v2) f1) (= (f27 (f28 f29 (f18 (f19 f24 ?v0) ?v2)) (f18 (f19 f24 ?v1) ?v2)) f1)))))
(assume a188 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 f25) ?v2) f1) (= (f3 (f15 f16 (f5 (f6 f8 ?v0) ?v2)) (f5 (f6 f8 ?v1) ?v2)) f1)))))
(assume a189 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 ?v0) f25) f1) (=> (= (f3 (f15 f16 ?v1) f25) f1) (= (f3 (f15 f16 f25) (f5 (f6 f8 ?v0) ?v1)) f1)))))
(assume a190 (forall ((?v0 S8) (?v1 S8)) (=> (= (f27 (f28 f29 ?v0) f26) f1) (=> (= (f27 (f28 f29 f26) ?v1) f1) (= (f27 (f28 f29 (f18 (f19 f24 ?v0) ?v1)) f26) f1)))))
(assume a191 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 ?v0) f25) f1) (=> (= (f3 (f15 f16 f25) ?v1) f1) (= (f3 (f15 f16 (f5 (f6 f8 ?v0) ?v1)) f25) f1)))))
(assume a192 (forall ((?v0 S8) (?v1 S8)) (=> (= (f27 (f28 f29 f26) ?v0) f1) (=> (= (f27 (f28 f29 ?v1) f26) f1) (= (f27 (f28 f29 (f18 (f19 f24 ?v1) ?v0)) f26) f1)))))
(assume a193 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 f25) ?v0) f1) (=> (= (f3 (f15 f16 ?v1) f25) f1) (= (f3 (f15 f16 (f5 (f6 f8 ?v1) ?v0)) f25) f1)))))
(assume a194 (forall ((?v0 S8) (?v1 S8)) (=> (= (f27 (f28 f29 f26) ?v0) f1) (=> (= (f27 (f28 f29 ?v1) f26) f1) (= (f27 (f28 f29 (f18 (f19 f24 ?v0) ?v1)) f26) f1)))))
(assume a195 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 f25) ?v0) f1) (=> (= (f3 (f15 f16 ?v1) f25) f1) (= (f3 (f15 f16 (f5 (f6 f8 ?v0) ?v1)) f25) f1)))))
(assume a196 (forall ((?v0 S8) (?v1 S8)) (=> (= (f27 (f28 f29 f26) ?v0) f1) (=> (= (f27 (f28 f29 f26) ?v1) f1) (= (f27 (f28 f29 f26) (f18 (f19 f24 ?v0) ?v1)) f1)))))
(assume a197 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 f25) ?v0) f1) (=> (= (f3 (f15 f16 f25) ?v1) f1) (= (f3 (f15 f16 f25) (f5 (f6 f8 ?v0) ?v1)) f1)))))
(assume a198 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 (f6 f8 ?v0) ?v1)) f25) f1) (or (and (= (f3 (f15 f16 f25) ?v0) f1) (= (f3 (f15 f16 ?v1) f25) f1)) (and (= (f3 (f15 f16 ?v0) f25) f1) (= (f3 (f15 f16 f25) ?v1) f1))))))
(assume a199 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 f25) (f5 (f6 f8 ?v0) ?v1)) f1) (or (and (= (f3 (f15 f16 f25) ?v0) f1) (= (f3 (f15 f16 f25) ?v1) f1)) (and (= (f3 (f15 f16 ?v0) f25) f1) (= (f3 (f15 f16 ?v1) f25) f1))))))
(assume a200 (forall ((?v0 S3)) (= (f3 (f15 f16 f25) (f5 (f6 f8 ?v0) ?v0)) f1)))
(assume a201 (= (f27 (f28 f29 f26) f21) f1))
(assume a202 (= (f3 @p_23 f14) f1))
(assume a203 (not (= (f27 (f28 f29 f21) f26) f1)))
(assume a204 (! (not (! (= (! (f3 @p_15 f25) :named @p_28) f1) :named @p_29)) :named @p_30))
(assume a205 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 f25) ?v0) f1) (=> (= (f3 (f15 f16 f25) ?v1) f1) (= (f3 (f15 f16 f25) (f5 (f6 f7 ?v0) ?v1)) f1)))))
(assume a206 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 f25) ?v0) f1) (=> (= (f3 (f15 f16 f25) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) f14) f1) (= (f3 (f15 f16 (f5 (f6 f8 ?v1) ?v0)) ?v0) f1))))))
(assume a207 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 f25) ?v0) f1) (=> (= (f3 (f15 f16 f25) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) f14) f1) (= (f3 (f15 f16 (f5 (f6 f8 ?v0) ?v1)) ?v0) f1))))))
(assume a208 (forall ((?v0 S8) (?v1 S8)) (=> (= (f27 (f28 f29 ?v0) f26) f1) (=> (= (f27 (f28 f29 ?v1) f26) f1) (= (f27 (f28 f29 (f18 (f19 f20 ?v0) ?v1)) f26) f1)))))
(assume a209 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 ?v0) f25) f1) (=> (= (f3 (f15 f16 ?v1) f25) f1) (= (f3 (f15 f16 (f5 (f6 f7 ?v0) ?v1)) f25) f1)))))
(assume a210 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f27 (f28 f29 f26) ?v0) f1) (=> (= (f27 (f28 f29 ?v1) ?v2) f1) (= (f27 (f28 f29 ?v1) (f18 (f19 f20 ?v2) ?v0)) f1)))))
(assume a211 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 f25) ?v0) f1) (=> (= (f3 (f15 f16 ?v1) ?v2) f1) (= (f3 (f15 f16 ?v1) (f5 (f6 f7 ?v2) ?v0)) f1)))))
(assume a212 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f27 (f28 f29 f26) ?v0) f1) (=> (= (f27 (f28 f29 ?v1) ?v2) f1) (= (f27 (f28 f29 ?v1) (f18 (f19 f20 ?v0) ?v2)) f1)))))
(assume a213 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 f25) ?v0) f1) (=> (= (f3 (f15 f16 ?v1) ?v2) f1) (= (f3 (f15 f16 ?v1) (f5 (f6 f7 ?v0) ?v2)) f1)))))
(assume a214 (forall ((?v0 S8)) (= (= f26 ?v0) (= ?v0 f26))))
(assume a215 (forall ((?v0 S3)) (= (= f25 ?v0) (= ?v0 f25))))
(assume a216 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (= (f18 (f19 f20 ?v0) ?v1) (f18 (f19 f20 ?v0) ?v2)) (= ?v1 ?v2))))
(assume a217 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (= (f5 (f6 f7 ?v0) ?v1) (f5 (f6 f7 ?v0) ?v2)) (= ?v1 ?v2))))
(assume a218 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (= (f18 (f19 f20 ?v0) ?v1) (f18 (f19 f20 ?v2) ?v1)) (= ?v0 ?v2))))
(assume a219 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (= (f5 (f6 f7 ?v0) ?v1) (f5 (f6 f7 ?v2) ?v1)) (= ?v0 ?v2))))
(assume a220 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f18 (f19 f20 ?v0) ?v1) (f18 (f19 f20 ?v0) ?v2)) (= ?v1 ?v2))))
(assume a221 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f5 (f6 f7 ?v0) ?v1) (f5 (f6 f7 ?v0) ?v2)) (= ?v1 ?v2))))
(assume a222 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f18 (f19 f20 ?v0) ?v1) (f18 (f19 f20 ?v2) ?v1)) (= ?v0 ?v2))))
(assume a223 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f5 (f6 f7 ?v0) ?v1) (f5 (f6 f7 ?v2) ?v1)) (= ?v0 ?v2))))
(assume a224 (forall ((?v0 S8)) (= (= f21 ?v0) (= ?v0 f21))))
(assume a225 (forall ((?v0 S3)) (= (= f14 ?v0) (= ?v0 f14))))
(assume a226 (forall ((?v0 S3)) (= (= f25 (f5 (f6 f7 ?v0) ?v0)) (= ?v0 f25))))
(assume a227 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (= (f27 (f28 f29 (f18 (f19 f20 ?v0) ?v1)) (f18 (f19 f20 ?v2) ?v1)) f1) (= (f27 (f28 f29 ?v0) ?v2) f1))))
(assume a228 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (= (f3 (f15 f16 (f5 (f6 f7 ?v0) ?v1)) (f5 (f6 f7 ?v2) ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v2) f1))))
(assume a229 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (= (f27 (f28 f29 (f18 (f19 f20 ?v0) ?v1)) (f18 (f19 f20 ?v0) ?v2)) f1) (= (f27 (f28 f29 ?v1) ?v2) f1))))
(assume a230 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (= (= (f3 (f15 f16 (f5 (f6 f7 ?v0) ?v1)) (f5 (f6 f7 ?v0) ?v2)) f1) (= (f3 (f15 f16 ?v1) ?v2) f1))))
(assume a231 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f27 (f28 f29 ?v0) ?v1) f1) (= (f27 (f28 f29 (f18 (f19 f20 ?v0) ?v2)) (f18 (f19 f20 ?v1) ?v2)) f1))))
(assume a232 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (= (f3 (f15 f16 (f5 (f6 f7 ?v0) ?v2)) (f5 (f6 f7 ?v1) ?v2)) f1))))
(assume a233 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f27 (f28 f29 ?v0) ?v1) f1) (= (f27 (f28 f29 (f18 (f19 f20 ?v2) ?v0)) (f18 (f19 f20 ?v2) ?v1)) f1))))
(assume a234 (forall ((?v0 S8) (?v1 S8) (?v2 S8) (?v3 S8)) (=> (= (f27 (f28 f29 ?v0) ?v1) f1) (=> (= (f27 (f28 f29 ?v2) ?v3) f1) (= (f27 (f28 f29 (f18 (f19 f20 ?v0) ?v2)) (f18 (f19 f20 ?v1) ?v3)) f1)))))
(assume a235 (forall ((?v0 S3) (?v1 S3) (?v2 S3) (?v3 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v2) ?v3) f1) (= (f3 (f15 f16 (f5 (f6 f7 ?v0) ?v2)) (f5 (f6 f7 ?v1) ?v3)) f1)))))
(assume a236 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f27 (f28 f29 (f18 (f19 f20 ?v0) ?v1)) (f18 (f19 f20 ?v2) ?v1)) f1) (= (f27 (f28 f29 ?v0) ?v2) f1))))
(assume a237 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 (f5 (f6 f7 ?v0) ?v1)) (f5 (f6 f7 ?v2) ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v2) f1))))
(assume a238 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f27 (f28 f29 (f18 (f19 f20 ?v0) ?v1)) (f18 (f19 f20 ?v0) ?v2)) f1) (= (f27 (f28 f29 ?v1) ?v2) f1))))
(assume a239 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 (f5 (f6 f7 ?v0) ?v1)) (f5 (f6 f7 ?v0) ?v2)) f1) (= (f3 (f15 f16 ?v1) ?v2) f1))))
(assume a240 (forall ((?v0 S3)) (= (= (f3 (f15 f16 f25) (f5 (f6 f7 ?v0) ?v0)) f1) (= (f3 (f15 f16 f25) ?v0) f1))))
(assume a241 (forall ((?v0 S3)) (= (= (f3 (f15 f16 (f5 (f6 f7 ?v0) ?v0)) f25) f1) (= (f3 (f15 f16 ?v0) f25) f1))))
(assume a242 (forall ((?v0 S8) (?v1 S8)) (=> (= (f27 (f28 f29 f26) ?v0) f1) (=> (= (f27 (f28 f29 f26) ?v1) f1) (= (f27 (f28 f29 f26) (f18 (f19 f20 ?v0) ?v1)) f1)))))
(assume a243 (forall ((?v0 S8) (?v1 S8)) (=> (= (f27 (f28 f29 f26) ?v0) f1) (=> (= (f27 (f28 f29 f26) ?v1) f1) (= (= (f18 (f19 f20 ?v0) ?v1) f26) (and (= ?v0 f26) (= ?v1 f26)))))))
(assume a244 (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 f25) ?v0) f1) (=> (= (f3 (f15 f16 f25) ?v1) f1) (= (= (f5 (f6 f7 ?v0) ?v1) f25) (and (= ?v0 f25) (= ?v1 f25)))))))
(assume a245 (forall ((?v0 S8)) (= (f27 (f28 f29 f26) ?v0) f1)))
(assume a246 (forall ((?v0 S8) (?v1 S8)) (=> (= ?v0 (f18 (f19 f24 ?v0) ?v1)) (or (= ?v1 f21) (= ?v0 f26)))))
(assume a247 (forall ((?v0 S8) (?v1 S8)) (=> (= (f27 (f28 f29 ?v0) ?v1) f1) (=> (= (f27 (f28 f29 ?v1) ?v0) f1) (= ?v0 ?v1)))))
(assume a248 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f27 (f28 f29 ?v0) ?v1) f1) (=> (= (f27 (f28 f29 ?v1) ?v2) f1) (= (f27 (f28 f29 ?v0) ?v2) f1)))))
(assume a249 (forall ((?v0 S8) (?v1 S8)) (=> (= ?v0 ?v1) (= (f27 (f28 f29 ?v0) ?v1) f1))))
(assume a250 (forall ((?v0 S8) (?v1 S8)) (or (= (f27 (f28 f29 ?v0) ?v1) f1) (= (f27 (f28 f29 ?v1) ?v0) f1))))
(assume a251 (forall ((?v0 S8)) (= (f27 (f28 f29 ?v0) ?v0) f1)))
(assume a252 (forall ((?v0 S8) (?v1 S8)) (= (= (f18 (f19 f20 ?v0) ?v1) f26) (and (= ?v0 f26) (= ?v1 f26)))))
(assume a253 (forall ((?v0 S8) (?v1 S8)) (=> (= (f18 (f19 f20 ?v0) ?v1) ?v0) (= ?v1 f26))))
(assume a254 (forall ((?v0 S8)) (= (= (f27 (f28 f29 f26) ?v0) f1) true)))
(assume a255 (forall ((?v0 S8)) (= (= (f27 (f28 f29 ?v0) f26) f1) (= ?v0 f26))))
(assume a256 (forall ((?v0 S8) (?v1 S8)) (= (= (f18 (f19 f24 ?v0) ?v1) f26) (or (= ?v0 f26) (= ?v1 f26)))))
(assume a257 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (= (f18 (f19 f24 ?v0) ?v1) (f18 (f19 f24 ?v0) ?v2)) (or (= ?v1 ?v2) (= ?v0 f26)))))
(assume a258 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (= (= (f18 (f19 f24 ?v0) ?v1) (f18 (f19 f24 ?v2) ?v1)) (or (= ?v0 ?v2) (= ?v1 f26)))))
(assume a259 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f27 (f28 f29 (f18 (f19 f20 ?v0) ?v1)) ?v2) f1) (=> (=> (= (f27 (f28 f29 ?v0) ?v2) f1) (=> (= (f27 (f28 f29 ?v1) ?v2) f1) false)) false))))
(assume a260 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f27 (f28 f29 (f18 (f19 f20 ?v0) ?v1)) ?v2) f1) (= (f27 (f28 f29 ?v0) ?v2) f1))))
(assume a261 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f27 (f28 f29 (f18 (f19 f20 ?v0) ?v1)) ?v2) f1) (= (f27 (f28 f29 ?v1) ?v2) f1))))
(assume a262 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f27 (f28 f29 ?v0) ?v1) f1) (= (f27 (f28 f29 ?v0) (f18 (f19 f20 ?v2) ?v1)) f1))))
(assume a263 (forall ((?v0 S8) (?v1 S8) (?v2 S8)) (=> (= (f27 (f28 f29 ?v0) ?v1) f1) (= (f27 (f28 f29 ?v0) (f18 (f19 f20 ?v1) ?v2)) f1))))
(step t0 (cl (not (! (= (! (or (! (not (! (not (! (= (! (f3 f4 (! (f5 (! (f6 f7 (! (f5 (! (f6 f8 (! (f5 f9 (! (f5 f10 (! (f5 f10 (! (f5 f11 f25) :named @p_31)) :named @p_33)) :named @p_34)) :named @p_35)) :named @p_36) f13) :named @p_37)) :named @p_38) (! (f5 f9 @p_31) :named @p_32)) :named @p_39)) :named @p_42) (! (f3 (! (f15 f16 (! (f5 f11 @p_33) :named @p_55)) :named @p_314) (! (f5 f11 f25) :named @p_31)) :named @p_315)) :named @p_316)) :named @p_317)) :named @p_370) (! (not (! (= @p_55 (! (f5 f9 @p_55) :named @p_56)) :named @p_205)) :named @p_288) (! (not (! (= @p_33 (! (f5 f9 @p_33) :named @p_53)) :named @p_54)) :named @p_197) (! (not (! (= @p_31 @p_32) :named @p_47)) :named @p_159) (! (not (! (= @p_32 @p_39) :named @p_206)) :named @p_368) (! (not (! (= (! (f3 f4 (! (f5 (! (f6 f7 (! (f5 (! (f6 f8 (! (f5 f9 (! (f5 f10 (! (f5 f10 (! (f5 f11 f25) :named @p_31)) :named @p_33)) :named @p_34)) :named @p_35)) :named @p_36) f13) :named @p_37)) :named @p_38) (! (f5 f9 @p_31) :named @p_32)) :named @p_39)) :named @p_42) (! (f3 (! (f15 f16 @p_56) :named @p_57) (! (f5 (! (f6 f7 (! (f5 (! (f6 f8 (! (f5 f9 (! (f5 f10 (! (f5 f10 (! (f5 f11 f25) :named @p_31)) :named @p_33)) :named @p_34)) :named @p_35)) :named @p_36) f13) :named @p_37)) :named @p_38) (! (f5 f9 @p_31) :named @p_32)) :named @p_39)) :named @p_58)) :named @p_59)) :named @p_287) (! (not (! (= (! (f3 f4 (! (f5 (! (f6 f7 (! (f5 (! (f6 f8 (! (f5 f9 (! (f5 f10 (! (f5 f10 (! (f5 f11 f25) :named @p_31)) :named @p_33)) :named @p_34)) :named @p_35)) :named @p_36) f13) :named @p_37)) :named @p_38) (! (f5 f9 @p_31) :named @p_32)) :named @p_39)) :named @p_42) (! (f3 (! (f15 f16 @p_39) :named @p_40) (! (f5 f9 @p_31) :named @p_32)) :named @p_41)) :named @p_43)) :named @p_158)) :named @p_371) (! (or @p_316 @p_288 @p_197 @p_159 @p_368 @p_287 @p_158) :named @p_369)) :named @p_380)) (not @p_371) @p_369) :rule equiv_pos2)
(step t1 (cl (! (= @p_370 @p_316) :named @p_381)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_381 1 5))
(step t2 (cl (= @p_288 @p_288)) :rule refl)
(step t3 (cl (= @p_197 @p_197)) :rule refl)
(step t4 (cl (= @p_159 @p_159)) :rule refl)
(step t5 (cl (= @p_368 @p_368)) :rule refl)
(step t6 (cl (= @p_287 @p_287)) :rule refl)
(step t7 (cl (= @p_158 @p_158)) :rule refl)
(step t8 (cl @p_380) :rule cong :premises (t1 t2 t3 t4 t5 t6 t7))
(step t9 (cl (! (=> (! (and @p_317 @p_205 @p_54 @p_47 @p_206 @p_59 @p_43) :named @p_373) false) :named @p_375) @p_373) :rule implies_neg1)
(anchor :step t10)
(assume t10.a0 @p_317)
(assume t10.a1 @p_205)
(assume t10.a2 @p_54)
(assume t10.a3 @p_47)
(assume t10.a4 @p_206)
(assume t10.a5 @p_59)
(assume t10.a6 @p_43)
(step t10.t0 (cl (not (! (= (! (= true false) :named @p_376) false) :named @p_379)) (not @p_376) false) :rule equiv_pos2)
(step t10.t1 (cl @p_379) :rule hole :args ("TRUST_THEORY_REWRITE" @p_379 1 5))
(step t10.t2 (cl (= (! (= @p_316 true) :named @p_378) @p_316)) :rule equiv_simplify)
(step t10.t3 (cl @p_378 @p_317) :rule equiv2 :premises (t10.t2))
(step t10.t4 (cl (not (! (= @p_14 @p_43) :named @p_44)) (not @p_14) @p_43) :rule equiv_pos2)
(step t10.t5 (cl (! (= @p_14 (! (= f1 @p_13) :named @p_45)) :named @p_46)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_46 2 6))
(step t10.t6 (cl (= f1 @p_10)) :rule symm :premises (a2))
(step t10.t7 (cl (= f4 f4)) :rule refl)
(step t10.t8 (cl (= f7 f7)) :rule refl)
(step t10.t9 (cl (= f8 f8)) :rule refl)
(step t10.t10 (cl (= f9 f9)) :rule refl)
(step t10.t11 (cl (= f10 f10)) :rule refl)
(step t10.t12 (cl (= f11 f11)) :rule refl)
(step t10.t13 (cl (= @p_2 @p_31)) :rule cong :premises (t10.t12 a104))
(step t10.t14 (cl (= @p_3 @p_33)) :rule cong :premises (t10.t11 t10.t13))
(step t10.t15 (cl (= @p_4 @p_34)) :rule cong :premises (t10.t11 t10.t14))
(step t10.t16 (cl (= @p_5 @p_35)) :rule cong :premises (t10.t10 t10.t15))
(step t10.t17 (cl (= @p_6 @p_36)) :rule cong :premises (t10.t9 t10.t16))
(step t10.t18 (cl (= f13 f13)) :rule refl)
(step t10.t19 (cl (= @p_7 @p_37)) :rule cong :premises (t10.t17 t10.t18))
(step t10.t20 (cl (= @p_8 @p_38)) :rule cong :premises (t10.t8 t10.t19))
(step t10.t21 (cl (= @p_18 @p_32)) :rule cong :premises (t10.t10 t10.t13))
(step t10.t22 (cl (= f14 @p_32)) :rule trans :premises (a80 t10.t21))
(step t10.t23 (cl (= @p_9 @p_39)) :rule cong :premises (t10.t20 t10.t22))
(step t10.t24 (cl (= @p_10 @p_42)) :rule cong :premises (t10.t7 t10.t23))
(step t10.t25 (cl (= f1 @p_42)) :rule trans :premises (t10.t6 t10.t24))
(step t10.t26 (cl (= f16 f16)) :rule refl)
(step t10.t27 (cl (= @p_12 @p_40)) :rule cong :premises (t10.t26 t10.t23))
(step t10.t28 (cl (= @p_13 @p_41)) :rule cong :premises (t10.t27 t10.t22))
(step t10.t29 (cl (= @p_45 @p_43)) :rule cong :premises (t10.t25 t10.t28))
(step t10.t30 (cl @p_44) :rule trans :premises (t10.t5 t10.t29))
(step t10.t31 (cl @p_43) :rule resolution :premises (t10.t4 t10.t30 a3))
(step t10.t32 (cl (! (= @p_41 @p_42) :named @p_236)) :rule symm :premises (t10.t31))
(step t10.t33 (cl (! (= @p_58 @p_42) :named @p_292)) :rule symm :premises (t10.a5))
(step t10.t34 (cl @p_59) :rule symm :premises (t10.t33))
(step t10.t35 (cl (or (! (not (forall ((?v0 S3)) (= ?v0 (f5 f9 ?v0)))) :named @p_52) @p_205)) :rule forall_inst :args (@p_55))
(step t10.t36 (cl @p_52 @p_205) :rule or :premises (t10.t35))
(step t10.t37 (cl (not (! (= (forall ((?v0 S3)) (= (f5 f9 ?v0) ?v0)) (forall ((?v0 S3)) (= ?v0 (f5 f9 ?v0)))) :named @p_48)) (not (forall ((?v0 S3)) (= (f5 f9 ?v0) ?v0))) (forall ((?v0 S3)) (= ?v0 (f5 f9 ?v0)))) :rule equiv_pos2)
(anchor :step t10.t38 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t10.t38.t0 (cl (! (= (= (! (f5 f9 ?v0) :named @p_50) ?v0) (= ?v0 @p_50)) :named @p_51)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_51 2 6))
(step t10.t38 (cl @p_48) :rule bind)
(step t10.t39 (cl (forall ((?v0 S3)) (= ?v0 (f5 f9 ?v0)))) :rule resolution :premises (t10.t37 t10.t38 a30))
(step t10.t40 (cl @p_205) :rule resolution :premises (t10.t36 t10.t39))
(step t10.t41 (cl (= @p_56 @p_55)) :rule symm :premises (t10.t40))
(step t10.t42 (cl (or @p_52 @p_54)) :rule forall_inst :args (@p_33))
(step t10.t43 (cl @p_52 @p_54) :rule or :premises (t10.t42))
(step t10.t44 (cl @p_54) :rule resolution :premises (t10.t43 t10.t39))
(step t10.t45 (cl (= @p_55 (! (f5 f11 @p_53) :named @p_244))) :rule cong :premises (t10.t12 t10.t44))
(step t10.t46 (cl (= @p_56 @p_244)) :rule trans :premises (t10.t41 t10.t45))
(step t10.t47 (cl (= @p_57 (! (f15 f16 @p_244) :named @p_245))) :rule cong :premises (t10.t26 t10.t46))
(step t10.t48 (cl (= @p_39 @p_39)) :rule refl)
(step t10.t49 (cl (= @p_58 (! (f3 @p_245 @p_39) :named @p_246))) :rule cong :premises (t10.t47 t10.t48))
(step t10.t50 (cl (= @p_53 @p_33)) :rule symm :premises (t10.t44))
(step t10.t51 (cl (= @p_244 @p_55)) :rule cong :premises (t10.t12 t10.t50))
(step t10.t52 (cl (= @p_244 @p_56)) :rule trans :premises (t10.t51 t10.t40))
(step t10.t53 (cl (= @p_245 @p_57)) :rule cong :premises (t10.t26 t10.t52))
(step t10.t54 (cl (= @p_57 @p_314)) :rule cong :premises (t10.t26 t10.t41))
(step t10.t55 (cl (= @p_245 @p_314)) :rule trans :premises (t10.t53 t10.t54))
(step t10.t56 (cl (! (= @p_39 @p_32) :named @p_233)) :rule symm :premises (t10.a4))
(step t10.t57 (cl (or @p_52 @p_47)) :rule forall_inst :args (@p_31))
(step t10.t58 (cl @p_52 @p_47) :rule or :premises (t10.t57))
(step t10.t59 (cl @p_47) :rule resolution :premises (t10.t58 t10.t39))
(step t10.t60 (cl (= @p_32 @p_31)) :rule symm :premises (t10.t59))
(step t10.t61 (cl (= @p_39 @p_31)) :rule trans :premises (t10.t56 t10.t60))
(step t10.t62 (cl (= @p_246 @p_315)) :rule cong :premises (t10.t55 t10.t61))
(step t10.t63 (cl (= @p_41 @p_315)) :rule trans :premises (t10.t32 t10.t34 t10.t49 t10.t62))
(step t10.t64 (cl @p_316) :rule trans :premises (t10.t31 t10.t63))
(step t10.t65 (cl @p_378) :rule resolution :premises (t10.t3 t10.t64))
(step t10.t66 (cl (= true @p_316)) :rule symm :premises (t10.t65))
(step t10.t67 (cl (= (! (= @p_316 false) :named @p_377) @p_317)) :rule equiv_simplify)
(step t10.t68 (cl @p_377 @p_370) :rule equiv2 :premises (t10.t67))
(step t10.t69 (cl (! (not @p_370) :named @p_372) @p_316) :rule not_not)
(step t10.t70 (cl @p_377 @p_316) :rule resolution :premises (t10.t68 t10.t69))
(step t10.t71 (cl (! (not (! (= (! (= @p_42 (! (f3 (! (f15 f16 @p_33) :named @p_203) f25) :named @p_318)) :named @p_319) @p_316) :named @p_320)) :named @p_367) @p_319 @p_317) :rule equiv_pos1)
(step t10.t72 (cl @p_319 @p_317 @p_367) :rule reordering :premises (t10.t71))
(step t10.t73 (cl (! (not (! (= @p_319 (! (= @p_42 (! (f3 (! (f15 f16 @p_31) :named @p_293) f25) :named @p_330)) :named @p_331)) :named @p_332)) :named @p_366) (! (not @p_319) :named @p_329) @p_331) :rule equiv_pos2)
(step t10.t74 (cl @p_331 @p_329 @p_366) :rule reordering :premises (t10.t73))
(step t10.t75 (cl (= (! (= @p_331 false) :named @p_363) (! (not @p_331) :named @p_362))) :rule equiv_simplify)
(step t10.t76 (cl (not @p_363) @p_362) :rule equiv1 :premises (t10.t75))
(step t10.t77 (cl (! (= @p_42 @p_42) :named @p_99)) :rule refl)
(step t10.t78 (cl (= @p_293 (! (f15 f16 @p_32) :named @p_115))) :rule cong :premises (t10.t26 t10.t59))
(step t10.t79 (cl (= f25 f25)) :rule refl)
(step t10.t80 (cl (= @p_330 (! (f3 @p_115 f25) :named @p_116))) :rule cong :premises (t10.t78 t10.t79))
(step t10.t81 (cl (= @p_331 (! (= @p_42 @p_116) :named @p_117))) :rule cong :premises (t10.t77 t10.t80))
(step t10.t82 (cl (= (! (= @p_117 false) :named @p_364) (! (not @p_117) :named @p_151))) :rule equiv_simplify)
(step t10.t83 (cl @p_364 (! (not @p_151) :named @p_365)) :rule equiv2 :premises (t10.t82))
(step t10.t84 (cl (not @p_365) @p_117) :rule not_not)
(step t10.t85 (cl @p_364 @p_117) :rule resolution :premises (t10.t83 t10.t84))
(step t10.t86 (cl (not (! (= @p_30 @p_151) :named @p_152)) (not @p_30) @p_151) :rule equiv_pos2)
(step t10.t87 (cl (! (= @p_29 (! (= f1 @p_28) :named @p_153)) :named @p_155)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_155 2 6))
(step t10.t88 (cl (= @p_30 (! (not @p_153) :named @p_154))) :rule cong :premises (t10.t87))
(step t10.t89 (cl (= @p_15 @p_115)) :rule cong :premises (t10.t26 t10.t22))
(step t10.t90 (cl (= @p_28 @p_116)) :rule cong :premises (t10.t89 t10.t79))
(step t10.t91 (cl (= @p_153 @p_117)) :rule cong :premises (t10.t25 t10.t90))
(step t10.t92 (cl (= @p_154 @p_151)) :rule cong :premises (t10.t91))
(step t10.t93 (cl @p_152) :rule trans :premises (t10.t88 t10.t92))
(step t10.t94 (cl @p_151) :rule resolution :premises (t10.t86 t10.t93 a204))
(step t10.t95 (cl @p_364) :rule resolution :premises (t10.t85 t10.t94))
(step t10.t96 (cl @p_363) :rule trans :premises (t10.t81 t10.t95))
(step t10.t97 (cl @p_362) :rule resolution :premises (t10.t76 t10.t96))
(step t10.t98 (cl (not (! (= (! (=> (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))) (! (= (! (= @p_330 @p_42) :named @p_355) @p_319) :named @p_356)) :named @p_357) (! (=> (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))) @p_332) :named @p_354)) :named @p_358)) (not @p_357) @p_354) :rule equiv_pos2)
(step t10.t99 (cl (= (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))) (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))))) :rule refl)
(step t10.t100 (cl (! (= @p_355 @p_331) :named @p_361)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_361 2 6))
(step t10.t101 (cl (= @p_319 @p_319)) :rule refl)
(step t10.t102 (cl (= @p_356 (! (= @p_331 @p_319) :named @p_359))) :rule cong :premises (t10.t100 t10.t101))
(step t10.t103 (cl (! (= @p_359 @p_332) :named @p_360)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_360 1 6))
(step t10.t104 (cl (= @p_356 @p_332)) :rule trans :premises (t10.t102 t10.t103))
(step t10.t105 (cl @p_358) :rule cong :premises (t10.t99 t10.t104))
(step t10.t106 (cl @p_357 (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25))))) :rule implies_neg1)
(anchor :step t10.t107)
(assume t10.t107.a0 (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))))
(step t10.t107.t0 (cl (or (! (not (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25))))) :named @p_353) @p_356)) :rule forall_inst :args (@p_31))
(step t10.t107.t1 (cl @p_353 @p_356) :rule or :premises (t10.t107.t0))
(step t10.t107.t2 (cl (not (! (= (forall ((?v0 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) f12) f1) (= (f3 (f15 f16 ?v0) f12) f1))) (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25))))) :named @p_333)) (not (forall ((?v0 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) f12) f1) (= (f3 (f15 f16 ?v0) f12) f1)))) (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25))))) :rule equiv_pos2)
(anchor :step t10.t107.t3 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t10.t107.t3.t0 (cl (! (= (! (= (! (f3 (! (f15 f16 (f5 f10 ?v0)) :named @p_176) f12) :named @p_346) f1) :named @p_350) (! (= f1 @p_346) :named @p_347)) :named @p_352)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_352 2 6))
(step t10.t107.t3.t1 (cl (! (= (! (= (! (f3 (! (f15 f16 ?v0) :named @p_124) f12) :named @p_344) f1) :named @p_349) (! (= f1 @p_344) :named @p_345)) :named @p_351)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_351 2 6))
(step t10.t107.t3.t2 (cl (= (= @p_350 @p_349) (! (= @p_347 @p_345) :named @p_348))) :rule cong :premises (t10.t107.t3.t0 t10.t107.t3.t1))
(step t10.t107.t3 (cl (= (forall ((?v0 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) f12) f1) (= (f3 (f15 f16 ?v0) f12) f1))) (forall ((?v0 S3)) (= (= f1 (f3 (f15 f16 (f5 f10 ?v0)) f12)) (= f1 (f3 (f15 f16 ?v0) f12)))))) :rule bind)
(anchor :step t10.t107.t4 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t10.t107.t4.t0 (cl (= @p_176 @p_176)) :rule refl)
(step t10.t107.t4.t1 (cl (= @p_346 (! (f3 @p_176 f25) :named @p_334))) :rule cong :premises (t10.t107.t4.t0 a104))
(step t10.t107.t4.t2 (cl (= @p_347 (! (= @p_42 @p_334) :named @p_335))) :rule cong :premises (t10.t25 t10.t107.t4.t1))
(step t10.t107.t4.t3 (cl (= @p_124 @p_124)) :rule refl)
(step t10.t107.t4.t4 (cl (= @p_344 (! (f3 @p_124 f25) :named @p_336))) :rule cong :premises (t10.t107.t4.t3 a104))
(step t10.t107.t4.t5 (cl (= @p_345 (! (= @p_42 @p_336) :named @p_339))) :rule cong :premises (t10.t25 t10.t107.t4.t4))
(step t10.t107.t4.t6 (cl (= @p_348 (! (= @p_335 @p_339) :named @p_340))) :rule cong :premises (t10.t107.t4.t2 t10.t107.t4.t5))
(step t10.t107.t4 (cl (= (forall ((?v0 S3)) (= (= f1 (f3 (f15 f16 (f5 f10 ?v0)) f12)) (= f1 (f3 (f15 f16 ?v0) f12)))) (forall ((?v0 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) f25)))))) :rule bind)
(anchor :step t10.t107.t5 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t10.t107.t5.t0 (cl (! (= @p_340 (! (= @p_339 @p_335) :named @p_341)) :named @p_343)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_343 1 5))
(step t10.t107.t5.t1 (cl (! (= @p_339 (! (= @p_336 @p_42) :named @p_337)) :named @p_342)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_342 2 6))
(step t10.t107.t5.t2 (cl (= @p_335 @p_335)) :rule refl)
(step t10.t107.t5.t3 (cl (= @p_341 (! (= @p_337 @p_335) :named @p_338))) :rule cong :premises (t10.t107.t5.t1 t10.t107.t5.t2))
(step t10.t107.t5.t4 (cl (= @p_340 @p_338)) :rule trans :premises (t10.t107.t5.t0 t10.t107.t5.t3))
(step t10.t107.t5 (cl (= (forall ((?v0 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) f25)))) (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))))) :rule bind)
(step t10.t107.t6 (cl (= (forall ((?v0 S3)) (= (= f1 (f3 (f15 f16 (f5 f10 ?v0)) f12)) (= f1 (f3 (f15 f16 ?v0) f12)))) (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))))) :rule trans :premises (t10.t107.t4 t10.t107.t5))
(step t10.t107.t7 (cl @p_333) :rule trans :premises (t10.t107.t3 t10.t107.t6))
(step t10.t107.t8 (cl (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25))))) :rule resolution :premises (t10.t107.t2 t10.t107.t7 a70))
(step t10.t107.t9 (cl @p_356) :rule resolution :premises (t10.t107.t1 t10.t107.t8))
(step t10.t107 (cl @p_353 @p_356) :rule subproof :discharge (t10.t107.a0))
(step t10.t108 (cl @p_357 @p_356) :rule resolution :premises (t10.t106 t10.t107))
(step t10.t109 (cl @p_357 (not @p_356)) :rule implies_neg2)
(step t10.t110 (cl @p_357 @p_357) :rule resolution :premises (t10.t108 t10.t109))
(step t10.t111 (cl @p_357) :rule contraction :premises (t10.t110))
(step t10.t112 (cl @p_354) :rule resolution :premises (t10.t98 t10.t105 t10.t111))
(step t10.t113 (cl @p_353 @p_332) :rule implies :premises (t10.t112))
(step t10.t114 (cl (not @p_333) (not (forall ((?v0 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) f12) f1) (= (f3 (f15 f16 ?v0) f12) f1)))) (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25))))) :rule equiv_pos2)
(anchor :step t10.t115 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t10.t115.t0 (cl @p_352) :rule hole :args ("TRUST_THEORY_REWRITE" @p_352 2 6))
(step t10.t115.t1 (cl @p_351) :rule hole :args ("TRUST_THEORY_REWRITE" @p_351 2 6))
(step t10.t115.t2 (cl (= (= @p_350 @p_349) @p_348)) :rule cong :premises (t10.t115.t0 t10.t115.t1))
(step t10.t115 (cl (= (forall ((?v0 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) f12) f1) (= (f3 (f15 f16 ?v0) f12) f1))) (forall ((?v0 S3)) (= (= f1 (f3 (f15 f16 (f5 f10 ?v0)) f12)) (= f1 (f3 (f15 f16 ?v0) f12)))))) :rule bind)
(anchor :step t10.t116 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t10.t116.t0 (cl (= @p_176 @p_176)) :rule refl)
(step t10.t116.t1 (cl (= @p_346 @p_334)) :rule cong :premises (t10.t116.t0 a104))
(step t10.t116.t2 (cl (= @p_347 @p_335)) :rule cong :premises (t10.t25 t10.t116.t1))
(step t10.t116.t3 (cl (= @p_124 @p_124)) :rule refl)
(step t10.t116.t4 (cl (= @p_344 @p_336)) :rule cong :premises (t10.t116.t3 a104))
(step t10.t116.t5 (cl (= @p_345 @p_339)) :rule cong :premises (t10.t25 t10.t116.t4))
(step t10.t116.t6 (cl (= @p_348 @p_340)) :rule cong :premises (t10.t116.t2 t10.t116.t5))
(step t10.t116 (cl (= (forall ((?v0 S3)) (= (= f1 (f3 (f15 f16 (f5 f10 ?v0)) f12)) (= f1 (f3 (f15 f16 ?v0) f12)))) (forall ((?v0 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) f25)))))) :rule bind)
(anchor :step t10.t117 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t10.t117.t0 (cl @p_343) :rule hole :args ("TRUST_THEORY_REWRITE" @p_343 1 5))
(step t10.t117.t1 (cl @p_342) :rule hole :args ("TRUST_THEORY_REWRITE" @p_342 2 6))
(step t10.t117.t2 (cl (= @p_335 @p_335)) :rule refl)
(step t10.t117.t3 (cl (= @p_341 @p_338)) :rule cong :premises (t10.t117.t1 t10.t117.t2))
(step t10.t117.t4 (cl (= @p_340 @p_338)) :rule trans :premises (t10.t117.t0 t10.t117.t3))
(step t10.t117 (cl (= (forall ((?v0 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) f25)))) (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))))) :rule bind)
(step t10.t118 (cl (= (forall ((?v0 S3)) (= (= f1 (f3 (f15 f16 (f5 f10 ?v0)) f12)) (= f1 (f3 (f15 f16 ?v0) f12)))) (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))))) :rule trans :premises (t10.t116 t10.t117))
(step t10.t119 (cl @p_333) :rule trans :premises (t10.t115 t10.t118))
(step t10.t120 (cl (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25))))) :rule resolution :premises (t10.t114 t10.t119 a70))
(step t10.t121 (cl @p_332) :rule resolution :premises (t10.t113 t10.t120))
(step t10.t122 (cl @p_329) :rule resolution :premises (t10.t74 t10.t97 t10.t121))
(step t10.t123 (cl (not (! (= (! (=> (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))) (! (= (! (= @p_318 @p_42) :named @p_323) (! (= @p_315 @p_42) :named @p_322)) :named @p_324)) :named @p_325) (! (=> (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))) @p_320) :named @p_321)) :named @p_326)) (not @p_325) @p_321) :rule equiv_pos2)
(step t10.t124 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule refl)
(step t10.t125 (cl (! (= @p_323 @p_319) :named @p_328)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_328 2 6))
(step t10.t126 (cl (! (= @p_322 @p_316) :named @p_327)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_327 2 6))
(step t10.t127 (cl (= @p_324 @p_320)) :rule cong :premises (t10.t125 t10.t126))
(step t10.t128 (cl @p_326) :rule cong :premises (t10.t124 t10.t127))
(step t10.t129 (cl @p_325 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule implies_neg1)
(anchor :step t10.t130)
(assume t10.t130.a0 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))
(step t10.t130.t0 (cl (or (! (not (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :named @p_142) @p_324)) :rule forall_inst :args (@p_33 f25))
(step t10.t130.t1 (cl @p_142 @p_324) :rule or :premises (t10.t130.t0))
(step t10.t130.t2 (cl (not (! (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :named @p_119)) (not (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1)))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule equiv_pos2)
(anchor :step t10.t130.t3 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t10.t130.t3.t0 (cl (! (= (! (= (! (= (! (f3 (f15 f16 (f5 f11 ?v0)) (! (f5 f11 ?v1) :named @p_121)) :named @p_122) f1) :named @p_136) (! (= (! (f3 @p_124 ?v1) :named @p_125) f1) :named @p_135)) :named @p_137) (! (= @p_135 @p_136) :named @p_138)) :named @p_141)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_141 1 5))
(step t10.t130.t3.t1 (cl (! (= @p_135 (! (= f1 @p_125) :named @p_133)) :named @p_140)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_140 2 6))
(step t10.t130.t3.t2 (cl (! (= @p_136 (! (= f1 @p_122) :named @p_132)) :named @p_139)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_139 2 6))
(step t10.t130.t3.t3 (cl (= @p_138 (! (= @p_133 @p_132) :named @p_134))) :rule cong :premises (t10.t130.t3.t1 t10.t130.t3.t2))
(step t10.t130.t3.t4 (cl (= @p_137 @p_134)) :rule trans :premises (t10.t130.t3.t0 t10.t130.t3.t3))
(step t10.t130.t3 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1))) (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))))) :rule bind)
(anchor :step t10.t130.t4 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t10.t130.t4.t0 (cl (= @p_125 @p_125)) :rule refl)
(step t10.t130.t4.t1 (cl (= @p_133 (! (= @p_42 @p_125) :named @p_128))) :rule cong :premises (t10.t25 t10.t130.t4.t0))
(step t10.t130.t4.t2 (cl (= @p_122 @p_122)) :rule refl)
(step t10.t130.t4.t3 (cl (= @p_132 (! (= @p_42 @p_122) :named @p_127))) :rule cong :premises (t10.t25 t10.t130.t4.t2))
(step t10.t130.t4.t4 (cl (= @p_134 (! (= @p_128 @p_127) :named @p_129))) :rule cong :premises (t10.t130.t4.t1 t10.t130.t4.t3))
(step t10.t130.t4 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))))) :rule bind)
(anchor :step t10.t130.t5 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t10.t130.t5.t0 (cl (! (= @p_128 (! (= @p_125 @p_42) :named @p_126)) :named @p_131)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_131 2 6))
(step t10.t130.t5.t1 (cl (! (= @p_127 (! (= @p_122 @p_42) :named @p_123)) :named @p_130)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_130 2 6))
(step t10.t130.t5.t2 (cl (= @p_129 (= @p_126 @p_123))) :rule cong :premises (t10.t130.t5.t0 t10.t130.t5.t1))
(step t10.t130.t5 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule bind)
(step t10.t130.t6 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule trans :premises (t10.t130.t4 t10.t130.t5))
(step t10.t130.t7 (cl @p_119) :rule trans :premises (t10.t130.t3 t10.t130.t6))
(step t10.t130.t8 (cl (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule resolution :premises (t10.t130.t2 t10.t130.t7 a41))
(step t10.t130.t9 (cl @p_324) :rule resolution :premises (t10.t130.t1 t10.t130.t8))
(step t10.t130 (cl @p_142 @p_324) :rule subproof :discharge (t10.t130.a0))
(step t10.t131 (cl @p_325 @p_324) :rule resolution :premises (t10.t129 t10.t130))
(step t10.t132 (cl @p_325 (not @p_324)) :rule implies_neg2)
(step t10.t133 (cl @p_325 @p_325) :rule resolution :premises (t10.t131 t10.t132))
(step t10.t134 (cl @p_325) :rule contraction :premises (t10.t133))
(step t10.t135 (cl @p_321) :rule resolution :premises (t10.t123 t10.t128 t10.t134))
(step t10.t136 (cl @p_142 @p_320) :rule implies :premises (t10.t135))
(step t10.t137 (cl (not @p_119) (not (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1)))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule equiv_pos2)
(anchor :step t10.t138 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t10.t138.t0 (cl @p_141) :rule hole :args ("TRUST_THEORY_REWRITE" @p_141 1 5))
(step t10.t138.t1 (cl @p_140) :rule hole :args ("TRUST_THEORY_REWRITE" @p_140 2 6))
(step t10.t138.t2 (cl @p_139) :rule hole :args ("TRUST_THEORY_REWRITE" @p_139 2 6))
(step t10.t138.t3 (cl (= @p_138 @p_134)) :rule cong :premises (t10.t138.t1 t10.t138.t2))
(step t10.t138.t4 (cl (= @p_137 @p_134)) :rule trans :premises (t10.t138.t0 t10.t138.t3))
(step t10.t138 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1))) (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))))) :rule bind)
(anchor :step t10.t139 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t10.t139.t0 (cl (= @p_125 @p_125)) :rule refl)
(step t10.t139.t1 (cl (= @p_133 @p_128)) :rule cong :premises (t10.t25 t10.t139.t0))
(step t10.t139.t2 (cl (= @p_122 @p_122)) :rule refl)
(step t10.t139.t3 (cl (= @p_132 @p_127)) :rule cong :premises (t10.t25 t10.t139.t2))
(step t10.t139.t4 (cl (= @p_134 @p_129)) :rule cong :premises (t10.t139.t1 t10.t139.t3))
(step t10.t139 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))))) :rule bind)
(anchor :step t10.t140 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t10.t140.t0 (cl @p_131) :rule hole :args ("TRUST_THEORY_REWRITE" @p_131 2 6))
(step t10.t140.t1 (cl @p_130) :rule hole :args ("TRUST_THEORY_REWRITE" @p_130 2 6))
(step t10.t140.t2 (cl (= @p_129 (= @p_126 @p_123))) :rule cong :premises (t10.t140.t0 t10.t140.t1))
(step t10.t140 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule bind)
(step t10.t141 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule trans :premises (t10.t139 t10.t140))
(step t10.t142 (cl @p_119) :rule trans :premises (t10.t138 t10.t141))
(step t10.t143 (cl (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule resolution :premises (t10.t137 t10.t142 a41))
(step t10.t144 (cl @p_320) :rule resolution :premises (t10.t136 t10.t143))
(step t10.t145 (cl @p_317) :rule resolution :premises (t10.t72 t10.t122 t10.t144))
(step t10.t146 (cl @p_377) :rule resolution :premises (t10.t70 t10.t145))
(step t10.t147 (cl @p_376) :rule trans :premises (t10.t66 t10.t146))
(step t10.t148 (cl false) :rule resolution :premises (t10.t0 t10.t1 t10.t147))
(step t10 (cl @p_370 @p_288 @p_197 @p_159 @p_368 @p_287 @p_158 false) :rule subproof :discharge (t10.a0 t10.a1 t10.a2 t10.a3 t10.a4 t10.a5 t10.a6))
(step t11 (cl (! (not @p_373) :named @p_374) @p_317) :rule and_pos :args (0))
(step t12 (cl @p_374 @p_205) :rule and_pos :args (1))
(step t13 (cl @p_374 @p_54) :rule and_pos :args (2))
(step t14 (cl @p_374 @p_47) :rule and_pos :args (3))
(step t15 (cl @p_374 @p_206) :rule and_pos :args (4))
(step t16 (cl @p_374 @p_59) :rule and_pos :args (5))
(step t17 (cl @p_374 @p_43) :rule and_pos :args (6))
(step t18 (cl false @p_374 @p_374 @p_374 @p_374 @p_374 @p_374 @p_374) :rule resolution :premises (t10 t11 t12 t13 t14 t15 t16 t17))
(step t19 (cl @p_374 @p_374 @p_374 @p_374 @p_374 @p_374 @p_374 false) :rule reordering :premises (t18))
(step t20 (cl @p_374 false) :rule contraction :premises (t19))
(step t21 (cl @p_375 false) :rule resolution :premises (t9 t20))
(step t22 (cl @p_375 @p_1) :rule implies_neg2)
(step t23 (cl @p_375 @p_375) :rule resolution :premises (t21 t22))
(step t24 (cl @p_375) :rule contraction :premises (t23))
(step t25 (cl (= @p_375 @p_374)) :rule implies_simplify)
(step t26 (cl (not @p_375) @p_374) :rule equiv1 :premises (t25))
(step t27 (cl @p_374) :rule resolution :premises (t24 t26))
(step t28 (cl @p_370 @p_288 @p_197 @p_159 @p_368 @p_287 @p_158) :rule not_and :premises (t27))
(step t29 (cl @p_371 @p_372) :rule or_neg :args (0))
(step t30 (cl @p_371 (not @p_288)) :rule or_neg :args (1))
(step t31 (cl @p_371 (not @p_197)) :rule or_neg :args (2))
(step t32 (cl @p_371 (not @p_159)) :rule or_neg :args (3))
(step t33 (cl @p_371 (not @p_368)) :rule or_neg :args (4))
(step t34 (cl @p_371 (not @p_287)) :rule or_neg :args (5))
(step t35 (cl @p_371 (not @p_158)) :rule or_neg :args (6))
(step t36 (cl @p_371 @p_371 @p_371 @p_371 @p_371 @p_371 @p_371) :rule resolution :premises (t28 t29 t30 t31 t32 t33 t34 t35))
(step t37 (cl @p_371) :rule contraction :premises (t36))
(step t38 (cl @p_369) :rule resolution :premises (t0 t8 t37))
(step t39 (cl @p_316 @p_288 @p_197 @p_159 @p_368 @p_287 @p_158) :rule or :premises (t38))
(step t40 (cl @p_158 @p_159 @p_197 @p_287 @p_368 @p_288 @p_316) :rule reordering :premises (t39))
(step t41 (cl @p_367 @p_319 @p_317) :rule equiv_pos1)
(step t42 (cl @p_319 @p_317 @p_367) :rule reordering :premises (t41))
(step t43 (cl @p_366 @p_329 @p_331) :rule equiv_pos2)
(step t44 (cl @p_331 @p_329 @p_366) :rule reordering :premises (t43))
(step t45 (cl (= @p_363 @p_362)) :rule equiv_simplify)
(step t46 (cl (not @p_363) @p_362) :rule equiv1 :premises (t45))
(step t47 (cl @p_99) :rule refl)
(step t48 (cl (= f16 f16)) :rule refl)
(step t49 (cl (or @p_52 @p_47)) :rule forall_inst :args (@p_31))
(step t50 (cl @p_52 @p_47) :rule or :premises (t49))
(step t51 (cl (not @p_48) (not (forall ((?v0 S3)) (= (f5 f9 ?v0) ?v0))) (forall ((?v0 S3)) (= ?v0 (f5 f9 ?v0)))) :rule equiv_pos2)
(anchor :step t52 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t52.t0 (cl @p_51) :rule hole :args ("TRUST_THEORY_REWRITE" @p_51 2 6))
(step t52 (cl @p_48) :rule bind)
(step t53 (cl (forall ((?v0 S3)) (= ?v0 (f5 f9 ?v0)))) :rule resolution :premises (t51 t52 a30))
(step t54 (cl @p_47) :rule resolution :premises (t50 t53))
(step t55 (cl (= @p_293 @p_115)) :rule cong :premises (t48 t54))
(step t56 (cl (= f25 f25)) :rule refl)
(step t57 (cl (= @p_330 @p_116)) :rule cong :premises (t55 t56))
(step t58 (cl (= @p_331 @p_117)) :rule cong :premises (t47 t57))
(step t59 (cl (= @p_364 @p_151)) :rule equiv_simplify)
(step t60 (cl @p_364 @p_365) :rule equiv2 :premises (t59))
(step t61 (cl (not @p_365) @p_117) :rule not_not)
(step t62 (cl @p_364 @p_117) :rule resolution :premises (t60 t61))
(step t63 (cl (not @p_152) (not @p_30) @p_151) :rule equiv_pos2)
(step t64 (cl @p_155) :rule hole :args ("TRUST_THEORY_REWRITE" @p_155 2 6))
(step t65 (cl (= @p_30 @p_154)) :rule cong :premises (t64))
(step t66 (cl (= f1 @p_10)) :rule symm :premises (a2))
(step t67 (cl (= f4 f4)) :rule refl)
(step t68 (cl (= f7 f7)) :rule refl)
(step t69 (cl (= f8 f8)) :rule refl)
(step t70 (cl (= f9 f9)) :rule refl)
(step t71 (cl (= f10 f10)) :rule refl)
(step t72 (cl (= f11 f11)) :rule refl)
(step t73 (cl (= @p_2 @p_31)) :rule cong :premises (t72 a104))
(step t74 (cl (= @p_3 @p_33)) :rule cong :premises (t71 t73))
(step t75 (cl (= @p_4 @p_34)) :rule cong :premises (t71 t74))
(step t76 (cl (= @p_5 @p_35)) :rule cong :premises (t70 t75))
(step t77 (cl (= @p_6 @p_36)) :rule cong :premises (t69 t76))
(step t78 (cl (= f13 f13)) :rule refl)
(step t79 (cl (= @p_7 @p_37)) :rule cong :premises (t77 t78))
(step t80 (cl (= @p_8 @p_38)) :rule cong :premises (t68 t79))
(step t81 (cl (= @p_18 @p_32)) :rule cong :premises (t70 t73))
(step t82 (cl (= f14 @p_32)) :rule trans :premises (a80 t81))
(step t83 (cl (= @p_9 @p_39)) :rule cong :premises (t80 t82))
(step t84 (cl (= @p_10 @p_42)) :rule cong :premises (t67 t83))
(step t85 (cl (= f1 @p_42)) :rule trans :premises (t66 t84))
(step t86 (cl (= @p_15 @p_115)) :rule cong :premises (t48 t82))
(step t87 (cl (= @p_28 @p_116)) :rule cong :premises (t86 t56))
(step t88 (cl (= @p_153 @p_117)) :rule cong :premises (t85 t87))
(step t89 (cl (= @p_154 @p_151)) :rule cong :premises (t88))
(step t90 (cl @p_152) :rule trans :premises (t65 t89))
(step t91 (cl @p_151) :rule resolution :premises (t63 t90 a204))
(step t92 (cl @p_364) :rule resolution :premises (t62 t91))
(step t93 (cl @p_363) :rule trans :premises (t58 t92))
(step t94 (cl @p_362) :rule resolution :premises (t46 t93))
(step t95 (cl (not @p_358) (not @p_357) @p_354) :rule equiv_pos2)
(step t96 (cl (= (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))) (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))))) :rule refl)
(step t97 (cl @p_361) :rule hole :args ("TRUST_THEORY_REWRITE" @p_361 2 6))
(step t98 (cl (= @p_319 @p_319)) :rule refl)
(step t99 (cl (= @p_356 @p_359)) :rule cong :premises (t97 t98))
(step t100 (cl @p_360) :rule hole :args ("TRUST_THEORY_REWRITE" @p_360 1 6))
(step t101 (cl (= @p_356 @p_332)) :rule trans :premises (t99 t100))
(step t102 (cl @p_358) :rule cong :premises (t96 t101))
(step t103 (cl @p_357 (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25))))) :rule implies_neg1)
(anchor :step t104)
(assume t104.a0 (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))))
(step t104.t0 (cl (or @p_353 @p_356)) :rule forall_inst :args (@p_31))
(step t104.t1 (cl @p_353 @p_356) :rule or :premises (t104.t0))
(step t104.t2 (cl (not @p_333) (not (forall ((?v0 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) f12) f1) (= (f3 (f15 f16 ?v0) f12) f1)))) (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25))))) :rule equiv_pos2)
(anchor :step t104.t3 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t104.t3.t0 (cl @p_352) :rule hole :args ("TRUST_THEORY_REWRITE" @p_352 2 6))
(step t104.t3.t1 (cl @p_351) :rule hole :args ("TRUST_THEORY_REWRITE" @p_351 2 6))
(step t104.t3.t2 (cl (= (= @p_350 @p_349) @p_348)) :rule cong :premises (t104.t3.t0 t104.t3.t1))
(step t104.t3 (cl (= (forall ((?v0 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) f12) f1) (= (f3 (f15 f16 ?v0) f12) f1))) (forall ((?v0 S3)) (= (= f1 (f3 (f15 f16 (f5 f10 ?v0)) f12)) (= f1 (f3 (f15 f16 ?v0) f12)))))) :rule bind)
(anchor :step t104.t4 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t104.t4.t0 (cl (= @p_176 @p_176)) :rule refl)
(step t104.t4.t1 (cl (= @p_346 @p_334)) :rule cong :premises (t104.t4.t0 a104))
(step t104.t4.t2 (cl (= @p_347 @p_335)) :rule cong :premises (t85 t104.t4.t1))
(step t104.t4.t3 (cl (= @p_124 @p_124)) :rule refl)
(step t104.t4.t4 (cl (= @p_344 @p_336)) :rule cong :premises (t104.t4.t3 a104))
(step t104.t4.t5 (cl (= @p_345 @p_339)) :rule cong :premises (t85 t104.t4.t4))
(step t104.t4.t6 (cl (= @p_348 @p_340)) :rule cong :premises (t104.t4.t2 t104.t4.t5))
(step t104.t4 (cl (= (forall ((?v0 S3)) (= (= f1 (f3 (f15 f16 (f5 f10 ?v0)) f12)) (= f1 (f3 (f15 f16 ?v0) f12)))) (forall ((?v0 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) f25)))))) :rule bind)
(anchor :step t104.t5 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t104.t5.t0 (cl @p_343) :rule hole :args ("TRUST_THEORY_REWRITE" @p_343 1 5))
(step t104.t5.t1 (cl @p_342) :rule hole :args ("TRUST_THEORY_REWRITE" @p_342 2 6))
(step t104.t5.t2 (cl (= @p_335 @p_335)) :rule refl)
(step t104.t5.t3 (cl (= @p_341 @p_338)) :rule cong :premises (t104.t5.t1 t104.t5.t2))
(step t104.t5.t4 (cl (= @p_340 @p_338)) :rule trans :premises (t104.t5.t0 t104.t5.t3))
(step t104.t5 (cl (= (forall ((?v0 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) f25)))) (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))))) :rule bind)
(step t104.t6 (cl (= (forall ((?v0 S3)) (= (= f1 (f3 (f15 f16 (f5 f10 ?v0)) f12)) (= f1 (f3 (f15 f16 ?v0) f12)))) (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))))) :rule trans :premises (t104.t4 t104.t5))
(step t104.t7 (cl @p_333) :rule trans :premises (t104.t3 t104.t6))
(step t104.t8 (cl (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25))))) :rule resolution :premises (t104.t2 t104.t7 a70))
(step t104.t9 (cl @p_356) :rule resolution :premises (t104.t1 t104.t8))
(step t104 (cl @p_353 @p_356) :rule subproof :discharge (t104.a0))
(step t105 (cl @p_357 @p_356) :rule resolution :premises (t103 t104))
(step t106 (cl @p_357 (not @p_356)) :rule implies_neg2)
(step t107 (cl @p_357 @p_357) :rule resolution :premises (t105 t106))
(step t108 (cl @p_357) :rule contraction :premises (t107))
(step t109 (cl @p_354) :rule resolution :premises (t95 t102 t108))
(step t110 (cl @p_353 @p_332) :rule implies :premises (t109))
(step t111 (cl (not @p_333) (not (forall ((?v0 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) f12) f1) (= (f3 (f15 f16 ?v0) f12) f1)))) (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25))))) :rule equiv_pos2)
(anchor :step t112 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t112.t0 (cl @p_352) :rule hole :args ("TRUST_THEORY_REWRITE" @p_352 2 6))
(step t112.t1 (cl @p_351) :rule hole :args ("TRUST_THEORY_REWRITE" @p_351 2 6))
(step t112.t2 (cl (= (= @p_350 @p_349) @p_348)) :rule cong :premises (t112.t0 t112.t1))
(step t112 (cl (= (forall ((?v0 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) f12) f1) (= (f3 (f15 f16 ?v0) f12) f1))) (forall ((?v0 S3)) (= (= f1 (f3 (f15 f16 (f5 f10 ?v0)) f12)) (= f1 (f3 (f15 f16 ?v0) f12)))))) :rule bind)
(anchor :step t113 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t113.t0 (cl (= @p_176 @p_176)) :rule refl)
(step t113.t1 (cl (= @p_346 @p_334)) :rule cong :premises (t113.t0 a104))
(step t113.t2 (cl (= @p_347 @p_335)) :rule cong :premises (t85 t113.t1))
(step t113.t3 (cl (= @p_124 @p_124)) :rule refl)
(step t113.t4 (cl (= @p_344 @p_336)) :rule cong :premises (t113.t3 a104))
(step t113.t5 (cl (= @p_345 @p_339)) :rule cong :premises (t85 t113.t4))
(step t113.t6 (cl (= @p_348 @p_340)) :rule cong :premises (t113.t2 t113.t5))
(step t113 (cl (= (forall ((?v0 S3)) (= (= f1 (f3 (f15 f16 (f5 f10 ?v0)) f12)) (= f1 (f3 (f15 f16 ?v0) f12)))) (forall ((?v0 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) f25)))))) :rule bind)
(anchor :step t114 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t114.t0 (cl @p_343) :rule hole :args ("TRUST_THEORY_REWRITE" @p_343 1 5))
(step t114.t1 (cl @p_342) :rule hole :args ("TRUST_THEORY_REWRITE" @p_342 2 6))
(step t114.t2 (cl (= @p_335 @p_335)) :rule refl)
(step t114.t3 (cl (= @p_341 @p_338)) :rule cong :premises (t114.t1 t114.t2))
(step t114.t4 (cl (= @p_340 @p_338)) :rule trans :premises (t114.t0 t114.t3))
(step t114 (cl (= (forall ((?v0 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) f25)))) (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))))) :rule bind)
(step t115 (cl (= (forall ((?v0 S3)) (= (= f1 (f3 (f15 f16 (f5 f10 ?v0)) f12)) (= f1 (f3 (f15 f16 ?v0) f12)))) (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25)))))) :rule trans :premises (t113 t114))
(step t116 (cl @p_333) :rule trans :premises (t112 t115))
(step t117 (cl (forall ((?v0 S3)) (= (= (f3 (f15 f16 ?v0) f25) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) f25))))) :rule resolution :premises (t111 t116 a70))
(step t118 (cl @p_332) :rule resolution :premises (t110 t117))
(step t119 (cl @p_329) :rule resolution :premises (t44 t94 t118))
(step t120 (cl (not @p_326) (not @p_325) @p_321) :rule equiv_pos2)
(step t121 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule refl)
(step t122 (cl @p_328) :rule hole :args ("TRUST_THEORY_REWRITE" @p_328 2 6))
(step t123 (cl @p_327) :rule hole :args ("TRUST_THEORY_REWRITE" @p_327 2 6))
(step t124 (cl (= @p_324 @p_320)) :rule cong :premises (t122 t123))
(step t125 (cl @p_326) :rule cong :premises (t121 t124))
(step t126 (cl @p_325 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule implies_neg1)
(anchor :step t127)
(assume t127.a0 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))
(step t127.t0 (cl (or @p_142 @p_324)) :rule forall_inst :args (@p_33 f25))
(step t127.t1 (cl @p_142 @p_324) :rule or :premises (t127.t0))
(step t127.t2 (cl (not @p_119) (not (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1)))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule equiv_pos2)
(anchor :step t127.t3 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t127.t3.t0 (cl @p_141) :rule hole :args ("TRUST_THEORY_REWRITE" @p_141 1 5))
(step t127.t3.t1 (cl @p_140) :rule hole :args ("TRUST_THEORY_REWRITE" @p_140 2 6))
(step t127.t3.t2 (cl @p_139) :rule hole :args ("TRUST_THEORY_REWRITE" @p_139 2 6))
(step t127.t3.t3 (cl (= @p_138 @p_134)) :rule cong :premises (t127.t3.t1 t127.t3.t2))
(step t127.t3.t4 (cl (= @p_137 @p_134)) :rule trans :premises (t127.t3.t0 t127.t3.t3))
(step t127.t3 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1))) (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))))) :rule bind)
(anchor :step t127.t4 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t127.t4.t0 (cl (= @p_125 @p_125)) :rule refl)
(step t127.t4.t1 (cl (= @p_133 @p_128)) :rule cong :premises (t85 t127.t4.t0))
(step t127.t4.t2 (cl (= @p_122 @p_122)) :rule refl)
(step t127.t4.t3 (cl (= @p_132 @p_127)) :rule cong :premises (t85 t127.t4.t2))
(step t127.t4.t4 (cl (= @p_134 @p_129)) :rule cong :premises (t127.t4.t1 t127.t4.t3))
(step t127.t4 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))))) :rule bind)
(anchor :step t127.t5 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t127.t5.t0 (cl @p_131) :rule hole :args ("TRUST_THEORY_REWRITE" @p_131 2 6))
(step t127.t5.t1 (cl @p_130) :rule hole :args ("TRUST_THEORY_REWRITE" @p_130 2 6))
(step t127.t5.t2 (cl (= @p_129 (= @p_126 @p_123))) :rule cong :premises (t127.t5.t0 t127.t5.t1))
(step t127.t5 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule bind)
(step t127.t6 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule trans :premises (t127.t4 t127.t5))
(step t127.t7 (cl @p_119) :rule trans :premises (t127.t3 t127.t6))
(step t127.t8 (cl (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule resolution :premises (t127.t2 t127.t7 a41))
(step t127.t9 (cl @p_324) :rule resolution :premises (t127.t1 t127.t8))
(step t127 (cl @p_142 @p_324) :rule subproof :discharge (t127.a0))
(step t128 (cl @p_325 @p_324) :rule resolution :premises (t126 t127))
(step t129 (cl @p_325 (not @p_324)) :rule implies_neg2)
(step t130 (cl @p_325 @p_325) :rule resolution :premises (t128 t129))
(step t131 (cl @p_325) :rule contraction :premises (t130))
(step t132 (cl @p_321) :rule resolution :premises (t120 t125 t131))
(step t133 (cl @p_142 @p_320) :rule implies :premises (t132))
(step t134 (cl (not @p_119) (not (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1)))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule equiv_pos2)
(anchor :step t135 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t135.t0 (cl @p_141) :rule hole :args ("TRUST_THEORY_REWRITE" @p_141 1 5))
(step t135.t1 (cl @p_140) :rule hole :args ("TRUST_THEORY_REWRITE" @p_140 2 6))
(step t135.t2 (cl @p_139) :rule hole :args ("TRUST_THEORY_REWRITE" @p_139 2 6))
(step t135.t3 (cl (= @p_138 @p_134)) :rule cong :premises (t135.t1 t135.t2))
(step t135.t4 (cl (= @p_137 @p_134)) :rule trans :premises (t135.t0 t135.t3))
(step t135 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1))) (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))))) :rule bind)
(anchor :step t136 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t136.t0 (cl (= @p_125 @p_125)) :rule refl)
(step t136.t1 (cl (= @p_133 @p_128)) :rule cong :premises (t85 t136.t0))
(step t136.t2 (cl (= @p_122 @p_122)) :rule refl)
(step t136.t3 (cl (= @p_132 @p_127)) :rule cong :premises (t85 t136.t2))
(step t136.t4 (cl (= @p_134 @p_129)) :rule cong :premises (t136.t1 t136.t3))
(step t136 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))))) :rule bind)
(anchor :step t137 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t137.t0 (cl @p_131) :rule hole :args ("TRUST_THEORY_REWRITE" @p_131 2 6))
(step t137.t1 (cl @p_130) :rule hole :args ("TRUST_THEORY_REWRITE" @p_130 2 6))
(step t137.t2 (cl (= @p_129 (= @p_126 @p_123))) :rule cong :premises (t137.t0 t137.t1))
(step t137 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule bind)
(step t138 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule trans :premises (t136 t137))
(step t139 (cl @p_119) :rule trans :premises (t135 t138))
(step t140 (cl (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule resolution :premises (t134 t139 a41))
(step t141 (cl @p_320) :rule resolution :premises (t133 t140))
(step t142 (cl @p_317) :rule resolution :premises (t42 t119 t141))
(step t143 (cl (! (not (! (or @p_158 (! (not (! (= @p_42 (! (f3 @p_115 @p_39) :named @p_207)) :named @p_208)) :named @p_209) @p_206) :named @p_210)) :named @p_313) @p_158 @p_209 @p_206) :rule or_pos)
(step t144 (cl @p_158 @p_209 @p_206 @p_313) :rule reordering :premises (t143))
(step t145 (cl (not @p_44) (not @p_14) @p_43) :rule equiv_pos2)
(step t146 (cl @p_46) :rule hole :args ("TRUST_THEORY_REWRITE" @p_46 2 6))
(step t147 (cl (= @p_12 @p_40)) :rule cong :premises (t48 t83))
(step t148 (cl (= @p_13 @p_41)) :rule cong :premises (t147 t82))
(step t149 (cl (= @p_45 @p_43)) :rule cong :premises (t85 t148))
(step t150 (cl @p_44) :rule trans :premises (t146 t149))
(step t151 (cl @p_43) :rule resolution :premises (t145 t150 a3))
(step t152 (cl (! (not (! (or (! (not (! (= @p_42 (! (f3 @p_115 @p_244) :named @p_249)) :named @p_250)) :named @p_251) (! (not (! (= @p_42 @p_246) :named @p_247)) :named @p_248) @p_208) :named @p_252)) :named @p_312) @p_251 @p_248 @p_208) :rule or_pos)
(step t153 (cl @p_208 @p_251 @p_248 @p_312) :rule reordering :premises (t152))
(step t154 (cl (! (not (! (= (! (= @p_42 (! (f3 @p_23 @p_53) :named @p_296)) :named @p_297) (! (= @p_42 (! (f3 @p_293 @p_244) :named @p_294)) :named @p_295)) :named @p_298)) :named @p_311) (! (not @p_297) :named @p_310) @p_295) :rule equiv_pos2)
(step t155 (cl @p_310 @p_295 @p_311) :rule reordering :premises (t154))
(step t156 (cl (not (! (= @p_25 @p_297) :named @p_307)) (not @p_25) @p_297) :rule equiv_pos2)
(step t157 (cl (! (= @p_25 (! (= f1 @p_24) :named @p_308)) :named @p_309)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_309 2 6))
(step t158 (cl (= @p_23 @p_23)) :rule refl)
(step t159 (cl (= @p_16 @p_53)) :rule cong :premises (t70 t74))
(step t160 (cl (= @p_24 @p_296)) :rule cong :premises (t158 t159))
(step t161 (cl (= @p_308 @p_297)) :rule cong :premises (t85 t160))
(step t162 (cl @p_307) :rule trans :premises (t157 t161))
(step t163 (cl @p_297) :rule resolution :premises (t156 t162 a155))
(step t164 (cl (not (! (= (! (=> (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))) (! (= (! (= @p_296 @p_42) :named @p_301) (! (= @p_294 @p_42) :named @p_300)) :named @p_302)) :named @p_303) (! (=> (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))) @p_298) :named @p_299)) :named @p_304)) (not @p_303) @p_299) :rule equiv_pos2)
(step t165 (cl (! (= @p_301 @p_297) :named @p_306)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_306 2 6))
(step t166 (cl (! (= @p_300 @p_295) :named @p_305)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_305 2 6))
(step t167 (cl (= @p_302 @p_298)) :rule cong :premises (t165 t166))
(step t168 (cl @p_304) :rule cong :premises (t121 t167))
(step t169 (cl @p_303 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule implies_neg1)
(anchor :step t170)
(assume t170.a0 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))
(step t170.t0 (cl (or @p_142 @p_302)) :rule forall_inst :args (f25 @p_53))
(step t170.t1 (cl @p_142 @p_302) :rule or :premises (t170.t0))
(step t170.t2 (cl @p_302) :rule resolution :premises (t170.t1 t140))
(step t170 (cl @p_142 @p_302) :rule subproof :discharge (t170.a0))
(step t171 (cl @p_303 @p_302) :rule resolution :premises (t169 t170))
(step t172 (cl @p_303 (not @p_302)) :rule implies_neg2)
(step t173 (cl @p_303 @p_303) :rule resolution :premises (t171 t172))
(step t174 (cl @p_303) :rule contraction :premises (t173))
(step t175 (cl @p_299) :rule resolution :premises (t164 t168 t174))
(step t176 (cl @p_142 @p_298) :rule implies :premises (t175))
(step t177 (cl @p_298) :rule resolution :premises (t176 t140))
(step t178 (cl @p_295) :rule resolution :premises (t155 t163 t177))
(step t179 (cl (= @p_244 @p_244)) :rule refl)
(step t180 (cl (= @p_294 @p_249)) :rule cong :premises (t55 t179))
(step t181 (cl @p_250) :rule trans :premises (t178 t180))
(step t182 (cl (! (and @p_54 @p_205 @p_59) :named @p_289) @p_197 @p_288 @p_287) :rule and_neg)
(step t183 (cl (! (=> @p_289 @p_247) :named @p_291) @p_289) :rule implies_neg1)
(anchor :step t184)
(assume t184.a0 @p_54)
(assume t184.a1 @p_205)
(assume t184.a2 @p_59)
(step t184.t0 (cl @p_292) :rule symm :premises (t184.a2))
(step t184.t1 (cl @p_59) :rule symm :premises (t184.t0))
(step t184.t2 (cl (or @p_52 @p_205)) :rule forall_inst :args (@p_55))
(step t184.t3 (cl @p_52 @p_205) :rule or :premises (t184.t2))
(step t184.t4 (cl @p_205) :rule resolution :premises (t184.t3 t53))
(step t184.t5 (cl (= @p_56 @p_55)) :rule symm :premises (t184.t4))
(step t184.t6 (cl (or @p_52 @p_54)) :rule forall_inst :args (@p_33))
(step t184.t7 (cl @p_52 @p_54) :rule or :premises (t184.t6))
(step t184.t8 (cl @p_54) :rule resolution :premises (t184.t7 t53))
(step t184.t9 (cl (= @p_55 @p_244)) :rule cong :premises (t72 t184.t8))
(step t184.t10 (cl (= @p_56 @p_244)) :rule trans :premises (t184.t5 t184.t9))
(step t184.t11 (cl (= @p_57 @p_245)) :rule cong :premises (t48 t184.t10))
(step t184.t12 (cl (= @p_39 @p_39)) :rule refl)
(step t184.t13 (cl (= @p_58 @p_246)) :rule cong :premises (t184.t11 t184.t12))
(step t184.t14 (cl @p_247) :rule trans :premises (t184.t1 t184.t13))
(step t184 (cl @p_197 @p_288 @p_287 @p_247) :rule subproof :discharge (t184.a0 t184.a1 t184.a2))
(step t185 (cl (! (not @p_289) :named @p_290) @p_54) :rule and_pos :args (0))
(step t186 (cl @p_290 @p_205) :rule and_pos :args (1))
(step t187 (cl @p_290 @p_59) :rule and_pos :args (2))
(step t188 (cl @p_247 @p_290 @p_290 @p_290) :rule resolution :premises (t184 t185 t186 t187))
(step t189 (cl @p_290 @p_290 @p_290 @p_247) :rule reordering :premises (t188))
(step t190 (cl @p_290 @p_247) :rule contraction :premises (t189))
(step t191 (cl @p_291 @p_247) :rule resolution :premises (t183 t190))
(step t192 (cl @p_291 @p_248) :rule implies_neg2)
(step t193 (cl @p_291 @p_291) :rule resolution :premises (t191 t192))
(step t194 (cl @p_291) :rule contraction :premises (t193))
(step t195 (cl @p_290 @p_247) :rule implies :premises (t194))
(step t196 (cl @p_197 @p_288 @p_287 @p_247) :rule resolution :premises (t182 t195))
(step t197 (cl @p_197 @p_287 @p_288 @p_247) :rule reordering :premises (t196))
(step t198 (cl (or @p_52 @p_54)) :rule forall_inst :args (@p_33))
(step t199 (cl @p_52 @p_54) :rule or :premises (t198))
(step t200 (cl @p_54) :rule resolution :premises (t199 t53))
(step t201 (cl (! (not (! (or (! (= @p_39 @p_53) :named @p_63) (! (= @p_39 (! (f5 f9 (! (f5 f11 @p_31) :named @p_60)) :named @p_61)) :named @p_62) @p_59) :named @p_64)) :named @p_204) @p_63 @p_62 @p_59) :rule or_pos)
(step t202 (cl @p_63 @p_62 @p_59 @p_204) :rule reordering :premises (t201))
(step t203 (cl (! (and @p_47 @p_54 @p_63 @p_43) :named @p_198) @p_159 @p_197 (! (not @p_63) :named @p_168) @p_158) :rule and_neg)
(step t204 (cl (! (=> @p_198 (! (= @p_42 (! (f3 (! (f15 f16 (! (f5 f10 @p_32) :named @p_169)) :named @p_170) @p_31) :named @p_171)) :named @p_172)) :named @p_200) @p_198) :rule implies_neg1)
(anchor :step t205)
(assume t205.a0 @p_47)
(assume t205.a1 @p_54)
(assume t205.a2 @p_63)
(assume t205.a3 @p_43)
(step t205.t0 (cl (= @p_40 @p_40)) :rule refl)
(step t205.t1 (cl (= @p_32 @p_31)) :rule symm :premises (t54))
(step t205.t2 (cl (= @p_41 (! (f3 @p_40 @p_31) :named @p_167))) :rule cong :premises (t205.t0 t205.t1))
(step t205.t3 (cl (= @p_53 @p_39)) :rule symm :premises (t205.a2))
(step t205.t4 (cl @p_63) :rule symm :premises (t205.t3))
(step t205.t5 (cl (= @p_53 @p_33)) :rule symm :premises (t200))
(step t205.t6 (cl (= @p_39 @p_33)) :rule trans :premises (t205.t4 t205.t5))
(step t205.t7 (cl (= @p_40 @p_203)) :rule cong :premises (t48 t205.t6))
(step t205.t8 (cl (= (! (f15 f16 @p_53) :named @p_201) @p_203)) :rule cong :premises (t48 t205.t5))
(step t205.t9 (cl (= @p_203 @p_201)) :rule symm :premises (t205.t8))
(step t205.t10 (cl (= @p_40 @p_201)) :rule trans :premises (t205.t7 t205.t9))
(step t205.t11 (cl (= @p_167 (! (f3 @p_201 @p_32) :named @p_202))) :rule cong :premises (t205.t10 t54))
(step t205.t12 (cl (= @p_33 @p_169)) :rule cong :premises (t71 t54))
(step t205.t13 (cl (= @p_203 @p_170)) :rule cong :premises (t48 t205.t12))
(step t205.t14 (cl (= @p_201 @p_170)) :rule trans :premises (t205.t8 t205.t13))
(step t205.t15 (cl (= @p_202 @p_171)) :rule cong :premises (t205.t14 t205.t1))
(step t205.t16 (cl @p_172) :rule trans :premises (t151 t205.t2 t205.t11 t205.t15))
(step t205 (cl @p_159 @p_197 @p_168 @p_158 @p_172) :rule subproof :discharge (t205.a0 t205.a1 t205.a2 t205.a3))
(step t206 (cl (! (not @p_198) :named @p_199) @p_47) :rule and_pos :args (0))
(step t207 (cl @p_199 @p_54) :rule and_pos :args (1))
(step t208 (cl @p_199 @p_63) :rule and_pos :args (2))
(step t209 (cl @p_199 @p_43) :rule and_pos :args (3))
(step t210 (cl @p_172 @p_199 @p_199 @p_199 @p_199) :rule resolution :premises (t205 t206 t207 t208 t209))
(step t211 (cl @p_199 @p_199 @p_199 @p_199 @p_172) :rule reordering :premises (t210))
(step t212 (cl @p_199 @p_172) :rule contraction :premises (t211))
(step t213 (cl @p_200 @p_172) :rule resolution :premises (t204 t212))
(step t214 (cl @p_200 (! (not @p_172) :named @p_173)) :rule implies_neg2)
(step t215 (cl @p_200 @p_200) :rule resolution :premises (t213 t214))
(step t216 (cl @p_200) :rule contraction :premises (t215))
(step t217 (cl @p_199 @p_172) :rule implies :premises (t216))
(step t218 (cl @p_159 @p_197 @p_168 @p_158 @p_172) :rule resolution :premises (t203 t217))
(step t219 (cl @p_158 @p_159 @p_197 @p_172 @p_168) :rule reordering :premises (t218))
(step t220 (cl (! (not (! (= @p_117 @p_172) :named @p_174)) :named @p_196) @p_117 @p_173) :rule equiv_pos1)
(step t221 (cl @p_117 @p_173 @p_196) :rule reordering :premises (t220))
(step t222 (cl (not (! (= (! (=> (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))) (! (= (! (= @p_116 @p_42) :named @p_145) (! (= @p_171 @p_42) :named @p_191)) :named @p_192)) :named @p_193) (! (=> (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))) @p_174) :named @p_190)) :named @p_194)) (not @p_193) @p_190) :rule equiv_pos2)
(step t223 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule refl)
(step t224 (cl (! (= @p_145 @p_117) :named @p_150)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_150 2 6))
(step t225 (cl (! (= @p_191 @p_172) :named @p_195)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_195 2 6))
(step t226 (cl (= @p_192 @p_174)) :rule cong :premises (t224 t225))
(step t227 (cl @p_194) :rule cong :premises (t223 t226))
(step t228 (cl @p_193 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule implies_neg1)
(anchor :step t229)
(assume t229.a0 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))
(step t229.t0 (cl (or (! (not (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :named @p_189) @p_192)) :rule forall_inst :args (@p_32 f25))
(step t229.t1 (cl @p_189 @p_192) :rule or :premises (t229.t0))
(step t229.t2 (cl (not (! (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :named @p_175)) (not (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1)))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule equiv_pos2)
(anchor :step t229.t3 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t229.t3.t0 (cl (! (= (! (= (! (= (! (f3 @p_176 @p_121) :named @p_177) f1) :named @p_184) @p_135) :named @p_185) (! (= @p_135 @p_184) :named @p_186)) :named @p_188)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_188 1 5))
(step t229.t3.t1 (cl @p_140) :rule hole :args ("TRUST_THEORY_REWRITE" @p_140 2 6))
(step t229.t3.t2 (cl (! (= @p_184 (! (= f1 @p_177) :named @p_182)) :named @p_187)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_187 2 6))
(step t229.t3.t3 (cl (= @p_186 (! (= @p_133 @p_182) :named @p_183))) :rule cong :premises (t229.t3.t1 t229.t3.t2))
(step t229.t3.t4 (cl (= @p_185 @p_183)) :rule trans :premises (t229.t3.t0 t229.t3.t3))
(step t229.t3 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1))) (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1))))))) :rule bind)
(anchor :step t229.t4 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t229.t4.t0 (cl (= @p_125 @p_125)) :rule refl)
(step t229.t4.t1 (cl (= @p_133 @p_128)) :rule cong :premises (t85 t229.t4.t0))
(step t229.t4.t2 (cl (= @p_177 @p_177)) :rule refl)
(step t229.t4.t3 (cl (= @p_182 (! (= @p_42 @p_177) :named @p_179))) :rule cong :premises (t85 t229.t4.t2))
(step t229.t4.t4 (cl (= @p_183 (! (= @p_128 @p_179) :named @p_180))) :rule cong :premises (t229.t4.t1 t229.t4.t3))
(step t229.t4 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1))))))) :rule bind)
(anchor :step t229.t5 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t229.t5.t0 (cl @p_131) :rule hole :args ("TRUST_THEORY_REWRITE" @p_131 2 6))
(step t229.t5.t1 (cl (! (= @p_179 (! (= @p_177 @p_42) :named @p_178)) :named @p_181)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_181 2 6))
(step t229.t5.t2 (cl (= @p_180 (= @p_126 @p_178))) :rule cong :premises (t229.t5.t0 t229.t5.t1))
(step t229.t5 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule bind)
(step t229.t6 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule trans :premises (t229.t4 t229.t5))
(step t229.t7 (cl @p_175) :rule trans :premises (t229.t3 t229.t6))
(step t229.t8 (cl (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule resolution :premises (t229.t2 t229.t7 a69))
(step t229.t9 (cl @p_192) :rule resolution :premises (t229.t1 t229.t8))
(step t229 (cl @p_189 @p_192) :rule subproof :discharge (t229.a0))
(step t230 (cl @p_193 @p_192) :rule resolution :premises (t228 t229))
(step t231 (cl @p_193 (not @p_192)) :rule implies_neg2)
(step t232 (cl @p_193 @p_193) :rule resolution :premises (t230 t231))
(step t233 (cl @p_193) :rule contraction :premises (t232))
(step t234 (cl @p_190) :rule resolution :premises (t222 t227 t233))
(step t235 (cl @p_189 @p_174) :rule implies :premises (t234))
(step t236 (cl (not @p_175) (not (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1)))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule equiv_pos2)
(anchor :step t237 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t237.t0 (cl @p_188) :rule hole :args ("TRUST_THEORY_REWRITE" @p_188 1 5))
(step t237.t1 (cl @p_140) :rule hole :args ("TRUST_THEORY_REWRITE" @p_140 2 6))
(step t237.t2 (cl @p_187) :rule hole :args ("TRUST_THEORY_REWRITE" @p_187 2 6))
(step t237.t3 (cl (= @p_186 @p_183)) :rule cong :premises (t237.t1 t237.t2))
(step t237.t4 (cl (= @p_185 @p_183)) :rule trans :premises (t237.t0 t237.t3))
(step t237 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) f1) (= (f3 (f15 f16 ?v0) ?v1) f1))) (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1))))))) :rule bind)
(anchor :step t238 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t238.t0 (cl (= @p_125 @p_125)) :rule refl)
(step t238.t1 (cl (= @p_133 @p_128)) :rule cong :premises (t85 t238.t0))
(step t238.t2 (cl (= @p_177 @p_177)) :rule refl)
(step t238.t3 (cl (= @p_182 @p_179)) :rule cong :premises (t85 t238.t2))
(step t238.t4 (cl (= @p_183 @p_180)) :rule cong :premises (t238.t1 t238.t3))
(step t238 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1))))))) :rule bind)
(anchor :step t239 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t239.t0 (cl @p_131) :rule hole :args ("TRUST_THEORY_REWRITE" @p_131 2 6))
(step t239.t1 (cl @p_181) :rule hole :args ("TRUST_THEORY_REWRITE" @p_181 2 6))
(step t239.t2 (cl (= @p_180 (= @p_126 @p_178))) :rule cong :premises (t239.t0 t239.t1))
(step t239 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1)) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule bind)
(step t240 (cl (= (forall ((?v0 S3) (?v1 S3)) (= (= f1 (f3 (f15 f16 ?v0) ?v1)) (= f1 (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule trans :premises (t238 t239))
(step t241 (cl @p_175) :rule trans :premises (t237 t240))
(step t242 (cl (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f10 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule resolution :premises (t236 t241 a69))
(step t243 (cl @p_174) :rule resolution :premises (t235 t242))
(step t244 (cl @p_173) :rule resolution :premises (t221 t91 t243))
(step t245 (cl @p_168) :rule resolution :premises (t219 t151 t54 t200 t244))
(step t246 (cl (! (and @p_47 (! (= @p_60 @p_61) :named @p_157) @p_62 @p_43) :named @p_161) @p_159 (! (not @p_157) :named @p_160) (! (not @p_62) :named @p_109) @p_158) :rule and_neg)
(step t247 (cl (! (=> @p_161 (! (= @p_42 (! (f3 (! (f15 f16 (! (f5 f11 @p_32) :named @p_110)) :named @p_111) @p_31) :named @p_112)) :named @p_113)) :named @p_163) @p_161) :rule implies_neg1)
(anchor :step t248)
(assume t248.a0 @p_47)
(assume t248.a1 @p_157)
(assume t248.a2 @p_62)
(assume t248.a3 @p_43)
(step t248.t0 (cl (= @p_40 @p_40)) :rule refl)
(step t248.t1 (cl (= @p_32 @p_31)) :rule symm :premises (t54))
(step t248.t2 (cl (= @p_41 @p_167)) :rule cong :premises (t248.t0 t248.t1))
(step t248.t3 (cl (= @p_61 @p_39)) :rule symm :premises (t248.a2))
(step t248.t4 (cl @p_62) :rule symm :premises (t248.t3))
(step t248.t5 (cl (or @p_52 @p_157)) :rule forall_inst :args (@p_60))
(step t248.t6 (cl @p_52 @p_157) :rule or :premises (t248.t5))
(step t248.t7 (cl @p_157) :rule resolution :premises (t248.t6 t53))
(step t248.t8 (cl (= @p_61 @p_60)) :rule symm :premises (t248.t7))
(step t248.t9 (cl (= @p_39 @p_60)) :rule trans :premises (t248.t4 t248.t8))
(step t248.t10 (cl (= @p_40 (! (f15 f16 @p_60) :named @p_166))) :rule cong :premises (t48 t248.t9))
(step t248.t11 (cl (= (! (f15 f16 @p_61) :named @p_164) @p_166)) :rule cong :premises (t48 t248.t8))
(step t248.t12 (cl (= @p_166 @p_164)) :rule symm :premises (t248.t11))
(step t248.t13 (cl (= @p_40 @p_164)) :rule trans :premises (t248.t10 t248.t12))
(step t248.t14 (cl (= @p_167 (! (f3 @p_164 @p_32) :named @p_165))) :rule cong :premises (t248.t13 t54))
(step t248.t15 (cl (= @p_60 @p_110)) :rule cong :premises (t72 t54))
(step t248.t16 (cl (= @p_166 @p_111)) :rule cong :premises (t48 t248.t15))
(step t248.t17 (cl (= @p_164 @p_111)) :rule trans :premises (t248.t11 t248.t16))
(step t248.t18 (cl (= @p_165 @p_112)) :rule cong :premises (t248.t17 t248.t1))
(step t248.t19 (cl @p_113) :rule trans :premises (t151 t248.t2 t248.t14 t248.t18))
(step t248 (cl @p_159 @p_160 @p_109 @p_158 @p_113) :rule subproof :discharge (t248.a0 t248.a1 t248.a2 t248.a3))
(step t249 (cl (! (not @p_161) :named @p_162) @p_47) :rule and_pos :args (0))
(step t250 (cl @p_162 @p_157) :rule and_pos :args (1))
(step t251 (cl @p_162 @p_62) :rule and_pos :args (2))
(step t252 (cl @p_162 @p_43) :rule and_pos :args (3))
(step t253 (cl @p_113 @p_162 @p_162 @p_162 @p_162) :rule resolution :premises (t248 t249 t250 t251 t252))
(step t254 (cl @p_162 @p_162 @p_162 @p_162 @p_113) :rule reordering :premises (t253))
(step t255 (cl @p_162 @p_113) :rule contraction :premises (t254))
(step t256 (cl @p_163 @p_113) :rule resolution :premises (t247 t255))
(step t257 (cl @p_163 (! (not @p_113) :named @p_114)) :rule implies_neg2)
(step t258 (cl @p_163 @p_163) :rule resolution :premises (t256 t257))
(step t259 (cl @p_163) :rule contraction :premises (t258))
(step t260 (cl @p_162 @p_113) :rule implies :premises (t259))
(step t261 (cl @p_159 @p_160 @p_109 @p_158 @p_113) :rule resolution :premises (t246 t260))
(step t262 (cl @p_158 @p_159 @p_160 @p_113 @p_109) :rule reordering :premises (t261))
(step t263 (cl (or @p_52 @p_157)) :rule forall_inst :args (@p_60))
(step t264 (cl @p_52 @p_157) :rule or :premises (t263))
(step t265 (cl @p_157) :rule resolution :premises (t264 t53))
(step t266 (cl (! (not (! (= @p_117 @p_113) :named @p_118)) :named @p_156) @p_117 @p_114) :rule equiv_pos1)
(step t267 (cl @p_117 @p_114 @p_156) :rule reordering :premises (t266))
(step t268 (cl (not (! (= (! (=> (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))) (! (= @p_145 (! (= @p_112 @p_42) :named @p_144)) :named @p_146)) :named @p_147) (! (=> (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))) @p_118) :named @p_143)) :named @p_148)) (not @p_147) @p_143) :rule equiv_pos2)
(step t269 (cl (! (= @p_144 @p_113) :named @p_149)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_149 2 6))
(step t270 (cl (= @p_146 @p_118)) :rule cong :premises (t224 t269))
(step t271 (cl @p_148) :rule cong :premises (t121 t270))
(step t272 (cl @p_147 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule implies_neg1)
(anchor :step t273)
(assume t273.a0 (forall ((?v0 S3) (?v1 S3)) (= (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))) (= (f3 (f15 f16 (f5 f11 ?v0)) (f5 f11 ?v1)) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))
(step t273.t0 (cl (or @p_142 @p_146)) :rule forall_inst :args (@p_32 f25))
(step t273.t1 (cl @p_142 @p_146) :rule or :premises (t273.t0))
(step t273.t2 (cl @p_146) :rule resolution :premises (t273.t1 t140))
(step t273 (cl @p_142 @p_146) :rule subproof :discharge (t273.a0))
(step t274 (cl @p_147 @p_146) :rule resolution :premises (t272 t273))
(step t275 (cl @p_147 (not @p_146)) :rule implies_neg2)
(step t276 (cl @p_147 @p_147) :rule resolution :premises (t274 t275))
(step t277 (cl @p_147) :rule contraction :premises (t276))
(step t278 (cl @p_143) :rule resolution :premises (t268 t271 t277))
(step t279 (cl @p_142 @p_118) :rule implies :premises (t278))
(step t280 (cl @p_118) :rule resolution :premises (t279 t140))
(step t281 (cl @p_114) :rule resolution :premises (t267 t91 t280))
(step t282 (cl @p_109) :rule resolution :premises (t262 t151 t54 t265 t281))
(step t283 (cl (not (! (= (! (=> (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0)))) (! (or (! (not @p_99) :named @p_100) @p_63 @p_62 @p_59) :named @p_101)) :named @p_102) (! (=> (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0)))) @p_64) :named @p_98)) :named @p_103)) (not @p_102) @p_98) :rule equiv_pos2)
(step t284 (cl (= (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0)))) (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0)))))) :rule refl)
(step t285 (cl (! (= @p_99 true) :named @p_108)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_108 2 5))
(step t286 (cl (= @p_100 (! (not true) :named @p_106))) :rule cong :premises (t285))
(step t287 (cl (! (= @p_106 false) :named @p_107)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_107 1 6))
(step t288 (cl (= @p_100 false)) :rule trans :premises (t286 t287))
(step t289 (cl (= @p_63 @p_63)) :rule refl)
(step t290 (cl (= @p_62 @p_62)) :rule refl)
(step t291 (cl (= @p_59 @p_59)) :rule refl)
(step t292 (cl (= @p_101 (! (or false @p_63 @p_62 @p_59) :named @p_104))) :rule cong :premises (t288 t289 t290 t291))
(step t293 (cl (! (= @p_104 @p_64) :named @p_105)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_105 1 6))
(step t294 (cl (= @p_101 @p_64)) :rule trans :premises (t292 t293))
(step t295 (cl @p_103) :rule cong :premises (t284 t294))
(step t296 (cl @p_102 (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0))))) :rule implies_neg1)
(anchor :step t297)
(assume t297.a0 (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0)))))
(step t297.t0 (cl (or (! (not (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0))))) :named @p_97) @p_101)) :rule forall_inst :args (@p_39))
(step t297.t1 (cl @p_97 @p_101) :rule or :premises (t297.t0))
(step t297.t2 (cl (not (! (= (forall ((?v0 S3)) (=> (= (f3 f4 ?v0) f1) (=> (not (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12))))) (=> (not (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12))))) (= (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0) f1))))) (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0))))) :named @p_65)) (not (forall ((?v0 S3)) (=> (= (f3 f4 ?v0) f1) (=> (not (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12))))) (=> (not (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12))))) (= (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0) f1)))))) (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0))))) :rule equiv_pos2)
(anchor :step t297.t3 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t297.t3.t0 (cl (! (= (! (= (! (f3 f4 ?v0) :named @p_70) f1) :named @p_94) (! (= f1 @p_70) :named @p_84)) :named @p_96)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_96 2 6))
(step t297.t3.t1 (cl (= (! (not (! (= ?v0 @p_16) :named @p_83)) :named @p_89) @p_89)) :rule refl)
(step t297.t3.t2 (cl (= (! (not (! (= ?v0 @p_27) :named @p_82)) :named @p_87) @p_87)) :rule refl)
(step t297.t3.t3 (cl (! (= (! (= (! (f3 (! (f15 f16 (! (f5 f9 (! (f5 f11 @p_3) :named @p_77)) :named @p_78)) :named @p_79) ?v0) :named @p_80) f1) :named @p_91) (! (= f1 @p_80) :named @p_81)) :named @p_95)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_95 2 6))
(step t297.t3.t4 (cl (= (! (=> @p_87 @p_91) :named @p_92) (! (=> @p_87 @p_81) :named @p_88))) :rule cong :premises (t297.t3.t2 t297.t3.t3))
(step t297.t3.t5 (cl (= (! (=> @p_89 @p_92) :named @p_93) (! (=> @p_89 @p_88) :named @p_90))) :rule cong :premises (t297.t3.t1 t297.t3.t4))
(step t297.t3.t6 (cl (= (=> @p_94 @p_93) (=> @p_84 @p_90))) :rule cong :premises (t297.t3.t0 t297.t3.t5))
(step t297.t3 (cl (= (forall ((?v0 S3)) (=> (= (f3 f4 ?v0) f1) (=> (not (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12))))) (=> (not (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12))))) (= (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0) f1))))) (forall ((?v0 S3)) (=> (= f1 (f3 f4 ?v0)) (=> (not (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12))))) (=> (not (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12))))) (= f1 (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0)))))))) :rule bind)
(step t297.t4 (cl (! (= (forall ((?v0 S3)) (=> (= f1 (f3 f4 ?v0)) (=> (not (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12))))) (=> (not (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12))))) (= f1 (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0)))))) (forall ((?v0 S3)) (or (not (= f1 (f3 f4 ?v0))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12)))) (= f1 (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0))))) :named @p_86)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_86 13 6))
(step t297.t5 (cl (= (forall ((?v0 S3)) (=> (= (f3 f4 ?v0) f1) (=> (not (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12))))) (=> (not (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12))))) (= (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0) f1))))) (forall ((?v0 S3)) (or (not (= f1 (f3 f4 ?v0))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12)))) (= f1 (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0)))))) :rule trans :premises (t297.t3 t297.t4))
(anchor :step t297.t6 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t297.t6.t0 (cl (= @p_70 @p_70)) :rule refl)
(step t297.t6.t1 (cl (= @p_84 (! (= @p_42 @p_70) :named @p_73))) :rule cong :premises (t85 t297.t6.t0))
(step t297.t6.t2 (cl (= (! (not @p_84) :named @p_85) (! (not @p_73) :named @p_74))) :rule cong :premises (t297.t6.t1))
(step t297.t6.t3 (cl (! (= ?v0 ?v0) :named @p_49)) :rule refl)
(step t297.t6.t4 (cl (= @p_83 (! (= ?v0 @p_53) :named @p_69))) :rule cong :premises (t297.t6.t3 t159))
(step t297.t6.t5 (cl (= @p_26 @p_60)) :rule cong :premises (t72 t73))
(step t297.t6.t6 (cl (= @p_27 @p_61)) :rule cong :premises (t70 t297.t6.t5))
(step t297.t6.t7 (cl (= @p_82 (! (= ?v0 @p_61) :named @p_68))) :rule cong :premises (t297.t6.t3 t297.t6.t6))
(step t297.t6.t8 (cl (= @p_77 @p_55)) :rule cong :premises (t72 t74))
(step t297.t6.t9 (cl (= @p_78 @p_56)) :rule cong :premises (t70 t297.t6.t8))
(step t297.t6.t10 (cl (= @p_79 @p_57)) :rule cong :premises (t48 t297.t6.t9))
(step t297.t6.t11 (cl (= @p_80 (! (f3 @p_57 ?v0) :named @p_66))) :rule cong :premises (t297.t6.t10 t297.t6.t3))
(step t297.t6.t12 (cl (= @p_81 (! (= @p_42 @p_66) :named @p_67))) :rule cong :premises (t85 t297.t6.t11))
(step t297.t6.t13 (cl (= (or @p_85 @p_83 @p_82 @p_81) (! (or @p_74 @p_69 @p_68 @p_67) :named @p_75))) :rule cong :premises (t297.t6.t2 t297.t6.t4 t297.t6.t7 t297.t6.t12))
(step t297.t6 (cl (= (forall ((?v0 S3)) (or (not (= f1 (f3 f4 ?v0))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12)))) (= f1 (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0)))) (forall ((?v0 S3)) (or (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 f4 ?v0))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0)))))) :rule bind)
(anchor :step t297.t7 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t297.t7.t0 (cl (! (= @p_73 (! (= @p_70 @p_42) :named @p_71)) :named @p_76)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_76 2 6))
(step t297.t7.t1 (cl (= @p_74 (! (not @p_71) :named @p_72))) :rule cong :premises (t297.t7.t0))
(step t297.t7.t2 (cl (= @p_69 @p_69)) :rule refl)
(step t297.t7.t3 (cl (= @p_68 @p_68)) :rule refl)
(step t297.t7.t4 (cl (= @p_67 @p_67)) :rule refl)
(step t297.t7.t5 (cl (= @p_75 (or @p_72 @p_69 @p_68 @p_67))) :rule cong :premises (t297.t7.t1 t297.t7.t2 t297.t7.t3 t297.t7.t4))
(step t297.t7 (cl (= (forall ((?v0 S3)) (or (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 f4 ?v0))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0)))) (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0)))))) :rule bind)
(step t297.t8 (cl (= (forall ((?v0 S3)) (or (not (= f1 (f3 f4 ?v0))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12)))) (= f1 (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0)))) (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0)))))) :rule trans :premises (t297.t6 t297.t7))
(step t297.t9 (cl @p_65) :rule trans :premises (t297.t5 t297.t8))
(step t297.t10 (cl (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0))))) :rule resolution :premises (t297.t2 t297.t9 a6))
(step t297.t11 (cl @p_101) :rule resolution :premises (t297.t1 t297.t10))
(step t297 (cl @p_97 @p_101) :rule subproof :discharge (t297.a0))
(step t298 (cl @p_102 @p_101) :rule resolution :premises (t296 t297))
(step t299 (cl @p_102 (not @p_101)) :rule implies_neg2)
(step t300 (cl @p_102 @p_102) :rule resolution :premises (t298 t299))
(step t301 (cl @p_102) :rule contraction :premises (t300))
(step t302 (cl @p_98) :rule resolution :premises (t283 t295 t301))
(step t303 (cl @p_97 @p_64) :rule implies :premises (t302))
(step t304 (cl (not @p_65) (not (forall ((?v0 S3)) (=> (= (f3 f4 ?v0) f1) (=> (not (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12))))) (=> (not (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12))))) (= (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0) f1)))))) (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0))))) :rule equiv_pos2)
(anchor :step t305 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t305.t0 (cl @p_96) :rule hole :args ("TRUST_THEORY_REWRITE" @p_96 2 6))
(step t305.t1 (cl (= @p_89 @p_89)) :rule refl)
(step t305.t2 (cl (= @p_87 @p_87)) :rule refl)
(step t305.t3 (cl @p_95) :rule hole :args ("TRUST_THEORY_REWRITE" @p_95 2 6))
(step t305.t4 (cl (= @p_92 @p_88)) :rule cong :premises (t305.t2 t305.t3))
(step t305.t5 (cl (= @p_93 @p_90)) :rule cong :premises (t305.t1 t305.t4))
(step t305.t6 (cl (= (=> @p_94 @p_93) (=> @p_84 @p_90))) :rule cong :premises (t305.t0 t305.t5))
(step t305 (cl (= (forall ((?v0 S3)) (=> (= (f3 f4 ?v0) f1) (=> (not (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12))))) (=> (not (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12))))) (= (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0) f1))))) (forall ((?v0 S3)) (=> (= f1 (f3 f4 ?v0)) (=> (not (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12))))) (=> (not (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12))))) (= f1 (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0)))))))) :rule bind)
(step t306 (cl @p_86) :rule hole :args ("TRUST_THEORY_REWRITE" @p_86 13 6))
(step t307 (cl (= (forall ((?v0 S3)) (=> (= (f3 f4 ?v0) f1) (=> (not (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12))))) (=> (not (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12))))) (= (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0) f1))))) (forall ((?v0 S3)) (or (not (= f1 (f3 f4 ?v0))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12)))) (= f1 (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0)))))) :rule trans :premises (t305 t306))
(anchor :step t308 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t308.t0 (cl (= @p_70 @p_70)) :rule refl)
(step t308.t1 (cl (= @p_84 @p_73)) :rule cong :premises (t85 t308.t0))
(step t308.t2 (cl (= @p_85 @p_74)) :rule cong :premises (t308.t1))
(step t308.t3 (cl @p_49) :rule refl)
(step t308.t4 (cl (= @p_83 @p_69)) :rule cong :premises (t308.t3 t159))
(step t308.t5 (cl (= @p_26 @p_60)) :rule cong :premises (t72 t73))
(step t308.t6 (cl (= @p_27 @p_61)) :rule cong :premises (t70 t308.t5))
(step t308.t7 (cl (= @p_82 @p_68)) :rule cong :premises (t308.t3 t308.t6))
(step t308.t8 (cl (= @p_77 @p_55)) :rule cong :premises (t72 t74))
(step t308.t9 (cl (= @p_78 @p_56)) :rule cong :premises (t70 t308.t8))
(step t308.t10 (cl (= @p_79 @p_57)) :rule cong :premises (t48 t308.t9))
(step t308.t11 (cl (= @p_80 @p_66)) :rule cong :premises (t308.t10 t308.t3))
(step t308.t12 (cl (= @p_81 @p_67)) :rule cong :premises (t85 t308.t11))
(step t308.t13 (cl (= (or @p_85 @p_83 @p_82 @p_81) @p_75)) :rule cong :premises (t308.t2 t308.t4 t308.t7 t308.t12))
(step t308 (cl (= (forall ((?v0 S3)) (or (not (= f1 (f3 f4 ?v0))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12)))) (= f1 (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0)))) (forall ((?v0 S3)) (or (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 f4 ?v0))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0)))))) :rule bind)
(anchor :step t309 :args ((?v0 S3) (:= (?v0 S3) ?v0)))
(step t309.t0 (cl @p_76) :rule hole :args ("TRUST_THEORY_REWRITE" @p_76 2 6))
(step t309.t1 (cl (= @p_74 @p_72)) :rule cong :premises (t309.t0))
(step t309.t2 (cl (= @p_69 @p_69)) :rule refl)
(step t309.t3 (cl (= @p_68 @p_68)) :rule refl)
(step t309.t4 (cl (= @p_67 @p_67)) :rule refl)
(step t309.t5 (cl (= @p_75 (or @p_72 @p_69 @p_68 @p_67))) :rule cong :premises (t309.t1 t309.t2 t309.t3 t309.t4))
(step t309 (cl (= (forall ((?v0 S3)) (or (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 f4 ?v0))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0)))) (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0)))))) :rule bind)
(step t310 (cl (= (forall ((?v0 S3)) (or (not (= f1 (f3 f4 ?v0))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f12)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f12)))) (= f1 (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f12))))) ?v0)))) (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0)))))) :rule trans :premises (t308 t309))
(step t311 (cl @p_65) :rule trans :premises (t307 t310))
(step t312 (cl (forall ((?v0 S3)) (or (not (= (f3 f4 ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 (f5 f9 (f5 f10 (f5 f11 f25)))) (= ?v0 (f5 f9 (f5 f11 (f5 f11 f25)))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 (f5 f9 (f5 f11 (f5 f10 (f5 f11 f25))))) ?v0))))) :rule resolution :premises (t304 t311 a6))
(step t313 (cl @p_64) :rule resolution :premises (t303 t312))
(step t314 (cl @p_59) :rule resolution :premises (t202 t245 t282 t313))
(step t315 (cl (or @p_52 @p_205)) :rule forall_inst :args (@p_55))
(step t316 (cl @p_52 @p_205) :rule or :premises (t315))
(step t317 (cl @p_205) :rule resolution :premises (t316 t53))
(step t318 (cl @p_247) :rule resolution :premises (t197 t200 t314 t317))
(step t319 (cl (not (! (= (! (=> (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))) (! (or (! (not (! (= @p_249 @p_42) :named @p_280)) :named @p_281) (! (not (! (= @p_246 @p_42) :named @p_278)) :named @p_279) (! (= @p_207 @p_42) :named @p_234)) :named @p_282)) :named @p_283) (! (=> (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))) @p_252) :named @p_277)) :named @p_284)) (not @p_283) @p_277) :rule equiv_pos2)
(step t320 (cl (= (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))) (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule refl)
(step t321 (cl (! (= @p_280 @p_250) :named @p_286)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_286 2 6))
(step t322 (cl (= @p_281 @p_251)) :rule cong :premises (t321))
(step t323 (cl (! (= @p_278 @p_247) :named @p_285)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_285 2 6))
(step t324 (cl (= @p_279 @p_248)) :rule cong :premises (t323))
(step t325 (cl (! (= @p_234 @p_208) :named @p_242)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_242 2 6))
(step t326 (cl (= @p_282 @p_252)) :rule cong :premises (t322 t324 t325))
(step t327 (cl @p_284) :rule cong :premises (t320 t326))
(step t328 (cl @p_283 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule implies_neg1)
(anchor :step t329)
(assume t329.a0 (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))
(step t329.t0 (cl (or (! (not (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :named @p_276) @p_282)) :rule forall_inst :args (@p_32 @p_244 @p_39))
(step t329.t1 (cl @p_276 @p_282) :rule or :premises (t329.t0))
(step t329.t2 (cl (not (! (= (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v2) f1) (= (f3 (f15 f16 ?v0) ?v2) f1)))) (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :named @p_253)) (not (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v2) f1) (= (f3 (f15 f16 ?v0) ?v2) f1))))) (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule equiv_pos2)
(anchor :step t329.t3 :args ((?v0 S3) (?v1 S3) (?v2 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1) (:= (?v2 S3) ?v2)))
(step t329.t3.t0 (cl @p_140) :rule hole :args ("TRUST_THEORY_REWRITE" @p_140 2 6))
(step t329.t3.t1 (cl (! (= (! (= (! (f3 (! (f15 f16 ?v1) :named @p_213) ?v2) :named @p_257) f1) :named @p_272) (! (= f1 @p_257) :named @p_267)) :named @p_275)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_275 2 6))
(step t329.t3.t2 (cl (! (= (! (= (! (f3 @p_124 ?v2) :named @p_255) f1) :named @p_271) (! (= f1 @p_255) :named @p_266)) :named @p_274)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_274 2 6))
(step t329.t3.t3 (cl (= (! (=> @p_272 @p_271) :named @p_273) (! (=> @p_267 @p_266) :named @p_270))) :rule cong :premises (t329.t3.t1 t329.t3.t2))
(step t329.t3.t4 (cl (= (=> @p_135 @p_273) (=> @p_133 @p_270))) :rule cong :premises (t329.t3.t0 t329.t3.t3))
(step t329.t3 (cl (= (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v2) f1) (= (f3 (f15 f16 ?v0) ?v2) f1)))) (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= f1 (f3 (f15 f16 ?v0) ?v1)) (=> (= f1 (f3 (f15 f16 ?v1) ?v2)) (= f1 (f3 (f15 f16 ?v0) ?v2))))))) :rule bind)
(step t329.t4 (cl (! (= (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= f1 (f3 (f15 f16 ?v0) ?v1)) (=> (= f1 (f3 (f15 f16 ?v1) ?v2)) (= f1 (f3 (f15 f16 ?v0) ?v2))))) (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= f1 (f3 (f15 f16 ?v0) ?v1))) (not (= f1 (f3 (f15 f16 ?v1) ?v2))) (= f1 (f3 (f15 f16 ?v0) ?v2))))) :named @p_269)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_269 13 6))
(step t329.t5 (cl (= (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v2) f1) (= (f3 (f15 f16 ?v0) ?v2) f1)))) (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= f1 (f3 (f15 f16 ?v0) ?v1))) (not (= f1 (f3 (f15 f16 ?v1) ?v2))) (= f1 (f3 (f15 f16 ?v0) ?v2)))))) :rule trans :premises (t329.t3 t329.t4))
(anchor :step t329.t6 :args ((?v0 S3) (?v1 S3) (?v2 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1) (:= (?v2 S3) ?v2)))
(step t329.t6.t0 (cl (= @p_125 @p_125)) :rule refl)
(step t329.t6.t1 (cl (= @p_133 @p_128)) :rule cong :premises (t85 t329.t6.t0))
(step t329.t6.t2 (cl (= (! (not @p_133) :named @p_225) (! (not @p_128) :named @p_220))) :rule cong :premises (t329.t6.t1))
(step t329.t6.t3 (cl (= @p_257 @p_257)) :rule refl)
(step t329.t6.t4 (cl (= @p_267 (! (= @p_42 @p_257) :named @p_261))) :rule cong :premises (t85 t329.t6.t3))
(step t329.t6.t5 (cl (= (! (not @p_267) :named @p_268) (! (not @p_261) :named @p_262))) :rule cong :premises (t329.t6.t4))
(step t329.t6.t6 (cl (= @p_255 @p_255)) :rule refl)
(step t329.t6.t7 (cl (= @p_266 (! (= @p_42 @p_255) :named @p_260))) :rule cong :premises (t85 t329.t6.t6))
(step t329.t6.t8 (cl (= (or @p_225 @p_268 @p_266) (! (or @p_220 @p_262 @p_260) :named @p_263))) :rule cong :premises (t329.t6.t2 t329.t6.t5 t329.t6.t7))
(step t329.t6 (cl (= (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= f1 (f3 (f15 f16 ?v0) ?v1))) (not (= f1 (f3 (f15 f16 ?v1) ?v2))) (= f1 (f3 (f15 f16 ?v0) ?v2)))) (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1))) (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v1) ?v2))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v2)))))) :rule bind)
(anchor :step t329.t7 :args ((?v0 S3) (?v1 S3) (?v2 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1) (:= (?v2 S3) ?v2)))
(step t329.t7.t0 (cl @p_131) :rule hole :args ("TRUST_THEORY_REWRITE" @p_131 2 6))
(step t329.t7.t1 (cl (= @p_220 (! (not @p_126) :named @p_217))) :rule cong :premises (t329.t7.t0))
(step t329.t7.t2 (cl (! (= @p_261 (! (= @p_257 @p_42) :named @p_258)) :named @p_265)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_265 2 6))
(step t329.t7.t3 (cl (= @p_262 (! (not @p_258) :named @p_259))) :rule cong :premises (t329.t7.t2))
(step t329.t7.t4 (cl (! (= @p_260 (! (= @p_255 @p_42) :named @p_256)) :named @p_264)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_264 2 6))
(step t329.t7.t5 (cl (= @p_263 (or @p_217 @p_259 @p_256))) :rule cong :premises (t329.t7.t1 t329.t7.t3 t329.t7.t4))
(step t329.t7 (cl (= (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1))) (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v1) ?v2))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v2)))) (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule bind)
(step t329.t8 (cl (= (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= f1 (f3 (f15 f16 ?v0) ?v1))) (not (= f1 (f3 (f15 f16 ?v1) ?v2))) (= f1 (f3 (f15 f16 ?v0) ?v2)))) (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule trans :premises (t329.t6 t329.t7))
(step t329.t9 (cl @p_253) :rule trans :premises (t329.t5 t329.t8))
(step t329.t10 (cl (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule resolution :premises (t329.t2 t329.t9 a25))
(step t329.t11 (cl @p_282) :rule resolution :premises (t329.t1 t329.t10))
(step t329 (cl @p_276 @p_282) :rule subproof :discharge (t329.a0))
(step t330 (cl @p_283 @p_282) :rule resolution :premises (t328 t329))
(step t331 (cl @p_283 (not @p_282)) :rule implies_neg2)
(step t332 (cl @p_283 @p_283) :rule resolution :premises (t330 t331))
(step t333 (cl @p_283) :rule contraction :premises (t332))
(step t334 (cl @p_277) :rule resolution :premises (t319 t327 t333))
(step t335 (cl @p_276 @p_252) :rule implies :premises (t334))
(step t336 (cl (not @p_253) (not (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v2) f1) (= (f3 (f15 f16 ?v0) ?v2) f1))))) (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule equiv_pos2)
(anchor :step t337 :args ((?v0 S3) (?v1 S3) (?v2 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1) (:= (?v2 S3) ?v2)))
(step t337.t0 (cl @p_140) :rule hole :args ("TRUST_THEORY_REWRITE" @p_140 2 6))
(step t337.t1 (cl @p_275) :rule hole :args ("TRUST_THEORY_REWRITE" @p_275 2 6))
(step t337.t2 (cl @p_274) :rule hole :args ("TRUST_THEORY_REWRITE" @p_274 2 6))
(step t337.t3 (cl (= @p_273 @p_270)) :rule cong :premises (t337.t1 t337.t2))
(step t337.t4 (cl (= (=> @p_135 @p_273) (=> @p_133 @p_270))) :rule cong :premises (t337.t0 t337.t3))
(step t337 (cl (= (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v2) f1) (= (f3 (f15 f16 ?v0) ?v2) f1)))) (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= f1 (f3 (f15 f16 ?v0) ?v1)) (=> (= f1 (f3 (f15 f16 ?v1) ?v2)) (= f1 (f3 (f15 f16 ?v0) ?v2))))))) :rule bind)
(step t338 (cl @p_269) :rule hole :args ("TRUST_THEORY_REWRITE" @p_269 13 6))
(step t339 (cl (= (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v2) f1) (= (f3 (f15 f16 ?v0) ?v2) f1)))) (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= f1 (f3 (f15 f16 ?v0) ?v1))) (not (= f1 (f3 (f15 f16 ?v1) ?v2))) (= f1 (f3 (f15 f16 ?v0) ?v2)))))) :rule trans :premises (t337 t338))
(anchor :step t340 :args ((?v0 S3) (?v1 S3) (?v2 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1) (:= (?v2 S3) ?v2)))
(step t340.t0 (cl (= @p_125 @p_125)) :rule refl)
(step t340.t1 (cl (= @p_133 @p_128)) :rule cong :premises (t85 t340.t0))
(step t340.t2 (cl (= @p_225 @p_220)) :rule cong :premises (t340.t1))
(step t340.t3 (cl (= @p_257 @p_257)) :rule refl)
(step t340.t4 (cl (= @p_267 @p_261)) :rule cong :premises (t85 t340.t3))
(step t340.t5 (cl (= @p_268 @p_262)) :rule cong :premises (t340.t4))
(step t340.t6 (cl (= @p_255 @p_255)) :rule refl)
(step t340.t7 (cl (= @p_266 @p_260)) :rule cong :premises (t85 t340.t6))
(step t340.t8 (cl (= (or @p_225 @p_268 @p_266) @p_263)) :rule cong :premises (t340.t2 t340.t5 t340.t7))
(step t340 (cl (= (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= f1 (f3 (f15 f16 ?v0) ?v1))) (not (= f1 (f3 (f15 f16 ?v1) ?v2))) (= f1 (f3 (f15 f16 ?v0) ?v2)))) (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1))) (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v1) ?v2))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v2)))))) :rule bind)
(anchor :step t341 :args ((?v0 S3) (?v1 S3) (?v2 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1) (:= (?v2 S3) ?v2)))
(step t341.t0 (cl @p_131) :rule hole :args ("TRUST_THEORY_REWRITE" @p_131 2 6))
(step t341.t1 (cl (= @p_220 @p_217)) :rule cong :premises (t341.t0))
(step t341.t2 (cl @p_265) :rule hole :args ("TRUST_THEORY_REWRITE" @p_265 2 6))
(step t341.t3 (cl (= @p_262 @p_259)) :rule cong :premises (t341.t2))
(step t341.t4 (cl @p_264) :rule hole :args ("TRUST_THEORY_REWRITE" @p_264 2 6))
(step t341.t5 (cl (= @p_263 (or @p_217 @p_259 @p_256))) :rule cong :premises (t341.t1 t341.t3 t341.t4))
(step t341 (cl (= (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1))) (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v1) ?v2))) (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v2)))) (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule bind)
(step t342 (cl (= (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= f1 (f3 (f15 f16 ?v0) ?v1))) (not (= f1 (f3 (f15 f16 ?v1) ?v2))) (= f1 (f3 (f15 f16 ?v0) ?v2)))) (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25))))))))) :rule trans :premises (t340 t341))
(step t343 (cl @p_253) :rule trans :premises (t339 t342))
(step t344 (cl (forall ((?v0 S3) (?v1 S3) (?v2 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= (f3 (f15 f16 ?v0) ?v2) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))))) :rule resolution :premises (t336 t343 a25))
(step t345 (cl @p_252) :rule resolution :premises (t335 t344))
(step t346 (cl @p_208) :rule resolution :premises (t153 t181 t318 t345))
(step t347 (cl (not (! (= (! (=> (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1))) (! (or (! (not @p_236) :named @p_237) (! (not @p_234) :named @p_235) @p_233) :named @p_238)) :named @p_239) (! (=> (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1))) @p_210) :named @p_232)) :named @p_240)) (not @p_239) @p_232) :rule equiv_pos2)
(step t348 (cl (= (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1))) (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1))))) :rule refl)
(step t349 (cl (! (= @p_236 @p_43) :named @p_243)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_243 2 6))
(step t350 (cl (= @p_237 @p_158)) :rule cong :premises (t349))
(step t351 (cl (= @p_235 @p_209)) :rule cong :premises (t325))
(step t352 (cl (! (= @p_233 @p_206) :named @p_241)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_241 2 6))
(step t353 (cl (= @p_238 @p_210)) :rule cong :premises (t350 t351 t352))
(step t354 (cl @p_240) :rule cong :premises (t348 t353))
(step t355 (cl @p_239 (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1)))) :rule implies_neg1)
(anchor :step t356)
(assume t356.a0 (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1))))
(step t356.t0 (cl (or (! (not (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1)))) :named @p_231) @p_238)) :rule forall_inst :args (@p_39 @p_32))
(step t356.t1 (cl @p_231 @p_238) :rule or :premises (t356.t0))
(step t356.t2 (cl (not (! (= (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v0) f1) (= ?v0 ?v1)))) (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1)))) :named @p_211)) (not (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v0) f1) (= ?v0 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1)))) :rule equiv_pos2)
(anchor :step t356.t3 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t356.t3.t0 (cl @p_140) :rule hole :args ("TRUST_THEORY_REWRITE" @p_140 2 6))
(step t356.t3.t1 (cl (! (= (! (= (! (f3 @p_213 ?v0) :named @p_214) f1) :named @p_228) (! (= f1 @p_214) :named @p_223)) :named @p_230)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_230 2 6))
(step t356.t3.t2 (cl (= (! (= ?v0 ?v1) :named @p_212) @p_212)) :rule refl)
(step t356.t3.t3 (cl (= (! (=> @p_228 @p_212) :named @p_229) (! (=> @p_223 @p_212) :named @p_227))) :rule cong :premises (t356.t3.t1 t356.t3.t2))
(step t356.t3.t4 (cl (= (=> @p_135 @p_229) (=> @p_133 @p_227))) :rule cong :premises (t356.t3.t0 t356.t3.t3))
(step t356.t3 (cl (= (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v0) f1) (= ?v0 ?v1)))) (forall ((?v0 S3) (?v1 S3)) (=> (= f1 (f3 (f15 f16 ?v0) ?v1)) (=> (= f1 (f3 (f15 f16 ?v1) ?v0)) (= ?v0 ?v1)))))) :rule bind)
(step t356.t4 (cl (! (= (forall ((?v0 S3) (?v1 S3)) (=> (= f1 (f3 (f15 f16 ?v0) ?v1)) (=> (= f1 (f3 (f15 f16 ?v1) ?v0)) (= ?v0 ?v1)))) (forall ((?v0 S3) (?v1 S3)) (or (not (= f1 (f3 (f15 f16 ?v0) ?v1))) (not (= f1 (f3 (f15 f16 ?v1) ?v0))) (= ?v0 ?v1)))) :named @p_226)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_226 13 6))
(step t356.t5 (cl (= (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v0) f1) (= ?v0 ?v1)))) (forall ((?v0 S3) (?v1 S3)) (or (not (= f1 (f3 (f15 f16 ?v0) ?v1))) (not (= f1 (f3 (f15 f16 ?v1) ?v0))) (= ?v0 ?v1))))) :rule trans :premises (t356.t3 t356.t4))
(anchor :step t356.t6 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t356.t6.t0 (cl (= @p_125 @p_125)) :rule refl)
(step t356.t6.t1 (cl (= @p_133 @p_128)) :rule cong :premises (t85 t356.t6.t0))
(step t356.t6.t2 (cl (= @p_225 @p_220)) :rule cong :premises (t356.t6.t1))
(step t356.t6.t3 (cl (= @p_214 @p_214)) :rule refl)
(step t356.t6.t4 (cl (= @p_223 (! (= @p_42 @p_214) :named @p_218))) :rule cong :premises (t85 t356.t6.t3))
(step t356.t6.t5 (cl (= (! (not @p_223) :named @p_224) (! (not @p_218) :named @p_219))) :rule cong :premises (t356.t6.t4))
(step t356.t6.t6 (cl (= @p_212 @p_212)) :rule refl)
(step t356.t6.t7 (cl (= (or @p_225 @p_224 @p_212) (! (or @p_220 @p_219 @p_212) :named @p_221))) :rule cong :premises (t356.t6.t2 t356.t6.t5 t356.t6.t6))
(step t356.t6 (cl (= (forall ((?v0 S3) (?v1 S3)) (or (not (= f1 (f3 (f15 f16 ?v0) ?v1))) (not (= f1 (f3 (f15 f16 ?v1) ?v0))) (= ?v0 ?v1))) (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1))) (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v1) ?v0))) (= ?v0 ?v1))))) :rule bind)
(anchor :step t356.t7 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t356.t7.t0 (cl @p_131) :rule hole :args ("TRUST_THEORY_REWRITE" @p_131 2 6))
(step t356.t7.t1 (cl (= @p_220 @p_217)) :rule cong :premises (t356.t7.t0))
(step t356.t7.t2 (cl (! (= @p_218 (! (= @p_214 @p_42) :named @p_215)) :named @p_222)) :rule hole :args ("TRUST_THEORY_REWRITE" @p_222 2 6))
(step t356.t7.t3 (cl (= @p_219 (! (not @p_215) :named @p_216))) :rule cong :premises (t356.t7.t2))
(step t356.t7.t4 (cl (= @p_212 @p_212)) :rule refl)
(step t356.t7.t5 (cl (= @p_221 (or @p_217 @p_216 @p_212))) :rule cong :premises (t356.t7.t1 t356.t7.t3 t356.t7.t4))
(step t356.t7 (cl (= (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1))) (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v1) ?v0))) (= ?v0 ?v1))) (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1))))) :rule bind)
(step t356.t8 (cl (= (forall ((?v0 S3) (?v1 S3)) (or (not (= f1 (f3 (f15 f16 ?v0) ?v1))) (not (= f1 (f3 (f15 f16 ?v1) ?v0))) (= ?v0 ?v1))) (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1))))) :rule trans :premises (t356.t6 t356.t7))
(step t356.t9 (cl @p_211) :rule trans :premises (t356.t5 t356.t8))
(step t356.t10 (cl (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1)))) :rule resolution :premises (t356.t2 t356.t9 a24))
(step t356.t11 (cl @p_238) :rule resolution :premises (t356.t1 t356.t10))
(step t356 (cl @p_231 @p_238) :rule subproof :discharge (t356.a0))
(step t357 (cl @p_239 @p_238) :rule resolution :premises (t355 t356))
(step t358 (cl @p_239 (not @p_238)) :rule implies_neg2)
(step t359 (cl @p_239 @p_239) :rule resolution :premises (t357 t358))
(step t360 (cl @p_239) :rule contraction :premises (t359))
(step t361 (cl @p_232) :rule resolution :premises (t347 t354 t360))
(step t362 (cl @p_231 @p_210) :rule implies :premises (t361))
(step t363 (cl (not @p_211) (not (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v0) f1) (= ?v0 ?v1))))) (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1)))) :rule equiv_pos2)
(anchor :step t364 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t364.t0 (cl @p_140) :rule hole :args ("TRUST_THEORY_REWRITE" @p_140 2 6))
(step t364.t1 (cl @p_230) :rule hole :args ("TRUST_THEORY_REWRITE" @p_230 2 6))
(step t364.t2 (cl (= @p_212 @p_212)) :rule refl)
(step t364.t3 (cl (= @p_229 @p_227)) :rule cong :premises (t364.t1 t364.t2))
(step t364.t4 (cl (= (=> @p_135 @p_229) (=> @p_133 @p_227))) :rule cong :premises (t364.t0 t364.t3))
(step t364 (cl (= (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v0) f1) (= ?v0 ?v1)))) (forall ((?v0 S3) (?v1 S3)) (=> (= f1 (f3 (f15 f16 ?v0) ?v1)) (=> (= f1 (f3 (f15 f16 ?v1) ?v0)) (= ?v0 ?v1)))))) :rule bind)
(step t365 (cl @p_226) :rule hole :args ("TRUST_THEORY_REWRITE" @p_226 13 6))
(step t366 (cl (= (forall ((?v0 S3) (?v1 S3)) (=> (= (f3 (f15 f16 ?v0) ?v1) f1) (=> (= (f3 (f15 f16 ?v1) ?v0) f1) (= ?v0 ?v1)))) (forall ((?v0 S3) (?v1 S3)) (or (not (= f1 (f3 (f15 f16 ?v0) ?v1))) (not (= f1 (f3 (f15 f16 ?v1) ?v0))) (= ?v0 ?v1))))) :rule trans :premises (t364 t365))
(anchor :step t367 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t367.t0 (cl (= @p_125 @p_125)) :rule refl)
(step t367.t1 (cl (= @p_133 @p_128)) :rule cong :premises (t85 t367.t0))
(step t367.t2 (cl (= @p_225 @p_220)) :rule cong :premises (t367.t1))
(step t367.t3 (cl (= @p_214 @p_214)) :rule refl)
(step t367.t4 (cl (= @p_223 @p_218)) :rule cong :premises (t85 t367.t3))
(step t367.t5 (cl (= @p_224 @p_219)) :rule cong :premises (t367.t4))
(step t367.t6 (cl (= @p_212 @p_212)) :rule refl)
(step t367.t7 (cl (= (or @p_225 @p_224 @p_212) @p_221)) :rule cong :premises (t367.t2 t367.t5 t367.t6))
(step t367 (cl (= (forall ((?v0 S3) (?v1 S3)) (or (not (= f1 (f3 (f15 f16 ?v0) ?v1))) (not (= f1 (f3 (f15 f16 ?v1) ?v0))) (= ?v0 ?v1))) (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1))) (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v1) ?v0))) (= ?v0 ?v1))))) :rule bind)
(anchor :step t368 :args ((?v0 S3) (?v1 S3) (:= (?v0 S3) ?v0) (:= (?v1 S3) ?v1)))
(step t368.t0 (cl @p_131) :rule hole :args ("TRUST_THEORY_REWRITE" @p_131 2 6))
(step t368.t1 (cl (= @p_220 @p_217)) :rule cong :premises (t368.t0))
(step t368.t2 (cl @p_222) :rule hole :args ("TRUST_THEORY_REWRITE" @p_222 2 6))
(step t368.t3 (cl (= @p_219 @p_216)) :rule cong :premises (t368.t2))
(step t368.t4 (cl (= @p_212 @p_212)) :rule refl)
(step t368.t5 (cl (= @p_221 (or @p_217 @p_216 @p_212))) :rule cong :premises (t368.t1 t368.t3 t368.t4))
(step t368 (cl (= (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v0) ?v1))) (not (= (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))) (f3 (f15 f16 ?v1) ?v0))) (= ?v0 ?v1))) (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1))))) :rule bind)
(step t369 (cl (= (forall ((?v0 S3) (?v1 S3)) (or (not (= f1 (f3 (f15 f16 ?v0) ?v1))) (not (= f1 (f3 (f15 f16 ?v1) ?v0))) (= ?v0 ?v1))) (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1))))) :rule trans :premises (t367 t368))
(step t370 (cl @p_211) :rule trans :premises (t366 t369))
(step t371 (cl (forall ((?v0 S3) (?v1 S3)) (or (not (= (f3 (f15 f16 ?v0) ?v1) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (not (= (f3 (f15 f16 ?v1) ?v0) (f3 f4 (f5 (f6 f7 (f5 (f6 f8 (f5 f9 (f5 f10 (f5 f10 (f5 f11 f25))))) f13)) (f5 f9 (f5 f11 f25)))))) (= ?v0 ?v1)))) :rule resolution :premises (t363 t370 a24))
(step t372 (cl @p_210) :rule resolution :premises (t362 t371))
(step t373 (cl @p_206) :rule resolution :premises (t144 t151 t346 t372))
(step t374 (cl) :rule resolution :premises (t40 t142 t373 t317 t314 t200 t54 t151))
