(assume t16 (or (not (= d (f b))) (not (not (= d (f a))))))
(assume t17 (or (or (not (= d (f b))) (not (not (= d (f a))))) (not (not (= d (f b))))))
(assume t18 (or (or (not (= d (f b))) (not (not (= d (f a))))) (not (not (not (= d (f a)))))))
(assume nt19.0 (not (or (not (= d (f b))) (not (not (= d (f a)))))))
(assume nt19.1 (not (or (not (= d (f b))) (not (not (= d (f a)))))))
(step t17' (cl (or (not (= d (f b))) (not (not (= d (f a))))) (not (not (= d (f b))))) :rule or :premises (t17))
(step t16' (cl (not (= d (f b))) (not (not (= d (f a))))) :rule or :premises (t16))
(step t18' (cl (or (not (= d (f b))) (not (not (= d (f a))))) (not (not (not (= d (f a)))))) :rule or :premises (t18))
(step t19 (cl (or (not (= d (f b))) (not (not (= d (f a))))) (or (not (= d (f b))) (not (not (= d (f a)))))) :rule resolution :premises (t16' t17' t18'))
(step t.end (cl) :rule resolution :premises (nt19.0 nt19.1 t19))
