(assume t33 (or (not (not (= d (f b)))) (not (= d (f a)))))
(assume t34 (or (or (not (not (= d (f b)))) (not (= d (f a)))) (not (not (not (= d (f b)))))))
(assume t35 (or (or (not (not (= d (f b)))) (not (= d (f a)))) (not (not (= d (f a))))))
(assume nt36.0 (not (or (not (not (= d (f b)))) (not (= d (f a))))))
(assume nt36.1 (not (or (not (not (= d (f b)))) (not (= d (f a))))))
(step t35' (cl (or (not (not (= d (f b)))) (not (= d (f a)))) (not (not (= d (f a))))) :rule or :premises (t35))
(step t34' (cl (or (not (not (= d (f b)))) (not (= d (f a)))) (not (not (not (= d (f b)))))) :rule or :premises (t34))
(step t33' (cl (not (not (= d (f b)))) (not (= d (f a)))) :rule or :premises (t33))
(step t36 (cl (or (not (not (= d (f b)))) (not (= d (f a)))) (or (not (not (= d (f b)))) (not (= d (f a))))) :rule resolution :premises (t33' t34' t35'))
(step t.end (cl) :rule resolution :premises (nt36.0 nt36.1 t36))
