unsat
(assume a0 (! (not (! (= y$n1s4 y$n15s4) :named @p_1)) :named @p_2))
(assume a1 (! (= y$2 (! (= y$n1s4 y$X) :named @p_3)) :named @p_4))
(assume a2 (! (= (! (= y$n15s4 y$X) :named @p_5) y$6) :named @p_6))
(assume a3 (! (= y$s$1_op (! (Add_4_4_4 y$X y$n1s4) :named @p_7)) :named @p_8))
(assume a4 (! (= y$s$4_op (! (ite y$6 y$n1s4 y$s$1_op) :named @p_9)) :named @p_10))
(assume a5 (! (= y$X$next_rhs_op (! (ite y$en y$s$4_op y$X) :named @p_11)) :named @p_12))
(assume a6 (! (= y$13 (! (= y$X$next y$X$next_rhs_op) :named @p_13)) :named @p_14))
(assume a7 (! (= y$27 (! (not (! (= y$n15s4 y$X$next) :named @p_15)) :named @p_16)) :named @p_17))
(assume a8 (! (= y$28 (! (= y$prop$next y$27) :named @p_18)) :named @p_19))
(assume a9 (! (= y$prop$next (! (not y$26) :named @p_20)) :named @p_21))
(assume a10 (= y$34 (= y$n1s4 y$X$next)))
(assume a11 (= y$s$1$next_op (Add_4_4_4 y$X$next y$n1s4)))
(assume a12 (= y$37 (= y$n15s4 y$s$1$next_op)))
(assume a13 (= y$43 (and y$34 y$37)))
(assume a14 (= y$43 (! (not y$45) :named @p_22)))
(assume a15 (! (= y$33 (! (= y$n15s4 y$s$1_op) :named @p_23)) :named @p_24))
(assume a16 (! (= y$42 (! (and y$2 y$33) :named @p_25)) :named @p_26))
(assume a17 (! (= y$42 (! (not y$44) :named @p_27)) :named @p_28))
(assume a18 (! (= y$58 (! (and y$2 y$13 y$28 y$26 y$45 y$44) :named @p_29)) :named @p_30))
(assume a19 y$58)
(step t0 (cl (! (and (! (= y$n1s4 (! (ite y$en (! (ite (! (= y$n15s4 y$n1s4) :named @p_31) y$n1s4 (! (Add_4_4_4 y$n1s4 y$n1s4) :named @p_33)) :named @p_34) y$n1s4) :named @p_35)) :named @p_101) (! (= y$n15s4 @p_35) :named @p_36)) :named @p_145) (! (not @p_101) :named @p_144) (! (not @p_36) :named @p_53)) :rule and_neg)
(step t1 (cl (! (=> @p_145 @p_31) :named @p_147) @p_145) :rule implies_neg1)
(anchor :step t2)
(assume t2.a0 @p_101)
(assume t2.a1 @p_36)
(step t2.t0 (cl (not (! (= y$58 @p_36) :named @p_37)) (! (not y$58) :named @p_69) @p_36) :rule equiv_pos2)
(step t2.t1 (cl (not (! (= @p_30 (! (= y$58 (! (and (! (= y$X y$n1s4) :named @p_65) (! (= y$X$next (! (ite y$en (! (ite (! (= y$X y$n15s4) :named @p_61) y$n1s4 @p_7) :named @p_62) y$X) :named @p_63)) :named @p_64) (! (= @p_20 (! (not (! (= y$X$next y$n15s4) :named @p_58)) :named @p_59)) :named @p_60) y$26 y$45 y$44) :named @p_66)) :named @p_89)) :named @p_90)) (! (not @p_30) :named @p_68) @p_89) :rule equiv_pos2)
(step t2.t2 (cl (= y$58 y$58)) :rule refl)
(step t2.t3 (cl (not (! (= @p_4 (! (= y$2 @p_65) :named @p_78)) :named @p_79)) (not @p_4) @p_78) :rule equiv_pos2)
(step t2.t4 (cl (= y$2 y$2)) :rule refl)
(step t2.t5 (cl (= @p_3 @p_65)) :rule rare_rewrite :args ("eq-symm" y$n1s4 y$X))
(step t2.t6 (cl @p_79) :rule cong :premises (t2.t4 t2.t5))
(step t2.t7 (cl @p_78) :rule resolution :premises (t2.t3 t2.t6 a1))
(step t2.t8 (cl (not (! (= @p_14 (! (= y$13 @p_64) :named @p_80)) :named @p_81)) (not @p_14) @p_80) :rule equiv_pos2)
(step t2.t9 (cl (= y$13 y$13)) :rule refl)
(step t2.t10 (cl (= y$X$next y$X$next)) :rule refl)
(step t2.t11 (cl (not (! (= @p_12 (! (= y$X$next_rhs_op @p_63) :named @p_82)) :named @p_83)) (not @p_12) @p_82) :rule equiv_pos2)
(step t2.t12 (cl (= y$X$next_rhs_op y$X$next_rhs_op)) :rule refl)
(step t2.t13 (cl (= y$en y$en)) :rule refl)
(step t2.t14 (cl (not (! (= @p_10 (! (= y$s$4_op @p_62) :named @p_84)) :named @p_85)) (not @p_10) @p_84) :rule equiv_pos2)
(step t2.t15 (cl (= y$s$4_op y$s$4_op)) :rule refl)
(step t2.t16 (cl (not (! (= @p_6 (! (= y$6 @p_61) :named @p_86)) :named @p_87)) (not @p_6) @p_86) :rule equiv_pos2)
(step t2.t17 (cl (= @p_6 (! (= y$6 @p_5) :named @p_88))) :rule rare_rewrite :args ("eq-symm" @p_5 y$6))
(step t2.t18 (cl (= y$6 y$6)) :rule refl)
(step t2.t19 (cl (= @p_5 @p_61)) :rule rare_rewrite :args ("eq-symm" y$n15s4 y$X))
(step t2.t20 (cl (= @p_88 @p_86)) :rule cong :premises (t2.t18 t2.t19))
(step t2.t21 (cl @p_87) :rule trans :premises (t2.t17 t2.t20))
(step t2.t22 (cl @p_86) :rule resolution :premises (t2.t16 t2.t21 a2))
(step t2.t23 (cl (! (= y$n1s4 y$n1s4) :named @p_48)) :rule refl)
(step t2.t24 (cl (= @p_9 @p_62)) :rule cong :premises (t2.t22 t2.t23 a3))
(step t2.t25 (cl @p_85) :rule cong :premises (t2.t15 t2.t24))
(step t2.t26 (cl @p_84) :rule resolution :premises (t2.t14 t2.t25 a4))
(step t2.t27 (cl (= y$X y$X)) :rule refl)
(step t2.t28 (cl (= @p_11 @p_63)) :rule cong :premises (t2.t13 t2.t26 t2.t27))
(step t2.t29 (cl @p_83) :rule cong :premises (t2.t12 t2.t28))
(step t2.t30 (cl @p_82) :rule resolution :premises (t2.t11 t2.t29 a5))
(step t2.t31 (cl (= @p_13 @p_64)) :rule cong :premises (t2.t10 t2.t30))
(step t2.t32 (cl @p_81) :rule cong :premises (t2.t9 t2.t31))
(step t2.t33 (cl @p_80) :rule resolution :premises (t2.t8 t2.t32 a6))
(step t2.t34 (cl (not (! (= (! (= y$28 (! (= y$27 y$prop$next) :named @p_93)) :named @p_94) (! (= y$28 (! (= y$prop$next @p_59) :named @p_91)) :named @p_92)) :named @p_96)) (not @p_94) @p_92) :rule equiv_pos2)
(step t2.t35 (cl (= y$28 y$28)) :rule refl)
(step t2.t36 (cl (not (! (= @p_17 (! (= y$27 @p_59) :named @p_99)) :named @p_100)) (not @p_17) @p_99) :rule equiv_pos2)
(step t2.t37 (cl (= y$27 y$27)) :rule refl)
(step t2.t38 (cl (= @p_15 @p_58)) :rule rare_rewrite :args ("eq-symm" y$n15s4 y$X$next))
(step t2.t39 (cl (= @p_16 @p_59)) :rule cong :premises (t2.t38))
(step t2.t40 (cl @p_100) :rule cong :premises (t2.t37 t2.t39))
(step t2.t41 (cl @p_99) :rule resolution :premises (t2.t36 t2.t40 a7))
(step t2.t42 (cl (= y$prop$next y$prop$next)) :rule refl)
(step t2.t43 (cl (= @p_93 (! (= @p_59 y$prop$next) :named @p_97))) :rule cong :premises (t2.t41 t2.t42))
(step t2.t44 (cl (= @p_94 (! (= y$28 @p_97) :named @p_98))) :rule cong :premises (t2.t35 t2.t43))
(step t2.t45 (cl (= @p_97 @p_91)) :rule rare_rewrite :args ("eq-symm" @p_59 y$prop$next))
(step t2.t46 (cl (= @p_98 @p_92)) :rule cong :premises (t2.t35 t2.t45))
(step t2.t47 (cl @p_96) :rule trans :premises (t2.t44 t2.t46))
(step t2.t48 (cl (not (! (= @p_19 @p_94) :named @p_95)) (not @p_19) @p_94) :rule equiv_pos2)
(step t2.t49 (cl (= @p_18 @p_93)) :rule rare_rewrite :args ("eq-symm" y$prop$next y$27))
(step t2.t50 (cl @p_95) :rule cong :premises (t2.t35 t2.t49))
(step t2.t51 (cl @p_94) :rule resolution :premises (t2.t48 t2.t50 a8))
(step t2.t52 (cl @p_92) :rule resolution :premises (t2.t34 t2.t47 t2.t51))
(step t2.t53 (cl (= @p_59 @p_59)) :rule refl)
(step t2.t54 (cl (= @p_91 @p_60)) :rule cong :premises (a9 t2.t53))
(step t2.t55 (cl (= y$28 @p_60)) :rule trans :premises (t2.t52 t2.t54))
(step t2.t56 (cl (= y$26 y$26)) :rule refl)
(step t2.t57 (cl (= y$45 y$45)) :rule refl)
(step t2.t58 (cl (= y$44 y$44)) :rule refl)
(step t2.t59 (cl (= @p_29 @p_66)) :rule cong :premises (t2.t7 t2.t33 t2.t55 t2.t56 t2.t57 t2.t58))
(step t2.t60 (cl @p_90) :rule cong :premises (t2.t2 t2.t59))
(step t2.t61 (cl @p_89) :rule resolution :premises (t2.t1 t2.t60 a18))
(step t2.t62 (cl (not @p_78) (not y$2) @p_65) :rule equiv_pos2)
(step t2.t63 (cl @p_68 @p_69 @p_29) :rule equiv_pos2)
(step t2.t64 (cl @p_29) :rule resolution :premises (t2.t63 a18 a19))
(step t2.t65 (cl y$2) :rule and :premises (t2.t64) :args (0))
(step t2.t66 (cl @p_65) :rule resolution :premises (t2.t62 t2.t7 t2.t65))
(step t2.t67 (cl (= @p_65 @p_48)) :rule cong :premises (t2.t66 t2.t23))
(step t2.t68 (cl (not (! (= y$13 (! (= y$X$next @p_35) :named @p_75)) :named @p_76)) (not y$13) @p_75) :rule equiv_pos2)
(step t2.t69 (cl (= y$n15s4 y$n15s4)) :rule refl)
(step t2.t70 (cl (= @p_61 @p_1)) :rule cong :premises (t2.t66 t2.t69))
(step t2.t71 (cl (= @p_7 @p_33)) :rule cong :premises (t2.t66 t2.t23))
(step t2.t72 (cl (= @p_62 (! (ite @p_1 y$n1s4 @p_33) :named @p_45))) :rule cong :premises (t2.t70 t2.t23 t2.t71))
(step t2.t73 (cl (= @p_63 (! (ite y$en @p_45 y$n1s4) :named @p_46))) :rule cong :premises (t2.t13 t2.t72 t2.t66))
(step t2.t74 (cl (= @p_64 (! (= y$X$next @p_46) :named @p_77))) :rule cong :premises (t2.t10 t2.t73))
(step t2.t75 (cl (= y$13 @p_77)) :rule trans :premises (t2.t33 t2.t74))
(step t2.t76 (cl (= @p_1 @p_31)) :rule rare_rewrite :args ("eq-symm" y$n1s4 y$n15s4))
(step t2.t77 (cl (= @p_33 @p_33)) :rule refl)
(step t2.t78 (cl (= @p_45 @p_34)) :rule cong :premises (t2.t76 t2.t23 t2.t77))
(step t2.t79 (cl (= @p_46 @p_35)) :rule cong :premises (t2.t13 t2.t78 t2.t23))
(step t2.t80 (cl (= @p_77 @p_75)) :rule cong :premises (t2.t10 t2.t79))
(step t2.t81 (cl @p_76) :rule trans :premises (t2.t75 t2.t80))
(step t2.t82 (cl y$13) :rule and :premises (t2.t64) :args (1))
(step t2.t83 (cl @p_75) :rule resolution :premises (t2.t68 t2.t81 t2.t82))
(step t2.t84 (cl (= @p_64 (! (= @p_35 @p_46) :named @p_47))) :rule cong :premises (t2.t83 t2.t73))
(step t2.t85 (cl (not (! (= y$26 (! (= y$26 true) :named @p_73)) :named @p_74)) @p_20 @p_73) :rule equiv_pos2)
(step t2.t86 (cl (= @p_73 y$26)) :rule rare_rewrite :args ("bool-eq-true" y$26))
(step t2.t87 (cl @p_74) :rule symm :premises (t2.t86))
(step t2.t88 (cl y$26) :rule and :premises (t2.t64) :args (3))
(step t2.t89 (cl @p_73) :rule resolution :premises (t2.t85 t2.t87 t2.t88))
(step t2.t90 (cl (= @p_20 (! (not true) :named @p_43))) :rule cong :premises (t2.t89))
(step t2.t91 (cl (= @p_58 (! (= @p_35 y$n15s4) :named @p_41))) :rule cong :premises (t2.t83 t2.t69))
(step t2.t92 (cl (= @p_59 (! (not @p_41) :named @p_42))) :rule cong :premises (t2.t91))
(step t2.t93 (cl (= @p_60 (! (= @p_43 @p_42) :named @p_44))) :rule cong :premises (t2.t90 t2.t92))
(step t2.t94 (cl (not (! (= y$45 (! (= y$45 true) :named @p_71)) :named @p_72)) @p_22 @p_71) :rule equiv_pos2)
(step t2.t95 (cl (= @p_71 y$45)) :rule rare_rewrite :args ("bool-eq-true" y$45))
(step t2.t96 (cl @p_72) :rule symm :premises (t2.t95))
(step t2.t97 (cl y$45) :rule and :premises (t2.t64) :args (4))
(step t2.t98 (cl @p_71) :rule resolution :premises (t2.t94 t2.t96 t2.t97))
(step t2.t99 (cl (not (! (= y$44 (! (= y$44 true) :named @p_67)) :named @p_70)) @p_27 @p_67) :rule equiv_pos2)
(step t2.t100 (cl (= @p_67 y$44)) :rule rare_rewrite :args ("bool-eq-true" y$44))
(step t2.t101 (cl @p_70) :rule symm :premises (t2.t100))
(step t2.t102 (cl y$44) :rule and :premises (t2.t64) :args (5))
(step t2.t103 (cl @p_67) :rule resolution :premises (t2.t99 t2.t101 t2.t102))
(step t2.t104 (cl (= @p_66 (! (and @p_48 @p_47 @p_44 true true true) :named @p_49))) :rule cong :premises (t2.t67 t2.t84 t2.t93 t2.t89 t2.t98 t2.t103))
(step t2.t105 (cl (= y$58 @p_49)) :rule trans :premises (t2.t61 t2.t104))
(step t2.t106 (cl (= @p_49 (! (and @p_48 @p_47 @p_44 true) :named @p_57))) :rule ac_simp)
(step t2.t107 (cl (= @p_57 (! (and @p_48 @p_47 @p_44) :named @p_50))) :rule and_simplify)
(step t2.t108 (cl (= @p_49 @p_50)) :rule trans :premises (t2.t106 t2.t107))
(step t2.t109 (cl (= @p_48 true)) :rule rare_rewrite :args ("eq-refl" y$n1s4))
(step t2.t110 (cl (! (= @p_35 @p_35) :named @p_38)) :rule refl)
(step t2.t111 (cl (= @p_47 @p_38)) :rule cong :premises (t2.t110 t2.t79))
(step t2.t112 (cl (! (= @p_38 true) :named @p_39)) :rule rare_rewrite :args ("eq-refl" @p_35))
(step t2.t113 (cl (= @p_47 true)) :rule trans :premises (t2.t111 t2.t112))
(step t2.t114 (cl (= @p_43 false)) :rule rare_rewrite :args ("evaluate"))
(step t2.t115 (cl (= @p_41 @p_36)) :rule rare_rewrite :args ("eq-symm" @p_35 y$n15s4))
(step t2.t116 (cl (= @p_42 @p_53)) :rule cong :premises (t2.t115))
(step t2.t117 (cl (= @p_44 (! (= false @p_53) :named @p_54))) :rule cong :premises (t2.t114 t2.t116))
(step t2.t118 (cl (= @p_54 (! (= @p_53 false) :named @p_55))) :rule rare_rewrite :args ("eq-symm" false @p_53))
(step t2.t119 (cl (= @p_55 (! (not @p_53) :named @p_56))) :rule rare_rewrite :args ("bool-eq-false" @p_53))
(step t2.t120 (cl (= @p_56 @p_36)) :rule rare_rewrite :args ("bool-double-not-elim" @p_36))
(step t2.t121 (cl (= @p_55 @p_36)) :rule trans :premises (t2.t119 t2.t120))
(step t2.t122 (cl (= @p_54 @p_36)) :rule trans :premises (t2.t118 t2.t121))
(step t2.t123 (cl (= @p_44 @p_36)) :rule trans :premises (t2.t117 t2.t122))
(step t2.t124 (cl (= @p_50 (! (and true true @p_36) :named @p_51))) :rule cong :premises (t2.t109 t2.t113 t2.t123))
(step t2.t125 (cl (= @p_51 (! (and true @p_36) :named @p_52))) :rule ac_simp)
(step t2.t126 (cl (= @p_52 @p_36)) :rule and_simplify)
(step t2.t127 (cl (= @p_51 @p_36)) :rule trans :premises (t2.t125 t2.t126))
(step t2.t128 (cl (= @p_50 @p_36)) :rule trans :premises (t2.t124 t2.t127))
(step t2.t129 (cl (= @p_49 @p_36)) :rule trans :premises (t2.t108 t2.t128))
(step t2.t130 (cl @p_37) :rule trans :premises (t2.t105 t2.t129))
(step t2.t131 (cl (= @p_39 @p_38)) :rule equiv_simplify)
(step t2.t132 (cl (not @p_39) @p_38) :rule equiv1 :premises (t2.t131))
(step t2.t133 (cl (! (= @p_38 @p_38) :named @p_40)) :rule refl)
(step t2.t134 (cl @p_38) :rule refl)
(step t2.t135 (cl @p_40) :rule cong :premises (t2.t110 t2.t134))
(step t2.t136 (cl @p_39) :rule trans :premises (t2.t133 t2.t135 t2.t112))
(step t2.t137 (cl @p_38) :rule resolution :premises (t2.t132 t2.t136))
(step t2.t138 (cl (= @p_36 @p_36)) :rule cong :premises (t2.t69 t2.t137))
(step t2.t139 (cl @p_37) :rule trans :premises (t2.t130 t2.t138))
(step t2.t140 (cl @p_36) :rule resolution :premises (t2.t0 t2.t139 a19))
(step t2.t141 (cl (not (! (= (! (ite y$en (! (= @p_35 @p_34) :named @p_103) (! (= @p_35 y$n1s4) :named @p_136)) :named @p_137) (! (ite y$en @p_103 @p_101) :named @p_135)) :named @p_140)) (! (not @p_137) :named @p_139) @p_135) :rule equiv_pos2)
(step t2.t142 (cl (= @p_103 (! (= @p_34 @p_35) :named @p_141))) :rule rare_rewrite :args ("eq-symm" @p_35 @p_34))
(step t2.t143 (cl (= @p_136 @p_101)) :rule rare_rewrite :args ("eq-symm" @p_35 y$n1s4))
(step t2.t144 (cl (= @p_137 (! (ite y$en @p_141 @p_101) :named @p_142))) :rule cong :premises (t2.t13 t2.t142 t2.t143))
(step t2.t145 (cl (= (! (= (! (= @p_34 @p_34) :named @p_130) true) :named @p_131) @p_130)) :rule equiv_simplify)
(step t2.t146 (cl (not @p_131) @p_130) :rule equiv1 :premises (t2.t145))
(step t2.t147 (cl (! (= @p_130 @p_130) :named @p_132)) :rule refl)
(step t2.t148 (cl @p_130) :rule refl)
(step t2.t149 (cl @p_130) :rule refl)
(step t2.t150 (cl @p_132) :rule cong :premises (t2.t148 t2.t149))
(step t2.t151 (cl @p_131) :rule rare_rewrite :args ("eq-refl" @p_34))
(step t2.t152 (cl @p_131) :rule trans :premises (t2.t147 t2.t150 t2.t151))
(step t2.t153 (cl @p_130) :rule resolution :premises (t2.t146 t2.t152))
(step t2.t154 (cl @p_38) :rule refl)
(step t2.t155 (cl (= @p_141 @p_141)) :rule cong :premises (t2.t153 t2.t154))
(step t2.t156 (cl (! (= @p_141 @p_103) :named @p_143)) :rule rare_rewrite :args ("eq-symm" @p_34 @p_35))
(step t2.t157 (cl @p_143) :rule trans :premises (t2.t155 t2.t156))
(step t2.t158 (cl (= @p_101 @p_101)) :rule refl)
(step t2.t159 (cl (= @p_142 @p_135)) :rule cong :premises (t2.t13 t2.t157 t2.t158))
(step t2.t160 (cl @p_140) :rule trans :premises (t2.t144 t2.t159))
(step t2.t161 (cl (not (! (= @p_137 @p_137) :named @p_138)) @p_139 @p_137) :rule equiv_pos2)
(step t2.t162 (cl (= @p_103 @p_103)) :rule cong :premises (t2.t137 t2.t148))
(step t2.t163 (cl (= @p_136 @p_136)) :rule cong :premises (t2.t137 t2.t23))
(step t2.t164 (cl @p_138) :rule cong :premises (t2.t13 t2.t162 t2.t163))
(step t2.t165 (cl @p_137) :rule rare_rewrite :args ("ite-eq" y$en @p_34 y$n1s4))
(step t2.t166 (cl @p_137) :rule resolution :premises (t2.t161 t2.t164 t2.t165))
(step t2.t167 (cl @p_135) :rule resolution :premises (t2.t141 t2.t160 t2.t166))
(step t2.t168 (cl y$en @p_101) :rule ite1 :premises (t2.t167))
(step t2.t169 (cl (! (not y$en) :named @p_102) @p_103) :rule ite2 :premises (t2.t167))
(step t2.t170 (cl @p_103 @p_102) :rule reordering :premises (t2.t169))
(step t2.t171 (cl (= (! (= @p_103 false) :named @p_105) (! (not @p_103) :named @p_104))) :rule equiv_simplify)
(step t2.t172 (cl (not @p_105) @p_104) :rule equiv1 :premises (t2.t171))
(step t2.t173 (cl @p_41) :rule symm :premises (t2.t140))
(step t2.t174 (cl (not (! (= (! (ite @p_31 (! (= @p_34 y$n1s4) :named @p_127) (! (= @p_34 @p_33) :named @p_123)) :named @p_128) (! (ite @p_31 (! (= y$n1s4 @p_34) :named @p_125) (! (= @p_33 @p_34) :named @p_124)) :named @p_126)) :named @p_134)) (! (not @p_128) :named @p_133) @p_126) :rule equiv_pos2)
(step t2.t175 (cl (= @p_31 @p_31)) :rule refl)
(step t2.t176 (cl (= @p_127 @p_125)) :rule rare_rewrite :args ("eq-symm" @p_34 y$n1s4))
(step t2.t177 (cl (= @p_123 @p_124)) :rule rare_rewrite :args ("eq-symm" @p_34 @p_33))
(step t2.t178 (cl @p_134) :rule cong :premises (t2.t175 t2.t176 t2.t177))
(step t2.t179 (cl (not (! (= @p_128 @p_128) :named @p_129)) @p_133 @p_128) :rule equiv_pos2)
(step t2.t180 (cl (= @p_127 @p_127)) :rule cong :premises (t2.t153 t2.t23))
(step t2.t181 (cl (= @p_123 @p_123)) :rule cong :premises (t2.t153 t2.t77))
(step t2.t182 (cl @p_129) :rule cong :premises (t2.t175 t2.t180 t2.t181))
(step t2.t183 (cl @p_128) :rule rare_rewrite :args ("ite-eq" @p_31 y$n1s4 @p_33))
(step t2.t184 (cl @p_128) :rule resolution :premises (t2.t179 t2.t182 t2.t183))
(step t2.t185 (cl @p_126) :rule resolution :premises (t2.t174 t2.t178 t2.t184))
(step t2.t186 (cl @p_31 @p_124) :rule ite1 :premises (t2.t185))
(step t2.t187 (cl (! (not @p_31) :named @p_32)) :rule not_symm :premises (a0))
(step t2.t188 (cl @p_124) :rule resolution :premises (t2.t186 t2.t187))
(step t2.t189 (cl @p_123) :rule symm :premises (t2.t188))
(step t2.t190 (cl (= @p_103 (! (= y$n15s4 @p_33) :named @p_106))) :rule cong :premises (t2.t173 t2.t189))
(step t2.t191 (cl (! (= (! (= @p_106 false) :named @p_107) (! (not @p_106) :named @p_108)) :named @p_114)) :rule equiv_simplify)
(step t2.t192 (cl @p_107 (! (not @p_108) :named @p_122)) :rule equiv2 :premises (t2.t191))
(step t2.t193 (cl (not @p_122) @p_106) :rule not_not)
(step t2.t194 (cl @p_107 @p_106) :rule resolution :premises (t2.t192 t2.t193))
(step t2.t195 (cl (not (! (= @p_28 @p_108) :named @p_109)) (not @p_28) @p_108) :rule equiv_pos2)
(step t2.t196 (cl (not (! (= @p_26 (! (= y$42 (! (and @p_65 (! (= y$n15s4 @p_7) :named @p_116)) :named @p_117)) :named @p_118)) :named @p_119)) (not @p_26) @p_118) :rule equiv_pos2)
(step t2.t197 (cl (= y$42 y$42)) :rule refl)
(step t2.t198 (cl (not (! (= @p_24 (! (= y$33 @p_116) :named @p_120)) :named @p_121)) (not @p_24) @p_120) :rule equiv_pos2)
(step t2.t199 (cl (= y$33 y$33)) :rule refl)
(step t2.t200 (cl (= @p_23 @p_116)) :rule cong :premises (t2.t69 a3))
(step t2.t201 (cl @p_121) :rule cong :premises (t2.t199 t2.t200))
(step t2.t202 (cl @p_120) :rule resolution :premises (t2.t198 t2.t201 a15))
(step t2.t203 (cl (= @p_25 @p_117)) :rule cong :premises (t2.t7 t2.t202))
(step t2.t204 (cl @p_119) :rule cong :premises (t2.t197 t2.t203))
(step t2.t205 (cl @p_118) :rule resolution :premises (t2.t196 t2.t204 a16))
(step t2.t206 (cl (= @p_116 @p_106)) :rule cong :premises (t2.t69 t2.t71))
(step t2.t207 (cl (= @p_117 (! (and @p_48 @p_106) :named @p_110))) :rule cong :premises (t2.t67 t2.t206))
(step t2.t208 (cl (= y$42 @p_110)) :rule trans :premises (t2.t205 t2.t207))
(step t2.t209 (cl (= @p_27 @p_43)) :rule cong :premises (t2.t103))
(step t2.t210 (cl (= @p_28 (! (= @p_110 @p_43) :named @p_111))) :rule cong :premises (t2.t208 t2.t209))
(step t2.t211 (cl (= @p_111 (! (= @p_43 @p_110) :named @p_112))) :rule rare_rewrite :args ("eq-symm" @p_110 @p_43))
(step t2.t212 (cl (= @p_106 @p_106)) :rule refl)
(step t2.t213 (cl (= @p_110 (! (and true @p_106) :named @p_115))) :rule cong :premises (t2.t109 t2.t212))
(step t2.t214 (cl (= @p_115 @p_106)) :rule and_simplify)
(step t2.t215 (cl (= @p_110 @p_106)) :rule trans :premises (t2.t213 t2.t214))
(step t2.t216 (cl (= @p_112 (! (= false @p_106) :named @p_113))) :rule cong :premises (t2.t114 t2.t215))
(step t2.t217 (cl (= @p_113 @p_107)) :rule rare_rewrite :args ("eq-symm" false @p_106))
(step t2.t218 (cl @p_114) :rule rare_rewrite :args ("bool-eq-false" @p_106))
(step t2.t219 (cl (= @p_113 @p_108)) :rule trans :premises (t2.t217 t2.t218))
(step t2.t220 (cl (= @p_112 @p_108)) :rule trans :premises (t2.t216 t2.t219))
(step t2.t221 (cl (= @p_111 @p_108)) :rule trans :premises (t2.t211 t2.t220))
(step t2.t222 (cl @p_109) :rule trans :premises (t2.t210 t2.t221))
(step t2.t223 (cl @p_108) :rule resolution :premises (t2.t195 t2.t222 a17))
(step t2.t224 (cl @p_107) :rule resolution :premises (t2.t194 t2.t223))
(step t2.t225 (cl @p_105) :rule trans :premises (t2.t190 t2.t224))
(step t2.t226 (cl @p_104) :rule resolution :premises (t2.t172 t2.t225))
(step t2.t227 (cl @p_102) :rule resolution :premises (t2.t170 t2.t226))
(step t2.t228 (cl @p_101) :rule resolution :premises (t2.t168 t2.t227))
(step t2.t229 (cl @p_136) :rule symm :premises (t2.t228))
(step t2.t230 (cl @p_31) :rule trans :premises (t2.t140 t2.t229))
(step t2 (cl @p_144 @p_53 @p_31) :rule subproof :discharge (t2.a0 t2.a1))
(step t3 (cl (! (not @p_145) :named @p_146) @p_101) :rule and_pos :args (0))
(step t4 (cl @p_146 @p_36) :rule and_pos :args (1))
(step t5 (cl @p_31 @p_146 @p_146) :rule resolution :premises (t2 t3 t4))
(step t6 (cl @p_146 @p_146 @p_31) :rule reordering :premises (t5))
(step t7 (cl @p_146 @p_31) :rule contraction :premises (t6))
(step t8 (cl @p_147 @p_31) :rule resolution :premises (t1 t7))
(step t9 (cl @p_147 @p_32) :rule implies_neg2)
(step t10 (cl @p_147 @p_147) :rule resolution :premises (t8 t9))
(step t11 (cl @p_147) :rule contraction :premises (t10))
(step t12 (cl @p_146 @p_31) :rule implies :premises (t11))
(step t13 (cl @p_144 @p_53 @p_31) :rule resolution :premises (t0 t12))
(step t14 (cl @p_31 @p_144 @p_53) :rule reordering :premises (t13))
(step t15 (cl (not @p_140) @p_139 @p_135) :rule equiv_pos2)
(step t16 (cl (= y$en y$en)) :rule refl)
(step t17 (cl (= @p_103 @p_141)) :rule rare_rewrite :args ("eq-symm" @p_35 @p_34))
(step t18 (cl (= @p_136 @p_101)) :rule rare_rewrite :args ("eq-symm" @p_35 y$n1s4))
(step t19 (cl (= @p_137 @p_142)) :rule cong :premises (t16 t17 t18))
(step t20 (cl (= @p_131 @p_130)) :rule equiv_simplify)
(step t21 (cl (not @p_131) @p_130) :rule equiv1 :premises (t20))
(step t22 (cl @p_132) :rule refl)
(step t23 (cl @p_130) :rule refl)
(step t24 (cl @p_130) :rule refl)
(step t25 (cl @p_132) :rule cong :premises (t23 t24))
(step t26 (cl @p_131) :rule rare_rewrite :args ("eq-refl" @p_34))
(step t27 (cl @p_131) :rule trans :premises (t22 t25 t26))
(step t28 (cl @p_130) :rule resolution :premises (t21 t27))
(step t29 (cl @p_38) :rule refl)
(step t30 (cl (= @p_141 @p_141)) :rule cong :premises (t28 t29))
(step t31 (cl @p_143) :rule rare_rewrite :args ("eq-symm" @p_34 @p_35))
(step t32 (cl @p_143) :rule trans :premises (t30 t31))
(step t33 (cl (= @p_101 @p_101)) :rule refl)
(step t34 (cl (= @p_142 @p_135)) :rule cong :premises (t16 t32 t33))
(step t35 (cl @p_140) :rule trans :premises (t19 t34))
(step t36 (cl (not @p_138) @p_139 @p_137) :rule equiv_pos2)
(step t37 (cl (= @p_39 @p_38)) :rule equiv_simplify)
(step t38 (cl (not @p_39) @p_38) :rule equiv1 :premises (t37))
(step t39 (cl @p_40) :rule refl)
(step t40 (cl @p_38) :rule refl)
(step t41 (cl @p_38) :rule refl)
(step t42 (cl @p_40) :rule cong :premises (t40 t41))
(step t43 (cl @p_39) :rule rare_rewrite :args ("eq-refl" @p_35))
(step t44 (cl @p_39) :rule trans :premises (t39 t42 t43))
(step t45 (cl @p_38) :rule resolution :premises (t38 t44))
(step t46 (cl (= @p_103 @p_103)) :rule cong :premises (t45 t23))
(step t47 (cl @p_48) :rule refl)
(step t48 (cl (= @p_136 @p_136)) :rule cong :premises (t45 t47))
(step t49 (cl @p_138) :rule cong :premises (t16 t46 t48))
(step t50 (cl @p_137) :rule rare_rewrite :args ("ite-eq" y$en @p_34 y$n1s4))
(step t51 (cl @p_137) :rule resolution :premises (t36 t49 t50))
(step t52 (cl @p_135) :rule resolution :premises (t15 t35 t51))
(step t53 (cl y$en @p_101) :rule ite1 :premises (t52))
(step t54 (cl @p_102 @p_103) :rule ite2 :premises (t52))
(step t55 (cl @p_103 @p_102) :rule reordering :premises (t54))
(step t56 (cl (= @p_105 @p_104)) :rule equiv_simplify)
(step t57 (cl (not @p_105) @p_104) :rule equiv1 :premises (t56))
(step t58 (cl (not @p_37) @p_69 @p_36) :rule equiv_pos2)
(step t59 (cl (not @p_90) @p_68 @p_89) :rule equiv_pos2)
(step t60 (cl (= y$58 y$58)) :rule refl)
(step t61 (cl (not @p_79) (not @p_4) @p_78) :rule equiv_pos2)
(step t62 (cl (= y$2 y$2)) :rule refl)
(step t63 (cl (= @p_3 @p_65)) :rule rare_rewrite :args ("eq-symm" y$n1s4 y$X))
(step t64 (cl @p_79) :rule cong :premises (t62 t63))
(step t65 (cl @p_78) :rule resolution :premises (t61 t64 a1))
(step t66 (cl (not @p_81) (not @p_14) @p_80) :rule equiv_pos2)
(step t67 (cl (= y$13 y$13)) :rule refl)
(step t68 (cl (= y$X$next y$X$next)) :rule refl)
(step t69 (cl (not @p_83) (not @p_12) @p_82) :rule equiv_pos2)
(step t70 (cl (= y$X$next_rhs_op y$X$next_rhs_op)) :rule refl)
(step t71 (cl (not @p_85) (not @p_10) @p_84) :rule equiv_pos2)
(step t72 (cl (= y$s$4_op y$s$4_op)) :rule refl)
(step t73 (cl (not @p_87) (not @p_6) @p_86) :rule equiv_pos2)
(step t74 (cl (= @p_6 @p_88)) :rule rare_rewrite :args ("eq-symm" @p_5 y$6))
(step t75 (cl (= y$6 y$6)) :rule refl)
(step t76 (cl (= @p_5 @p_61)) :rule rare_rewrite :args ("eq-symm" y$n15s4 y$X))
(step t77 (cl (= @p_88 @p_86)) :rule cong :premises (t75 t76))
(step t78 (cl @p_87) :rule trans :premises (t74 t77))
(step t79 (cl @p_86) :rule resolution :premises (t73 t78 a2))
(step t80 (cl (= @p_9 @p_62)) :rule cong :premises (t79 t47 a3))
(step t81 (cl @p_85) :rule cong :premises (t72 t80))
(step t82 (cl @p_84) :rule resolution :premises (t71 t81 a4))
(step t83 (cl (= y$X y$X)) :rule refl)
(step t84 (cl (= @p_11 @p_63)) :rule cong :premises (t16 t82 t83))
(step t85 (cl @p_83) :rule cong :premises (t70 t84))
(step t86 (cl @p_82) :rule resolution :premises (t69 t85 a5))
(step t87 (cl (= @p_13 @p_64)) :rule cong :premises (t68 t86))
(step t88 (cl @p_81) :rule cong :premises (t67 t87))
(step t89 (cl @p_80) :rule resolution :premises (t66 t88 a6))
(step t90 (cl (not @p_96) (not @p_94) @p_92) :rule equiv_pos2)
(step t91 (cl (= y$28 y$28)) :rule refl)
(step t92 (cl (not @p_100) (not @p_17) @p_99) :rule equiv_pos2)
(step t93 (cl (= y$27 y$27)) :rule refl)
(step t94 (cl (= @p_15 @p_58)) :rule rare_rewrite :args ("eq-symm" y$n15s4 y$X$next))
(step t95 (cl (= @p_16 @p_59)) :rule cong :premises (t94))
(step t96 (cl @p_100) :rule cong :premises (t93 t95))
(step t97 (cl @p_99) :rule resolution :premises (t92 t96 a7))
(step t98 (cl (= y$prop$next y$prop$next)) :rule refl)
(step t99 (cl (= @p_93 @p_97)) :rule cong :premises (t97 t98))
(step t100 (cl (= @p_94 @p_98)) :rule cong :premises (t91 t99))
(step t101 (cl (= @p_97 @p_91)) :rule rare_rewrite :args ("eq-symm" @p_59 y$prop$next))
(step t102 (cl (= @p_98 @p_92)) :rule cong :premises (t91 t101))
(step t103 (cl @p_96) :rule trans :premises (t100 t102))
(step t104 (cl (not @p_95) (not @p_19) @p_94) :rule equiv_pos2)
(step t105 (cl (= @p_18 @p_93)) :rule rare_rewrite :args ("eq-symm" y$prop$next y$27))
(step t106 (cl @p_95) :rule cong :premises (t91 t105))
(step t107 (cl @p_94) :rule resolution :premises (t104 t106 a8))
(step t108 (cl @p_92) :rule resolution :premises (t90 t103 t107))
(step t109 (cl (= @p_59 @p_59)) :rule refl)
(step t110 (cl (= @p_91 @p_60)) :rule cong :premises (a9 t109))
(step t111 (cl (= y$28 @p_60)) :rule trans :premises (t108 t110))
(step t112 (cl (= y$26 y$26)) :rule refl)
(step t113 (cl (= y$45 y$45)) :rule refl)
(step t114 (cl (= y$44 y$44)) :rule refl)
(step t115 (cl (= @p_29 @p_66)) :rule cong :premises (t65 t89 t111 t112 t113 t114))
(step t116 (cl @p_90) :rule cong :premises (t60 t115))
(step t117 (cl @p_89) :rule resolution :premises (t59 t116 a18))
(step t118 (cl (not @p_78) (not y$2) @p_65) :rule equiv_pos2)
(step t119 (cl @p_68 @p_69 @p_29) :rule equiv_pos2)
(step t120 (cl @p_29) :rule resolution :premises (t119 a18 a19))
(step t121 (cl y$2) :rule and :premises (t120) :args (0))
(step t122 (cl @p_65) :rule resolution :premises (t118 t65 t121))
(step t123 (cl (= @p_65 @p_48)) :rule cong :premises (t122 t47))
(step t124 (cl (not @p_76) (not y$13) @p_75) :rule equiv_pos2)
(step t125 (cl (= y$n15s4 y$n15s4)) :rule refl)
(step t126 (cl (= @p_61 @p_1)) :rule cong :premises (t122 t125))
(step t127 (cl (= @p_7 @p_33)) :rule cong :premises (t122 t47))
(step t128 (cl (= @p_62 @p_45)) :rule cong :premises (t126 t47 t127))
(step t129 (cl (= @p_63 @p_46)) :rule cong :premises (t16 t128 t122))
(step t130 (cl (= @p_64 @p_77)) :rule cong :premises (t68 t129))
(step t131 (cl (= y$13 @p_77)) :rule trans :premises (t89 t130))
(step t132 (cl (= @p_1 @p_31)) :rule rare_rewrite :args ("eq-symm" y$n1s4 y$n15s4))
(step t133 (cl (= @p_33 @p_33)) :rule refl)
(step t134 (cl (= @p_45 @p_34)) :rule cong :premises (t132 t47 t133))
(step t135 (cl (= @p_46 @p_35)) :rule cong :premises (t16 t134 t47))
(step t136 (cl (= @p_77 @p_75)) :rule cong :premises (t68 t135))
(step t137 (cl @p_76) :rule trans :premises (t131 t136))
(step t138 (cl y$13) :rule and :premises (t120) :args (1))
(step t139 (cl @p_75) :rule resolution :premises (t124 t137 t138))
(step t140 (cl (= @p_64 @p_47)) :rule cong :premises (t139 t129))
(step t141 (cl (not @p_74) @p_20 @p_73) :rule equiv_pos2)
(step t142 (cl (= @p_73 y$26)) :rule rare_rewrite :args ("bool-eq-true" y$26))
(step t143 (cl @p_74) :rule symm :premises (t142))
(step t144 (cl y$26) :rule and :premises (t120) :args (3))
(step t145 (cl @p_73) :rule resolution :premises (t141 t143 t144))
(step t146 (cl (= @p_20 @p_43)) :rule cong :premises (t145))
(step t147 (cl (= @p_58 @p_41)) :rule cong :premises (t139 t125))
(step t148 (cl (= @p_59 @p_42)) :rule cong :premises (t147))
(step t149 (cl (= @p_60 @p_44)) :rule cong :premises (t146 t148))
(step t150 (cl (not @p_72) @p_22 @p_71) :rule equiv_pos2)
(step t151 (cl (= @p_71 y$45)) :rule rare_rewrite :args ("bool-eq-true" y$45))
(step t152 (cl @p_72) :rule symm :premises (t151))
(step t153 (cl y$45) :rule and :premises (t120) :args (4))
(step t154 (cl @p_71) :rule resolution :premises (t150 t152 t153))
(step t155 (cl (not @p_70) @p_27 @p_67) :rule equiv_pos2)
(step t156 (cl (= @p_67 y$44)) :rule rare_rewrite :args ("bool-eq-true" y$44))
(step t157 (cl @p_70) :rule symm :premises (t156))
(step t158 (cl y$44) :rule and :premises (t120) :args (5))
(step t159 (cl @p_67) :rule resolution :premises (t155 t157 t158))
(step t160 (cl (= @p_66 @p_49)) :rule cong :premises (t123 t140 t149 t145 t154 t159))
(step t161 (cl (= y$58 @p_49)) :rule trans :premises (t117 t160))
(step t162 (cl (= @p_49 @p_57)) :rule ac_simp)
(step t163 (cl (= @p_57 @p_50)) :rule and_simplify)
(step t164 (cl (= @p_49 @p_50)) :rule trans :premises (t162 t163))
(step t165 (cl (= @p_48 true)) :rule rare_rewrite :args ("eq-refl" y$n1s4))
(step t166 (cl (= @p_47 @p_38)) :rule cong :premises (t40 t135))
(step t167 (cl (= @p_47 true)) :rule trans :premises (t166 t43))
(step t168 (cl (= @p_43 false)) :rule rare_rewrite :args ("evaluate"))
(step t169 (cl (= @p_41 @p_36)) :rule rare_rewrite :args ("eq-symm" @p_35 y$n15s4))
(step t170 (cl (= @p_42 @p_53)) :rule cong :premises (t169))
(step t171 (cl (= @p_44 @p_54)) :rule cong :premises (t168 t170))
(step t172 (cl (= @p_54 @p_55)) :rule rare_rewrite :args ("eq-symm" false @p_53))
(step t173 (cl (= @p_55 @p_56)) :rule rare_rewrite :args ("bool-eq-false" @p_53))
(step t174 (cl (= @p_56 @p_36)) :rule rare_rewrite :args ("bool-double-not-elim" @p_36))
(step t175 (cl (= @p_55 @p_36)) :rule trans :premises (t173 t174))
(step t176 (cl (= @p_54 @p_36)) :rule trans :premises (t172 t175))
(step t177 (cl (= @p_44 @p_36)) :rule trans :premises (t171 t176))
(step t178 (cl (= @p_50 @p_51)) :rule cong :premises (t165 t167 t177))
(step t179 (cl (= @p_51 @p_52)) :rule ac_simp)
(step t180 (cl (= @p_52 @p_36)) :rule and_simplify)
(step t181 (cl (= @p_51 @p_36)) :rule trans :premises (t179 t180))
(step t182 (cl (= @p_50 @p_36)) :rule trans :premises (t178 t181))
(step t183 (cl (= @p_49 @p_36)) :rule trans :premises (t164 t182))
(step t184 (cl @p_37) :rule trans :premises (t161 t183))
(step t185 (cl (= @p_36 @p_36)) :rule cong :premises (t125 t45))
(step t186 (cl @p_37) :rule trans :premises (t184 t185))
(step t187 (cl @p_36) :rule resolution :premises (t58 t186 a19))
(step t188 (cl @p_41) :rule symm :premises (t187))
(step t189 (cl (not @p_134) @p_133 @p_126) :rule equiv_pos2)
(step t190 (cl (= @p_31 @p_31)) :rule refl)
(step t191 (cl (= @p_127 @p_125)) :rule rare_rewrite :args ("eq-symm" @p_34 y$n1s4))
(step t192 (cl (= @p_123 @p_124)) :rule rare_rewrite :args ("eq-symm" @p_34 @p_33))
(step t193 (cl @p_134) :rule cong :premises (t190 t191 t192))
(step t194 (cl (not @p_129) @p_133 @p_128) :rule equiv_pos2)
(step t195 (cl (= @p_127 @p_127)) :rule cong :premises (t28 t47))
(step t196 (cl (= @p_123 @p_123)) :rule cong :premises (t28 t133))
(step t197 (cl @p_129) :rule cong :premises (t190 t195 t196))
(step t198 (cl @p_128) :rule rare_rewrite :args ("ite-eq" @p_31 y$n1s4 @p_33))
(step t199 (cl @p_128) :rule resolution :premises (t194 t197 t198))
(step t200 (cl @p_126) :rule resolution :premises (t189 t193 t199))
(step t201 (cl @p_31 @p_124) :rule ite1 :premises (t200))
(step t202 (cl @p_32) :rule not_symm :premises (a0))
(step t203 (cl @p_124) :rule resolution :premises (t201 t202))
(step t204 (cl @p_123) :rule symm :premises (t203))
(step t205 (cl (= @p_103 @p_106)) :rule cong :premises (t188 t204))
(step t206 (cl @p_114) :rule equiv_simplify)
(step t207 (cl @p_107 @p_122) :rule equiv2 :premises (t206))
(step t208 (cl (not @p_122) @p_106) :rule not_not)
(step t209 (cl @p_107 @p_106) :rule resolution :premises (t207 t208))
(step t210 (cl (not @p_109) (not @p_28) @p_108) :rule equiv_pos2)
(step t211 (cl (not @p_119) (not @p_26) @p_118) :rule equiv_pos2)
(step t212 (cl (= y$42 y$42)) :rule refl)
(step t213 (cl (not @p_121) (not @p_24) @p_120) :rule equiv_pos2)
(step t214 (cl (= y$33 y$33)) :rule refl)
(step t215 (cl (= @p_23 @p_116)) :rule cong :premises (t125 a3))
(step t216 (cl @p_121) :rule cong :premises (t214 t215))
(step t217 (cl @p_120) :rule resolution :premises (t213 t216 a15))
(step t218 (cl (= @p_25 @p_117)) :rule cong :premises (t65 t217))
(step t219 (cl @p_119) :rule cong :premises (t212 t218))
(step t220 (cl @p_118) :rule resolution :premises (t211 t219 a16))
(step t221 (cl (= @p_116 @p_106)) :rule cong :premises (t125 t127))
(step t222 (cl (= @p_117 @p_110)) :rule cong :premises (t123 t221))
(step t223 (cl (= y$42 @p_110)) :rule trans :premises (t220 t222))
(step t224 (cl (= @p_27 @p_43)) :rule cong :premises (t159))
(step t225 (cl (= @p_28 @p_111)) :rule cong :premises (t223 t224))
(step t226 (cl (= @p_111 @p_112)) :rule rare_rewrite :args ("eq-symm" @p_110 @p_43))
(step t227 (cl (= @p_106 @p_106)) :rule refl)
(step t228 (cl (= @p_110 @p_115)) :rule cong :premises (t165 t227))
(step t229 (cl (= @p_115 @p_106)) :rule and_simplify)
(step t230 (cl (= @p_110 @p_106)) :rule trans :premises (t228 t229))
(step t231 (cl (= @p_112 @p_113)) :rule cong :premises (t168 t230))
(step t232 (cl (= @p_113 @p_107)) :rule rare_rewrite :args ("eq-symm" false @p_106))
(step t233 (cl @p_114) :rule rare_rewrite :args ("bool-eq-false" @p_106))
(step t234 (cl (= @p_113 @p_108)) :rule trans :premises (t232 t233))
(step t235 (cl (= @p_112 @p_108)) :rule trans :premises (t231 t234))
(step t236 (cl (= @p_111 @p_108)) :rule trans :premises (t226 t235))
(step t237 (cl @p_109) :rule trans :premises (t225 t236))
(step t238 (cl @p_108) :rule resolution :premises (t210 t237 a17))
(step t239 (cl @p_107) :rule resolution :premises (t209 t238))
(step t240 (cl @p_105) :rule trans :premises (t205 t239))
(step t241 (cl @p_104) :rule resolution :premises (t57 t240))
(step t242 (cl @p_102) :rule resolution :premises (t55 t241))
(step t243 (cl @p_101) :rule resolution :premises (t53 t242))
(step t244 (cl) :rule resolution :premises (t14 t243 t187 t202))
