(assume nt288.0 (not (= (f6 c_0) (f5 c_0))))
(assume nt288.1 (not (not (and (= c_0 (f5 c_0)) (= c_0 (f6 c_0))))))
(assume nt288.2 (not (not (and (= c_0 (f5 c_0)) (= c_0 (f6 c_0))))))
(assume t285 (or (not (= c_0 (f5 c_0))) (not (= c_0 (f6 c_0))) (= (f6 c_0) (f5 c_0))))
(assume t286 (or (not (and (= c_0 (f5 c_0)) (= c_0 (f6 c_0)))) (= c_0 (f5 c_0))))
(assume t287 (or (not (and (= c_0 (f5 c_0)) (= c_0 (f6 c_0)))) (= c_0 (f6 c_0))))
(step t286' (cl (not (and (= c_0 (f5 c_0)) (= c_0 (f6 c_0)))) (= c_0 (f5 c_0))) :rule or :premises (t286))
(step t285' (cl (not (= c_0 (f5 c_0))) (not (= c_0 (f6 c_0))) (= (f6 c_0) (f5 c_0))) :rule or :premises (t285))
(step t287' (cl (not (and (= c_0 (f5 c_0)) (= c_0 (f6 c_0)))) (= c_0 (f6 c_0))) :rule or :premises (t287))
(step t288 (cl (= (f6 c_0) (f5 c_0)) (not (and (= c_0 (f5 c_0)) (= c_0 (f6 c_0)))) (not (and (= c_0 (f5 c_0)) (= c_0 (f6 c_0))))) :rule resolution :premises (t285' t286' t287'))
(step t.end (cl) :rule resolution :premises (nt288.0 nt288.1 nt288.2 t288))
