unsat
(assume a0 (distinct y$n0s8 y$n2s8 y$n3s8 y$n5s8 y$n4s8))
(assume a1 (= y$L1 (not y$2)))
(assume a2 (= y$L2 (not y$4)))
(assume a3 (= y$L3 (not y$6)))
(assume a4 (= y$L4 (not y$8)))
(assume a5 (= y$L5 (not y$10)))
(assume a6 (= y$L6 (not y$12)))
(assume a7 (= y$L7 (not y$14)))
(assume a8 (= y$18 (= y$n0s8 y$X)))
(assume a9 (= y$20 (= y$n0s8 y$Y)))
(assume a10 (= y$prop (not y$215)))
(assume a11 (= y$205 (= y$14 y$prop)))
(assume a12 (= y$219 (and y$L0 y$2 y$4 y$6 y$8 y$10 y$12 y$14 y$LoneHot y$18 y$20 y$215 y$205)))
(assume a13 y$219)
(step t0 (cl (not (= y$205 (= y$14 y$prop))) (not y$205) (= y$14 y$prop)) :rule equiv_pos2)
(step t1 (cl (not (= y$219 (and y$L0 y$2 y$4 y$6 y$8 y$10 y$12 y$14 y$LoneHot y$18 y$20 y$215 y$205))) (not y$219) (and y$L0 y$2 y$4 y$6 y$8 y$10 y$12 y$14 y$LoneHot y$18 y$20 y$215 y$205)) :rule equiv_pos2)
(step t2 (cl (and y$L0 y$2 y$4 y$6 y$8 y$10 y$12 y$14 y$LoneHot y$18 y$20 y$215 y$205)) :rule resolution :premises (t1 a12 a13))
(step t3 (cl y$205) :rule and :premises (t2) :args (12))
(step t4 (cl (= y$14 y$prop)) :rule resolution :premises (t0 a11 t3))
(step t5 (cl (not (= y$14 y$prop)) (not y$14) y$prop) :rule equiv_pos2)
(step t6 (cl y$14) :rule and :premises (t2) :args (7))
(step t7 (cl (not y$prop) (not y$215)) :rule equiv1 :premises (a10))
(step t8 (cl y$215) :rule and :premises (t2) :args (11))
(step t9 (cl (not y$prop)) :rule resolution :premises (t7 t8))
(step t10 (cl (not (= y$14 y$prop))) :rule resolution :premises (t5 t6 t9))
(step t11 (cl) :rule resolution :premises (t4 t10))
